# Lab N: Interaction

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

## Setup

```python
import builtins
import math
import warnings
from datascience import Table, make_array
from statsmodels.formula.api import ols

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Advertising & Pricing Beauty Products

### Business Decision

Tyler Beauty Products advertises and sells beauty products.  It recently conducted a set of marketing programs at various retail stores, where products were priced at 3 different levels: \\$2.00, \\$2.50, or \\$3.00.  Each of these stores advertised at 2 different budget levels: \\$50,000 or \\$100,000.  The product manager must now decide on advertising budgets of \\$50,000 or \\$100,000 at each of several new stores.  For various strategic reasons, each of the new stores has already decided whether to price the beauty products at \\$2.00, \\$2.50, or \\$3.00.

### Retrieve Data

Retrieve data from file 'Tyler.csv'. Show ALL marketing program results.

```python
data = ...
data
```

```python
grader.check("q1")
```

### Analysis

#### Model 1

Build a multiple linear regression model to predict 'Sales' based on 'Price' and 'Budget'.  Show the model goodness of fit (R^2).  Show the model parameter values (intercept and coefficients).

```python
model_formula = ...
model1 = ...
model1.rsquared
model1.params
```

```python
grader.check("q2")
```

Build a report that shows what sales and profits the model would predict for each combination of price and budget.
Calculate profit as sales minus budget.

```python
report1 = Table().with_columns('Price', make_array(2, 2.5, 3))
report1 = report1.with_column('Budget', make_array(50, 50, 50))

sales_m1_predicted_50 = ...
report1 = report1.with_column('sales_predicted_50', sales_m1_predicted_50)

profit_m1_50 = ...
report1 = report1.with_column('profit_50', profit_m1_50)

report1 = report1.with_column('Budget', make_array(100, 100, 100)) # replace Budget amount

sales_m1_predicted_100 = ...
report1 = report1.with_column('sales_predicted_100', sales_m1_predicted_100)

profit_m1_100 = ...
report1 = report1.with_column('profit_100', profit_m1_100)

report1 = report1.drop('Budget')

report1
```

```python
grader.check("q3")
```

Visualize predicted profits at budget \$50,000 and budget \$100,000 vs. price, as a line plot.

```python
x_axis_column_name = ...
y_axis_column_name = ...
y_axis_column_name_2 = ...
report1.select(x_axis_column_name, y_axis_column_name, y_axis_column_name_2).plot(x_axis_column_name)
```

```python
grader.check("q4")
```

For each price, determine which budget results in the best predicted profit.

```python
for price in make_array(2, 2.5, 3):
    budget = make_array('Budget 50', 'Budget 100')
    profit_m1_50  = ...
    profit_m1_100 = ...
    larger_m1_profit = ...
    (price, budget.item((profit_m1_50<profit_m1_100)*1), larger_m1_profit)
```

```python
grader.check("q5")
```

#### Model 2

Add a synthetic numeric variable called 'priceXbudget' to the data, with values that are the product of 'Price' and
'Budget'. This is an INTERACTION variable. Show ALL the marketing program results.

```python
syth_col_data = ...
data2 = data.with_column('priceXbudget', syth_col_data)
data2.show()
```

```python
grader.check("q6")
```

Build a multiple linear regression model to predict 'Sales' based on 'Price', 'Budget', and 'priceXbudget'.  Show the model goodness of fit (R^2).  Show the model parameter values (intercept and coefficients).

```python
model_formula = ...
model2 = ols(model_formula, data2).fit()
r_squared = ...
model2.params
```

```python
grader.check("q7")
```

Build a report2that shows what sales and profits the model would predict for each combination of price and budget.
Calculate profit as sales minus budget.

```python
report2= Table().with_columns('Price', make_array(2, 2.5, 3))
report2= report2.with_column('Budget', make_array(50, 50, 50))

array_price_50_budget = ...
report2= report2.with_column('priceXbudget', array_price_50_budget)

sales_m2_predicted_50 = ...
report2= report2.with_column('sales_predicted_50', sales_m2_predicted_50)

profit_m2_50 = ...
report2= report2.with_column('profit_50', profit_m2_50)


report2= report2.with_column('Budget', make_array(100, 100, 100)) # replace Budget amount

array_price_100_budget = ...
report2= report2.with_column('priceXbudget', array_price_100_budget)

sales_m2_predicted_100 = ...
report2= report2.with_column('sales_predicted_100', sales_m2_predicted_100)

profit_m2_100 = ...
report2= report2.with_column('profit_100', profit_m2_100)

report2= report2.drop('Budget')
report2= report2.drop('priceXbudget')

report2
```

```python
grader.check("q8")
```

Visualize predicted profits at budget \$50,000 and budget \$100,000 vs. price, as a line plot.

```python
x_axis_column_name_2 = ...
y_axis_column_name_2 = ...
y_axis_column_name_3 = ...
report2.select(x_axis_column_name_2, y_axis_column_name_2, y_axis_column_name_3).plot(x_axis_column_name_2)
```

```python
grader.check("q9")
```

For each price, determine which budget results in the best predicted profit.

```python
for price in make_array(2, 2.5, 3):
    budget = make_array('Budget 50', 'Budget 100')
    profit_m2_50  = ...
    profit_m2_100 = ...
    larger_m2_profit = ...
    (price, budget.item((profit_m2_50<profit_m2_100)*1), larger_m2_profit)
```

```python
grader.check("q10")
```

Calculate the price where it becomes more profitable to budget \$50,000 than to budget \$100,000.

$$
\textit{predicted profit at budget 50,000} = \big( i + (c_1)(\textit{price}) + (c_2)(50) + (c_3)(\textit{price})(50) \big) - 50
$$

$$
\textit{predicted profit at budget 100,000} = \big( i + (c_1)(\textit{price}) + (c_2)(100) + (c_3)(\textit{price})(100) \big) - 100
$$

$$
i + (c_1)(\textit{price}) + (c_2)(50) + (c_3)(\textit{price})(50) - 50 = i + (c_1)(\textit{price}) + (c_2)(100) + (c_3)(\textit{price})(100) - 100
$$

$$
(c_2)(50) + (c_3)(\textit{price})(50) - 50 = (c_2)(100) + (c_3)(\textit{price})(100) - 100
$$

$$
(c_2)(50) - (c_2)(100) - 50 + 100 =  (c_3)(\textit{price})(100) - (c_3)(\textit{price})(50)
$$

$$
(c_2)(50) - (c_2)(100) - 50 + 100 = (\textit{price}) \big((c_3)(100) - (c_3)(50)\big)
$$

$$
\cfrac{(c_2)(50) - (c_2)(100) - 50 + 100}{(c_3)(100) - (c_3)(50)}  = \textit{price}
$$


```python
c2 = ...
c3 = ...

crossover_numerator = ...
crossover_denominator = ...
crossover = crossover_numerator / crossover_denominator
crossover
```

```python
grader.check("q11")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Use the interaction model to predict sales at price $2.50 and budget $100,000.

```python
sales_c1 = ...
sales_c1
```

```python
grader.check("qc_1")
```

**Question 2**. According to the interaction model, calculate the difference in profit between a $50,000 budget and a $100,000 budget at price $3.00.

```python
profit50 = ...
profit100 = ...
profit_c2 = ...
profit_c2
```

```python
grader.check("qc_2")
```

**Question 3**. According to the interaction model, calculate the range of profits across all 6 price/budget combinations.

```python
profit_20_50 = ...
profit_25_50 = ...
profit_30_50 = ...
profit_20_100 = ...
profit_25_100 = ...
profit_30_100 = ...
lowest = ...
highest = ...
range = ...
range
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised November 24, 2023
</span>
</p>

Tests for q1 (points: None)
```
>>> data.labels == ('Store', 'Price', 'Budget', 'Sales')
True
```
```
>>> data.num_rows == 24
True
```


Tests for q10 (points: None)
```
>>> any(profit_m2_100 == report2.column('profit_100'))
True
```
```
>>> any(profit_m2_100 == report2.column('profit_100'))
True
```
```
>>> any(profit_m2_50 == report2.column('profit_50'))
True
```
```
>>> any(larger_m2_profit == report2.column('profit_100'))
False
```


Tests for q11 (points: None)
```
>>> c2 == model2.params[2]
True
```
```
>>> c3 == model2.params[3]
True
```
```
>>> crossover_numerator == c2 * 50 - c2 * 100 - 50 + 100
True
```
```
>>> crossover_denominator == c3 * 100 - c3 * 50
True
```


Tests for q2 (points: None)
```
>>> 'Budget' in model_formula
True
```
```
>>> math.isclose(model1.rsquared, 0.791, abs_tol=0.01)
True
```
```
>>> hasattr(model1, 'rsquared') and hasattr(model1, 'params') and (len(model1.params) == 3)
True
```


Tests for q3 (points: None)
```
>>> report1 = report1.with_column('Budget', make_array(50, 50, 50))
>>> all(sales_m1_predicted_50 == model1.predict(report1))
True
```
```
>>> report1 = report1.with_column('Budget', make_array(100, 100, 100))
>>> all(sales_m1_predicted_100 == model1.predict(report1))
True
```


Tests for q4 (points: None)
```
>>> x_axis_column_name == 'Price'
True
```
```
>>> y_axis_column_name == 'profit_50'
True
```
```
>>> x_axis_column_name == 'Price'
True
```


Tests for q5 (points: None)
```
>>> assert profit_m1_100 is not Ellipsis
>>> any(profit_m1_100 == report1.column('profit_100'))
True
```
```
>>> assert profit_m1_100 is not Ellipsis
>>> any(profit_m1_100 == report1.column('profit_100'))
True
```
```
>>> assert profit_m1_50 is not Ellipsis
>>> any(profit_m1_50 == report1.column('profit_50'))
True
```
```
>>> assert larger_m1_profit is not Ellipsis
>>> any(larger_m1_profit == report1.column('profit_100'))
True
```


Tests for q6 (points: None)
```
>>> all(syth_col_data == data2.column('Price') * data2.column('Budget'))
True
```


Tests for q7 (points: None)
```
>>> 'priceXbudget' in model_formula
True
```
```
>>> import math
>>> math.isclose(r_squared, 0.97, abs_tol=0.05)
True
```


Tests for q8 (points: None)
```
>>> report2 = report2.with_column('Budget', make_array(50, 50, 50))
>>> report2 = report2.with_column('priceXbudget', make_array(50 * 2, 50 * 2.5, 50 * 3))
>>> all(sales_m2_predicted_50 == model2.predict(report2))
True
```
```
>>> report2 = report2.with_column('Budget', make_array(100, 100, 100))
>>> report2 = report2.with_column('priceXbudget', make_array(100 * 2, 100 * 2.5, 100 * 3))
>>> all(sales_m2_predicted_100 == model2.predict(report2))
True
```


Tests for q9 (points: None)
```
>>> x_axis_column_name_2 == 'Price'
True
```
```
>>> y_axis_column_name_2 == 'profit_50'
True
```
```
>>> y_axis_column_name_3 == 'profit_100'
True
```


Tests for qc_1 (points: None)
```
>>> all(sales_c1 == model2.predict(Table().with_columns('Price', 2.5, 'Budget', 100, 'priceXbudget', 2.5 * 100)))
True
```


Tests for qc_2 (points: None)
```
>>> all(profit50 == model2.predict(Table().with_columns('Price', 3, 'Budget', 50, 'priceXbudget', 3 * 50)) - 50)
True
```
```
>>> all(profit100 == model2.predict(Table().with_columns('Price', 3, 'Budget', 100, 'priceXbudget', 3 * 100)) - 100)
True
```
```
>>> all(profit_c2 == profit50 - profit100)
True
```


Tests for qc_3 (points: None)
```
>>> all(profit_20_50 == model2.predict(Table().with_columns('Price', 2.0, 'Budget', 50, 'priceXbudget', 2 * 50)) - 50)
True
```
```
>>> all(profit_25_50 == model2.predict(Table().with_columns('Price', 2.5, 'Budget', 50, 'priceXbudget', 2.5 * 50)) - 50)
True
```
```
>>> all(profit_30_50 == model2.predict(Table().with_columns('Price', 3.0, 'Budget', 50, 'priceXbudget', 3 * 50)) - 50)
True
```
```
>>> all(profit_20_100 == model2.predict(Table().with_columns('Price', 2.0, 'Budget', 100, 'priceXbudget', 2 * 100)) - 100)
True
```
```
>>> all(profit_25_100 == model2.predict(Table().with_columns('Price', 2.5, 'Budget', 100, 'priceXbudget', 2.5 * 100)) - 100)
True
```
```
>>> all(profit_30_100 == model2.predict(Table().with_columns('Price', 3.0, 'Budget', 100, 'priceXbudget', 3 * 100)) - 100)
True
```
```
>>> lowest == min(make_array(profit_20_50, profit_25_50, profit_30_50, profit_20_100, profit_25_100, profit_30_100)).item(0)
True
```
```
>>> highest == max(make_array(profit_20_50, profit_25_50, profit_30_50, profit_20_100, profit_25_100, profit_30_100)).item(0)
True
```
```
>>> range == highest - lowest
True
```
