# Lab I: A/B Causality Test

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

## Setup

```python
import warnings
from datascience import Table, make_array, are, percentile
from numpy import mean, arange, append
from numpy.random import seed, choice

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Website Click-Throughs

### Business Decision

A software and services company wants to drive its website visitors from the splash page to commerce pages that describe and offer additional for-fee services.  The regular splash page includes graphics, text, and a link to the commerce pages.  At the start of each month, a new splash page design is introduced on a trial basis for a week, which includes certain new graphics, new text, and a link to the commerce pages.  When a visitor goes to the company's website, that visitor is presented with either the regular splash page or the new splash page, selected by machine randomly with equal probabilities.  After a week, if the new splash page seems more effective, then it replaces the regular splash page for the remainder of the month.

Assume that you are the Vice President of Marketing.  Does Design A cause more visitors to go to the commerce pages than Design B does?  Should you revise the splash page to present Design A instead of Design B to all visitors?

### Data

Retrieve the website splash page click-through dataset of all visitors during the trial period (Splash_Page_1.csv).  
Each observation represents one visitor interaction: 
* customer_id is a unique identifier for the visitor
* design is A or B, and indicates which design was presented
* clickthru is 0 or 1, where 0 indicates that the visitor did not go to the commerce pages, and 1 indicates that the visitor did

Show the first few observations.

```python
splash_page = ...
...
```

```python
grader.check("q1")
```

### Analysis

#### Actual Outcome

Explore what actually happened during the trial period.

For each design, show the number of visitors that were presented that design during the trial period.

```python
trial_visitors = ...
trial_visitors
```

```python
grader.check("q2_1")
```

For each design, show the probability that a visitor went to the commerce pages during the trial period.

```python
trial_prob = ...
trial_prob
```

```python
grader.check("q2_2")
```

Show the difference between the probabilities.

```python
trial_diff = ...
trial_diff
```

```python
grader.check("q2_3")
```

#### Hypothesized Outcome

Note that visitors were presented with one of two possible designs by chance.  Perhaps these results are due to chance, too.  In other words, perhaps some visitors would have gone to the commerce pages regardless of which design was presented.  Similarly, perhaps some vistors would have not gone regardless.

To test for this possible chance occurrence, hypothesize that the designs had no effect on the visitors' actions, and then look for evidence that suggests that your hypothesis is unlikely.

Under the assumption that the designs have no effect, you could have scrambled which visitors were presented with which designs, and the visitors' actions would have been the same.

Construct an array of design choices (A or B), in the same proportion as in the actual trial period,
but randomly ordered.  _To do this, use the choice() function with argument replace=False._

```python
seed(1)
design_data = ...
design_total = ...

scramble = choice(design_data, splash_page.num_rows, replace=False)
scramble
```

```python
grader.check("q3_1")
```

Construct a table like the one you just did showing the assumed clickthru values, except
replace the design column with the one new array (for scrambled design choices).

```python
data_hypo = ...
data_hypo
```

```python
grader.check("q3_2")
```

For each design, show the probability that a visitor would have gone to the commerce pages during the trial period
(one probability for Design A and another probability for Design B), assuming the scrambled design choices.

```python
prob = ...
prob
```

```python
grader.check("q3_3")
```

Show the difference between the probabilities.

```python
diff = ...
diff
```

```python
grader.check("q3_4")
```

Explore the range of results you could have seen if the design choice had no effect on visitors' actions.

Get 1000 different scrambles of design choices.  For each scramble, calculate what would have been
the probabilities of a vistor going to the commerce pages (one probability for Design A and another
probability for Design B).  Calculate and show a histogram of all 1000 probability differences (30 bins, range -0.1 to 0.1).

```python
seed(1)

dist_array = make_array()

for i in arange(1000):
    scramble_i = choice(splash_page.column('design'), splash_page.num_rows, replace=False)
    data_hypo_i = ...
    prob_i = ...
    diff_i = ...
    dist_array = append(dist_array, diff_i)

# Name the column of probability differences 'diff'
dist = ...

dist_hist = dist.hist(bins=30, range=make_array(-0.1, 0.1))
dist_hist
```

```python
grader.check("q3_5")
```

Calculate and show the probability that the trial probability difference would be >= 0.0512
if the design has no effect on visitors' actions. (This is the p-value.)
Also show the histogram of probability differences with the area corresponding to the probability highlighted.

```python
p_value = ...
p_value

left_end_values = ...

dist.hist(bins=30, range=make_array(-0.1, 0.1), left_end=left_end_values, right_end=0.1)
```

```python
grader.check("q3_6")
```

Calculate and show the lower bound of the probability difference corresponding to the highest 5%
of results, based on what you could have seen if the design choice has no effect on visitors' actions.
(This is the critical value @ significance level 0.05.)
Also show the histogram of probability differences with the area corresponding to the significance level highlighted.

```python
cv = ...
cv

left_end_values2 = ...

dist.hist(bins=30, range=make_array(-0.1, 0.1), left_end=left_end_values2, right_end=0.1)
```

```python
grader.check("q3_7")
```

Show that you should conclude that the hypothesis is correct (i.e., conclude that the design makes no difference), at significance level 0.05. You should only use comparison symbols (>, <, =) in the blank. Define reason_1 and reason_2 as the full string expressing the statement. For example, reason_1 = "p_value = sig_level". Don't include any extra spaces.

```python
# Reason 1: "p_value ___ 0.05"
reason_1 = ...

# Reason 2: "trial_diff ___ cv"
reason_2 = ...

reason_1, reason_2
```

```python
grader.check("q3_8")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Calculate the proportion of visitors who were actually presented with Design A.

```python
count_A = ...
prop_A = ...
prop_A
```

```python
grader.check("qc_1")
```

**Question 2**. Assume a significance level of 0.10, rather than 0.05. Calculate the critical value. 

```python
cv_c2 = ...
cv_c2
```

```python
grader.check("qc_2")
```

**Question 3**. Assume a significance level of 0.10, rather than 0.05. Calculate whether the manager should conclude that design has no effect on clickthrough rate.
(Answer is True or False.)

```python
cv_c3 = ...
conclusion_c3 = ...
conclusion_c3
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised October 20, 2023
</span>
</p>

Tests for q1 (points: None)
```
>>> splash_page.labels == ('customer_id', 'design', 'clickthru')
True
```
```
>>> splash_page.num_rows == 1014
True
```


Tests for q2_1 (points: None)
```
>>> trial_visitors.labels == ('design', 'count')
True
```
```
>>> trial_visitors.num_rows == 2
True
```


Tests for q2_2 (points: None)
```
>>> trial_prob.labels == ('design', 'clickthru mean')
True
```
```
>>> trial_prob.num_rows == 2
True
```


Tests for q2_3 (points: None)
```
>>> round(trial_diff, 2) == 0.05
True
```


Tests for q3_1 (points: None)
```
>>> all(design_data == splash_page.column('design'))
True
```
```
>>> design_total == splash_page.num_rows
True
```


Tests for q3_2 (points: None)
```
>>> data_hypo.labels == ('customer_id', 'design', 'clickthru')
True
```
```
>>> data_hypo.num_rows == 1014
True
```


Tests for q3_3 (points: None)
```
>>> prob.labels == ('design', 'clickthru mean')
True
```
```
>>> prob.num_rows == 2
True
```


Tests for q3_4 (points: None)
```
>>> round(diff, 2) == 0.03
True
```


Tests for q3_5 (points: None)
```
>>> data_hypo_i.labels == ('customer_id', 'design', 'clickthru') and data_hypo_i.num_rows == 1014
True
```
```
>>> prob_i.labels == ('design', 'clickthru mean') and prob_i.num_rows == 2
True
```
```
>>> round(diff_i, 3) == -0.008
True
```
```
>>> 'diff' in dist.labels and all(dist.column('diff') == dist_array)
True
```


Tests for q3_6 (points: None)
```
>>> p_value == 0.065
True
```
```
>>> left_end_values == trial_diff
True
```


Tests for q3_7 (points: None)
```
>>> round(cv, 3) == 0.055
True
```
```
>>> left_end_values2 == cv
True
```


Tests for q3_8 (points: None)
```
>>> (reason_1 == 'p_value > 0.05') | (reason_1 == 'p_value>0.05')
True
```
```
>>> (reason_2 == 'trial_diff < cv') | (reason_2 == 'trial_diff<cv')
True
```


Tests for qc_1 (points: None)
```
>>> all(count_A == splash_page.group('design').column('count'))
True
```
```
>>> prop_A == count_A.item(0) / sum(count_A)
True
```


Tests for qc_2 (points: None)
```
>>> cv_c2 == percentile(90, dist.column('diff'))
True
```


Tests for qc_3 (points: None)
```
>>> conclusion_c3 == False
True
```
