# Lab M2: Seasonality

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

## Setup

```python
import warnings
from datascience import Table
from numpy import mean, mod, sqrt
from statsmodels.formula.api import ols

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Investing in Smartphones

### Business Decision

Investors in a smartphone retailer want to assess the company's growth potential.  They are especially interested in how stable revenue growth is.

Assume that you are the director of financial planning and analysis for the smartphone company.  How would you estimate next year's sales?  How confident would you be about your estimate?

Assume that you are an investor in the smartphone company.  How confident are you about the company's growth potential?

### Data

Retrieve a dataset from file 'Smartphones.csv'.  Show the first few observations.
Visualize the dataset as a scatterplot.

```python
smartphones = ...
smartphones

smartphones.scatter('quarter')
```

```python
grader.check("q1_1")
```

Visualize the dataset as a lineplot.

```python
smartphones_plot = smartphones.plot('quarter')
smartphones_plot
```

### Analysis: Linear Regression with Seasonality Variables

#### Model That Accounts for Trend

Build a linear regression model to predict sales based on quarter.
Show the model goodness of fit (R^2).
Show the model parameters.

```python
model_1 = ...
model_1.rsquared
model_1.params
```

```python
grader.check("q2_1")
```

Add a variable to the dataset for predicted sales.
Show the resulting dataset.

```python
predicted_sales_1 = ...

smartphones_predictions = smartphones.with_column('sales_predicted', predicted_sales_1)
smartphones_predictions
```

```python
grader.check("q2_2")
```

Show the RMSE calculated based on the dataset.  
Visualize the performance of the model as a lineplot of sales and predicted sales vs. quarter.

```python
RMSE_1 = ...
RMSE_1

smartphones_plot = smartphones_predictions.plot('quarter')
smartphones_plot
```

```python
grader.check("q2_3")
```

#### Model That Accounts for Trend & Seasonality

Reset dataset to include only sales and quarter variables.  Show the first few observations.

```python
smartphones = smartphones.select('quarter','sales')
smartphones
```

Add a variable to the dataset for quarter name.  
1 indicates a 1st quarter, 2 indicates a 2nd quarter, 3 indicates a 3rd quarter, 4 indicates a 4th quarter).  

Show the first few observations of the resulting dataset.  
You can use the `mod` function for modulo arithmetic, e.g., `mod(data.column('quarter'),4)+1`.

```python
quarter_name = ...

smartphones_quarter = smartphones.with_column('quarter_name', quarter_name)
smartphones_quarter
```

```python
grader.check("q3_1")
```

Convert the quarter name variable to dummy variables.  Show the first few observations.

```python
q2_dummy = ...
q3_dummy = ...
q4_dummy = ...

smartphones_dummy = smartphones_quarter.with_columns('q2', q2_dummy.astype(int),
                         'q3', q3_dummy.astype(int),
                         'q4', q4_dummy.astype(int))
smartphones_dummy
```

```python
grader.check("q3_2")
```

Build a linear regression model to predict sales based on quarter and quarter name dummy variables.  
Show the model goodness of fit (R^2).  
Show the model parameters.

```python
model_2 = ...
model_2.rsquared
model_2.params
```

```python
grader.check("q3_3")
```

Add a variable to the dataset for predicted sales.

```python
predicted_sales_2 = ...

smartphones_dummy_preds = smartphones_dummy.with_column('sales_predicted', predicted_sales_2)
smartphones_dummy_preds
```

```python
grader.check("q3_4")
```

Show the RMSE calculated based on the dataset.  
Visualize the performance of the model as a lineplot of sales and predicted sales vs. quarter.

```python
RMSE_2 = ...
RMSE_2

dummy_preds_plot = smartphones_dummy_preds.select('quarter','sales','sales_predicted').plot('quarter')
dummy_preds_plot
```

```python
grader.check("q3_4")
```

Predict the sales in quarter 32, which is the 1st quarter of the year.

```python
q32_predictions = model_2.predict(Table().with_columns('quarter', 32, 'q2', 0, 'q3', 0, 'q4', 0))
q32_predictions
```

```python
grader.check("q3_5")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Use the model accounting for trend and seasonality (model_2) to predict sales in quarter 33.

```python
model_2_preds = ...
model_2_preds
```

```python
grader.check("qc_1")
```

**Question 2**. Use the model accounting for trend and seasonality (model_2) to calculate the difference between actual sales and predicted sales in quarter 7.

```python
residual_q7 = ...
residual_q7
```

```python
grader.check("qc_2")
```

**Question 3**. According to the model accounting for trend and seasonality (model_2) and the dataset, calculate the difference between the RMSE and the mean absolute error between sales and predicted sales.

```python
RMSE_percent_improvement = ...
RMSE_percent_improvement
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised November 18, 2023
</span>
</p>

Tests for q1_1 (points: None)
```
>>> smartphones.labels == ('quarter', 'sales') and smartphones.num_rows == 32
True
```


Tests for q2_1 (points: None)
```
>>> cond1 = list(model_1.model.data.frame.labels) == list(smartphones.labels)
>>> cond2 = model_1.model.data.frame.num_rows == smartphones.num_rows
>>> cond3 = all((model_1.model.data.frame.row(i) == smartphones.row(i) for i in range(smartphones.num_rows)))
>>> cond1 and cond2 and cond3
True
```
```
>>> hasattr(model_1, 'rsquared') and hasattr(model_1, 'params') and (len(model_1.params) == 2)
True
```


Tests for q2_2 (points: None)
```
>>> all(predicted_sales_1 == model_1.predict(smartphones))
True
```


Tests for q2_3 (points: None)
```
>>> RMSE_1 == sqrt(mean(model_1.resid ** 2))
True
```


Tests for q3_1 (points: None)
```
>>> all(quarter_name == mod(smartphones.column('quarter'), 4) + 1)
True
```


Tests for q3_2 (points: None)
```
>>> all((row.q2 == (row.quarter_name == 2) for row in smartphones_dummy.rows))
True
```
```
>>> all((row.q3 == (row.quarter_name == 3) for row in smartphones_dummy.rows))
True
```
```
>>> all((row.q4 == (row.quarter_name == 4) for row in smartphones_dummy.rows))
True
```


Tests for q3_3 (points: None)
```
>>> cond1 = list(model_2.model.data.frame.labels) == list(smartphones_dummy.labels)
>>> cond2 = model_2.model.data.frame.num_rows == smartphones_dummy.num_rows
>>> cond3 = all((model_2.model.data.frame.row(i) == smartphones_dummy.row(i) for i in range(smartphones_dummy.num_rows)))
>>> cond1 and cond2 and cond3
True
```
```
>>> hasattr(model_2, 'rsquared') and hasattr(model_2, 'params') and (len(model_2.params) == 5)
True
```


Tests for q3_4 (points: None)
```
>>> all(predicted_sales_2 == model_2.predict(smartphones_dummy))
True
```
```
>>> RMSE_2 == sqrt(mean(model_2.resid ** 2))
True
```


Tests for q3_5 (points: None)
```
>>> all(q32_predictions == model_2.predict(Table().with_columns('quarter', 32, 'q2', 0, 'q3', 0, 'q4', 0)))
True
```


Tests for qc_1 (points: None)
```
>>> all(model_2_preds == model_2.predict(Table().with_columns('quarter', 33, 'q2', 1, 'q3', 0, 'q4', 0)))
True
```


Tests for qc_2 (points: None)
```
>>> residual_q7 == smartphones_dummy_preds.column('sales').item(7) - smartphones_dummy_preds.column('sales_predicted').item(7)
True
```


Tests for qc_3 (points: None)
```
>>> RMSE_percent_improvement == (RMSE_1 - RMSE_2) / RMSE_1 * 100
True
```
