# Lab C2: Weighted Average

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

## Setup

```python
# Import some useful functions
import warnings
from datascience import Table
from numpy import mean

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Hide some unnecessary warning messages
warnings.filterwarnings("ignore")
```

## Doctor Surveys

### Business Decision

A medical device company has surveyed customer satisfaction on a random sample of doctors that use its products.  The company partitions the doctors into three categories depending on their hospital size: small (1-50 beds), medium (51-500 beds), and large (501+ beds).  As with many, perhaps most, surveys, many of the the survey recipients did not respond.  Still, the Vice President of Customer Service needs to interpret the survey results and might need to adjust service and make product change recommendations accordingly.

### Data

The dataset lists the doctors that were sent the survey.  Each observation corresponds to a doctor, and indicates ...
* a unique, anonymized doctor identifier,
* whether the doctor responded to the survey,
* the doctor’s reported satisfaction with the company’s devices (measured on a 1-5 scale, where ‘1’ indicates ‘not satisfied at all’ and ‘5’ indicates ‘completely satisfied’), if the doctor responded, and
* the doctor’s hospital size, which historically has been a key distinguishing feature of the company’s customers.

Retrieve the hospital survey data, show the number of observations and the first few observations.

```python
hospital_survey = ...
...
hospital_survey
```

```python
grader.check("q1")
```

### Analysis | Weighted Average by Group, missing data, proportions known

Calculate the number of doctors surveyed.  
Calculate the number of responses.  
Calculate the response rate.  
Calculate the mean satisfaction.

```python
# Calculate the number of doctors surveyed
n_pop = ...
n_pop

# Calculate the number of responses
n_respond = ...
n_respond

# Calculate the response rate
rate = ...
rate

# Calculate the mean satisfaction
avg_sat = ...
avg_sat
```

```python
grader.check("q2_1")
```

Build a table comprising the observations at small hospitals, show the first few observations.  
Calculate the number of doctors surveyed at small hospitals.  
Calculate the number of responses at small hospitals.  
Calculate the response rate for small hospitals.  
Calculate the mean satisfaction for small hospitals.

```python
# Build a table comprising the observations at small hospitals, show the first few observations
data_small = ...
data_small

# Calculate the number of doctors surveyed at small hospitals
n_pop_small = ...
n_pop_small

# Calculate the number of responses at small hospitals
n_respond_small = ...
n_respond_small

# Calculate the response rate for small hospitals
rate_small = ...
rate_small

# Calculate the mean satisfaction for small hospitals
avg_sat_small = ...
avg_sat_small
```

```python
grader.check("q2_2")
```

Build a table comprising the observations at medium hospitals, show the first few observations.  
Calculate the number of doctors surveyed at medium hospitals.  
Calculate the number of responses at medium hospitals.  
Calculate the response rate for medium hospitals.  
Calculate the mean satisfaction for medium hospitals.

```python
# Build a table comprising the observations at medium hospitals, show the first few observations
data_medium = hospital_survey.where('hospital_type','medium')
data_medium

# Calculate the number of doctors surveyed at medium hospitals
n_pop_medium = data_medium.num_rows
n_pop_medium

# Calculate the number of responses at medium hospitals
n_respond_medium = sum(data_medium.column('response'))
n_respond_medium

# Calculate the response rate for medium hospitals
rate_medium = mean(data_medium.column('response'))
rate_medium

# Calculate the mean satisfaction for medium hospitals
avg_sat_medium = mean(data_medium.where('response',1).column('satisfaction'))
avg_sat_medium
```

```python
grader.check("q2_3")
```

Build a table comprising the observations at large hospitals.  
Calculate the number of doctors surveyed at large hospitals.  
Calculate the number of responses at large hospitals.  
Calculate the response rate for large hospitals.  
Calculate the mean satisfaction for large hospitals.

```python
# Build a table comprising the observations at large hospitals
data_large = hospital_survey.where('hospital_type','large')
data_large

# Calculate the number of doctors surveyed at large hospitals
n_pop_large = data_large.num_rows
n_pop_large

# Calculate the number of responses at large hospitals
n_respond_large = sum(data_large.column('response'))
n_respond_large

# Calculate the response rate for large hospitals
rate_large = mean(data_large.column('response'))
rate_large

# Calculate the mean satisfaction for large hospitals
avg_sat_large = mean(data_large.where('response',1).column('satisfaction'))
avg_sat_large
```

```python
grader.check("q2_4")
```

Calculate the weighted average satisfaction by hospital type.

```python
weighted_average = ...
weighted_average
```

```python
grader.check("q2_5")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Calculate the proportion of respondents (doctors that responded) that work in small hospitals.

```python
prop_respond_small = ...
prop_respond_small
```

```python
grader.check("qc_1")
```

**Question 2**. Calculate the proportion of respondents (doctors that responded) that work in small hospitals and scored satisfaction 5.

```python
prop_respond_small_happy = ...
prop_respond_small_happy
```

```python
grader.check("qc_2")
```

**Question 3**. Calculate the proportion of doctors that did not respond.

```python
prop_no_respond = (n_pop - (n_respond_small + n_respond_medium + n_respond_large)) / n_pop
prop_no_respond
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised September 1, 2023
</span>
</p>

Tests for q1 (points: None)
```
>>> hospital_survey.num_rows == 10300
True
```
```
>>> hospital_survey.labels == ('doctor_id', 'hospital_type', 'response', 'satisfaction')
True
```


Tests for q2_1 (points: None)
```
>>> n_pop == 10300
True
```
```
>>> n_respond == 4312
True
```
```
>>> round(rate, 2) == 0.42
True
```
```
>>> round(avg_sat, 2) == 3.75
True
```


Tests for q2_2 (points: None)
```
>>> all(data_small['hospital_type'] == 'small')
True
```
```
>>> n_pop_small == 3396
True
```
```
>>> n_respond_small == 520
True
```
```
>>> round(rate_small, 2) == 0.15
True
```
```
>>> round(avg_sat_small, 2) == 2.67
True
```


Tests for q2_3 (points: None)
```
>>> all(data_medium['hospital_type'] == 'medium')
True
```
```
>>> n_pop_medium == 3484
True
```
```
>>> n_respond_medium == 1395
True
```
```
>>> round(rate_medium, 1) == 0.4
True
```
```
>>> round(avg_sat_medium, 1) == 3.4
True
```


Tests for q2_4 (points: None)
```
>>> all(data_large['hospital_type'] == 'large')
True
```
```
>>> n_pop_large == 3420
True
```
```
>>> n_respond_large == 2397
True
```
```
>>> round(rate_large, 1) == 0.7
True
```
```
>>> round(avg_sat_large, 2) == 4.18
True
```


Tests for q2_5 (points: None)
```
>>> round(weighted_average, 2) == 3.42
True
```


Tests for qc_1 (points: None)
```
>>> prop_respond_small == n_respond_small / n_respond
True
```


Tests for qc_2 (points: None)
```
>>> prop_respond_small_happy == data_small.where('response', 1).where('satisfaction', 5).num_rows / n_respond
True
```


Tests for qc_3 (points: None)
```
>>> prop_no_respond == (n_pop - (n_respond_small + n_respond_medium + n_respond_large)) / n_pop
True
```
