# Lab M1: Discontinuity

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

## Setup

```python
import warnings
from datascience import Table, are
from numpy import mean, sqrt
from statsmodels.formula.api import ols

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Sales Experience

### Business Decision

Reynolds Products is assessing the effectiveness of its new sales staff, which is hired contigent on sales performance during an intial 90-day probation period.  Management suspects that average daily sales for an individual salesperson increases with experience during and beyond the probation period.

Assume that you are the manager of new hire sales for Reynolds Products.  What do you think about the effectiveness of the probation period for new new salespeople?

### Data

Retrieve a dataset from file 'Reynolds.csv'.  Show the first few observations.
Visualize the dataset as a scatterplot.

```python
reynolds = ...
reynolds

reynolds.scatter('Employed')
```

```python
grader.check("q1")
```

### Analysis: One-Piece Model

Build a linear regression model to predict a salesperson's average sales based on number of days employed.
Show the model goodness of fit (R^2).  
Show the model parameters.

```python
model_1 = ...
model_1.rsquared
model_1.params
```

```python
grader.check("q2_1")
```

Add a variable for the predicted average sales. 
Add a variable for the prediction errors. _(these are called "residuals", use the .resid function)_

```python
predicted_sales_1 = ...
prediction_errors_1 = ...

reynolds_predictions = reynolds.with_columns('sales_predicted', predicted_sales_1,
                         'error', prediction_errors_1)
reynolds_predictions
```

```python
grader.check("q2_2")
```

RMSE is $\textit{sqrt}(\textit{mean}(\textit{error}^2))$.  It is a measure of predictive peformance, similar to R^2.

Show the RMSE calculated based on the dataset.  
Visualize the model performance as a scatterplot of average sales and predicted average sales vs. number of days employed.

```python
RMSE_1 = ...
RMSE_1

selected_cols = ...
selected_cols.scatter('Employed')
```

```python
grader.check("q2_3")
```

#### Predict

Predict the average sales of a salesperson that has been employed for 80 days.

```python
sales_80_prediction_1 = model_1.predict(Table().with_columns('Employed', 80))
sales_80_prediction_1
```

```python
grader.check("q2_4")
```

### Analysis: Two-Piece Model

Reset the dataset to include only average sales and number of days employed variables.  
Visualize the dataset as a scatterplot.

```python
sales_employed = ...

sales_employed.scatter('Employed')
```

```python
grader.check("q3_1")
```

Set the breakpoint to 90.  This is the predictor variable value at which an apparent discontinuity occurs.

```python
breakpoint = 90
breakpoint
```

#### Model 1

Filter the dataset to include only sales associated with number of days employed less than or equal to the breakpoint.  
Show the resulting filtered dataset.

```python
sales_employed_below_break = ...
sales_employed_below_break.show()
```

```python
grader.check("q3_2")
```

Build a linear regression model to predict a salesperson's average sales based on number of days employed (use the filtered dataset).  
Show the model goodness of fit (R^2).  
Show the model parameters.

```python
model_below_break = ...
model_below_break.rsquared
model_below_break.params
```

```python
grader.check("q3_3")
```

#### Model 2

Filter the (unfiltered) data to include only sales associated with number of days employed greater than the breakpoint.  
Show the resulting filtered dataset.

```python
sales_employed_above_break = ...
sales_employed_above_break.show()
```

```python
grader.check("q3_4")
```

Build a linear regression model to predict a salesperson's average sales based on number of days employed (use the filtered dataset).  
Show the model goodness of fit (R^2).  
Show the model parameters.

```python
model_above_break = ...
model_above_break.rsquared
model_above_break.params
```

```python
grader.check("q3_5")
```

#### Recombine

Add a variable for predicted sales to each filtered dataset.  
Add a variable for errors to each filtered dataset.  
Build a new dataset that combines the filtered datasets.  
Show the resulting dataset.

```python
predicted_sales_below_break = ...
errors_below_break = ...

predicted_sales_above_break = ...
errors_above_break = ...

below_break_preds = sales_employed_below_break.with_columns('sales_predicted', predicted_sales_below_break, 'error', errors_below_break)
above_break_preds = sales_employed_above_break.with_columns('sales_predicted', predicted_sales_above_break, 'error', errors_above_break)
data_combo = ...

data_combo.show()
```

```python
grader.check("q3_6")
```

Show the RMSE calculated based on the new dataset.  
Visualize the performance of the 2-piece model as a scatterplot of average sales and predicted average sales
vs. number of days employed.

```python
RMSE_combo = ...
RMSE_combo

selected_cols_combo = ...
selected_cols_combo.scatter('Employed')
```

```python
grader.check("q3_7")
```

#### Predict

Predict the average sales of a salesperson that has been employed for 80 days.

```python
sales_80_prediction = ...
sales_80_prediction
```

```python
grader.check("q3_8")
```

Predict the average sales of a salesperson that has been employed for 100 days.

```python
sales_100_prediction = ...
sales_100_prediction
```

```python
grader.check("q3_9")
```

### Analysis: Piecewise Model

Reset dataset to include average sales and number of days employed variables only.
Visualize the dataset as a scatterplot.

```python
sales_employed_2 = ...

sales_employed_2.scatter('Employed')
```

```python
grader.check("q4_1")
```

Set the breakpoint to 90.  This is the predictor variable value at which an apparent discontinuity occurs.  
Add a variable to the datset for switching (1 means number of days employed is greater than breakpoint, 0 means otherwise).  
Add a variable to the dataset for adjustment, like this: $(\textit{number of days employed} - \textit{breakpoint}) \times \textit{switch}$  
Show the resulting dataset.

```python
breakpoint = 90

switch = ...
switch_adjustment = sales_employed_2.with_column('switch', (switch).astype(int))

adjustment = ...
switch_adjustment = switch_adjustment.with_column('adjust', adjustment)

switch_adjustment.show()
```

```python
grader.check("q4_2")
```

Build a linear regression model to predict a salesperson's average sales based on number of days employed and an adjustment.  
Show the model goodness of fit (R^2).  
Show the model parameters.

```python
model_2 = ...
model_2.rsquared
model_2.params
```

```python
grader.check("q4_3")
```

Add a variable to the dataset for predicted sales.  Show the resulting dataset.

```python
predicted_sales_2 = ...

switch_adjustment_preds = switch_adjustment.with_column('sales_predicted', predicted_sales_2)
switch_adjustment_preds.show()
```

```python
grader.check("q4_4")
```

Show the RMSE calculated based on the dataset.
Visualize the performance of the model as a scatterplot of average sales and predicted average sales
vs. number of days employed.

```python
RMSE_2 = ...
RMSE_2

selected_cols_switch_adjust = switch_adjustment_preds.select('Employed','Sales', 'sales_predicted')
selected_cols_switch_adjust.scatter('Employed')
```

```python
grader.check("q4_5")
```

Predict the average sales of a salesperson that has been employed for 80 days.

```python
sales_80_prediction_2 = model_2.predict(Table().with_columns('Employed', 80, 'adjust', (100-breakpoint)*0))
sales_80_prediction_2
```

```python
grader.check("q4_6")
```

Predict the average sales of a salesperson that has been employed for 100 days.

```python
sales_100_prediction_2 = model_2.predict(Table().with_columns('Employed', 100, 'adjust', (100-breakpoint)*1))
sales_100_prediction_2
```

```python
grader.check("q4_7")
```

Predict the average sales of a salesperson that has been employed for 90 days.

```python
sales_90_prediction = model_2.predict(Table().with_columns('Employed', 90, 'adjust', (90-breakpoint)*1))
sales_90_prediction
```

```python
grader.check("q4_8")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. According to the dataset, show the highest actual average sales.

```python
highest_actual = max(reynolds_predictions.column('Sales'))
highest_actual
```

```python
grader.check("qc_1")
```

**Question 2**. According to the piecewise model (model_2), calculate the highest predicted average sales.

```python
highest_predicted = ...
highest_predicted
```

```python
grader.check("qc_2")
```

**Question 3**. According to the piecewise model (model_2) and the dataset, calculate the difference between the RMSE and the mean absolute error between sales and predicted sales.

```python
diff_c3 = ...
diff_c3
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised November 18, 2023
</span>
</p>

Tests for q1 (points: None)
```
>>> reynolds.labels == ('Employed', 'Sales') and reynolds.num_rows == 18
True
```


Tests for q2_1 (points: None)
```
>>> cond1 = list(model_1.model.data.frame.labels) == list(reynolds.labels)
>>> cond2 = model_1.model.data.frame.num_rows == reynolds.num_rows
>>> cond3 = all((model_1.model.data.frame.row(i) == reynolds.row(i) for i in range(reynolds.num_rows)))
>>> cond1 and cond2 and cond3
True
```
```
>>> hasattr(model_1, 'rsquared') and hasattr(model_1, 'params') and (len(model_1.params) == 2)
True
```


Tests for q2_2 (points: None)
```
>>> all(predicted_sales_1 == model_1.predict(reynolds))
True
```
```
>>> all(prediction_errors_1 == model_1.resid)
True
```


Tests for q2_3 (points: None)
```
>>> RMSE_1 == sqrt(mean(reynolds_predictions.column('error') ** 2))
True
```
```
>>> selected_cols.labels == ('Sales', 'sales_predicted', 'Employed')
True
```


Tests for q2_4 (points: None)
```
>>> all(sales_80_prediction_1 == model_1.predict(Table().with_columns('Employed', 80)))
True
```


Tests for q3_1 (points: None)
```
>>> sales_employed.labels == ('Employed', 'Sales') and sales_employed.num_rows == 18
True
```


Tests for q3_2 (points: None)
```
>>> sales_employed_below_break.column('Employed').max() <= breakpoint
True
```
```
>>> sales_employed_below_break.labels == ('Employed', 'Sales') and sales_employed_below_break.num_rows == 12
True
```


Tests for q3_3 (points: None)
```
>>> cond1 = list(model_below_break.model.data.frame.labels) == list(sales_employed_below_break.labels)
>>> cond2 = model_below_break.model.data.frame.num_rows == sales_employed_below_break.num_rows
>>> cond3 = all((model_below_break.model.data.frame.row(i) == sales_employed_below_break.row(i) for i in range(sales_employed_below_break.num_rows)))
>>> cond1 and cond2 and cond3
True
```
```
>>> hasattr(model_below_break, 'rsquared') and hasattr(model_below_break, 'params') and (len(model_below_break.params) == 2)
True
```


Tests for q3_4 (points: None)
```
>>> sales_employed_above_break.column('Employed').min() > breakpoint
True
```


Tests for q3_5 (points: None)
```
>>> cond1 = list(model_above_break.model.data.frame.labels) == list(sales_employed_above_break.labels)
>>> cond2 = model_above_break.model.data.frame.num_rows == sales_employed_above_break.num_rows
>>> cond3 = all((model_above_break.model.data.frame.row(i) == sales_employed_above_break.row(i) for i in range(sales_employed_above_break.num_rows)))
>>> cond1 and cond2 and cond3
True
```
```
>>> hasattr(model_above_break, 'rsquared') and hasattr(model_above_break, 'params') and (len(model_above_break.params) == 2)
True
```


Tests for q3_6 (points: None)
```
>>> all(predicted_sales_below_break == model_below_break.predict(sales_employed_below_break))
True
```
```
>>> all(errors_below_break == model_below_break.resid)
True
```
```
>>> all(predicted_sales_above_break == model_above_break.predict(sales_employed_above_break))
True
```
```
>>> all(errors_above_break == model_above_break.resid)
True
```
```
>>> data_combo.labels == ('Employed', 'Sales', 'sales_predicted', 'error') and data_combo.num_rows == 18
True
```


Tests for q3_7 (points: None)
```
>>> RMSE_combo == sqrt(mean(data_combo.column('error') ** 2))
True
```
```
>>> selected_cols_combo.labels == ('Employed', 'Sales', 'sales_predicted')
True
```


Tests for q3_8 (points: None)
```
>>> all(sales_80_prediction == model_below_break.predict(Table().with_columns('Employed', 80)))
True
```


Tests for q3_9 (points: None)
```
>>> all(sales_100_prediction == model_above_break.predict(Table().with_columns('Employed', 100)))
True
```


Tests for q4_1 (points: None)
```
>>> sales_employed_2.labels == ('Employed', 'Sales') and sales_employed_2.num_rows == 18
True
```


Tests for q4_2 (points: None)
```
>>> all(switch == (switch_adjustment.column('Employed') > breakpoint))
True
```
```
>>> all(adjustment == (switch_adjustment.column('Employed') - breakpoint) * switch_adjustment.column('switch'))
True
```


Tests for q4_3 (points: None)
```
>>> cond1 = list(model_2.model.data.frame.labels) == list(switch_adjustment.labels)
>>> cond2 = model_2.model.data.frame.num_rows == switch_adjustment.num_rows
>>> cond3 = all((model_2.model.data.frame.row(i) == switch_adjustment.row(i) for i in range(switch_adjustment.num_rows)))
>>> cond1 and cond2 and cond3
True
```
```
>>> hasattr(model_2, 'rsquared') and hasattr(model_2, 'params') and (len(model_2.params) == 3)
True
```


Tests for q4_4 (points: None)
```
>>> all(predicted_sales_2 == model_2.predict(switch_adjustment))
True
```


Tests for q4_5 (points: None)
```
>>> RMSE_2 == sqrt(mean(model_2.resid ** 2))
True
```
```
>>> selected_cols_switch_adjust.labels == ('Employed', 'Sales', 'sales_predicted')
True
```


Tests for q4_6 (points: None)
```
>>> all(sales_80_prediction_2 == model_2.predict(Table().with_columns('Employed', 80, 'adjust', (100 - breakpoint) * 0)))
True
```


Tests for q4_7 (points: None)
```
>>> all(sales_100_prediction_2 == model_2.predict(Table().with_columns('Employed', 100, 'adjust', (100 - breakpoint) * 1)))
True
```


Tests for q4_8 (points: None)
```
>>> all(sales_90_prediction == model_2.predict(Table().with_columns('Employed', 90, 'adjust', (90 - breakpoint) * 1)))
True
```


Tests for qc_1 (points: None)
```
>>> highest_actual == 376
True
```


Tests for qc_2 (points: None)
```
>>> all(highest_predicted == model_2.predict(Table().with_columns('Employed', 90, 'adjust', (90 - breakpoint) * 0)))
True
```


Tests for qc_3 (points: None)
```
>>> diff_c3 == RMSE_2 - mean(abs(switch_adjustment_preds.column('Sales') - switch_adjustment_preds.column('sales_predicted')))
True
```
