# Lab H: Confidence Interval

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

## Setup

```python
import warnings
from datascience import Table, make_array, are, percentile
from numpy import mean, std, sqrt, float64
from numpy.random import seed, standard_t

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

```python
sample_size_example = 81

dist_example = Table().with_column('t', standard_t(sample_size_example-1, 1000000))
dist_example.hist(bins=50, range=make_array(-4,4))

cv = percentile(2.5, dist_example.column('t'))
cv

75 + 2*8/sqrt(81)
75 - 2*8/sqrt(81)
```

## Insurance Claims

### Business Decision

An insurance company has a special team of 10 analysts to process complex claims.   The Director of Complex Claims Processing has just learned that the volume of complex claims coming to his team might increase from the current 1500 claims per month to 1700 claims per month, or even as high as 2000 claims per month.  The General Manager asks the Director what impact he expects this will have on the operation.

Should the director accept 1700 claims per month?  Should the director accept 2000 claims per month?

### Data

Here is the record of average labor (in minutes) to process a claim, one average for each of 50 days.  
70, 72, 70, 73, 73, 30, 70, 81, 60, 60, 70, 62, 20, 50, 70, 45, 60, 75, 10, 84,  
21, 60, 66, 81, 32, 24, 60, 67, 67, 56, 45, 49, 80, 45, 39, 67, 80, 87, 50, 51,  
21, 60, 66, 81, 42, 24, 60, 67, 99, 81

Show the sample size, sample mean, sample standard deviation, and a histogram of the sampled averages (10 bins) with range 52.9166 to 64.4033 highlighted.

```python
sample = Table().with_column('average labor', ...)
sample_size = ...
sample_mean = mean(...)
sample_std  = std(..., ddof=1)

print("The sample size is", sample_size)
print("The sample mean is", sample_mean)
print("The sample standard deviation is", sample_std)
sample.hist(bins=10, left_end=52.9166, right_end=64.4033)
```

```python
grader.check("q11")
```

### Analysis: Hypothesis Tests

Get 1,000,000 values from the standard t distribution for sample size 50.
Show a histogram of all the values (50 bins, range -4 to 4).

```python
seed(0)
dist = Table().with_column('t', standard_t(..., ...))
dist.hist(bins=50, range=make_array(-4,4))
```

```python
grader.check("q12")
```

Hypothesize that the population mean is $52.9166$ or lower.

Calculate and show the sample `t`, assuming the hypothesized population mean.

Calculate and show the p-value for the sample, assuming this sample t.
This is the probability of a getting the sample (or one with an equal or higher t value)
if the actual population mean is below or equal to 52.9166. 

```python
pop_mean_hypo_1 = ...
sample_t_1 = ...
p_value_1 = ...

print("The assumed population mean is", pop_mean_hypo_1)
print("The sample t is", sample_t_1)
print("The p-value is", p_value_1)
```

```python
grader.check("q13")
```

Hypothesize that the population mean is 64.4033 or higher.

Calculate and show the sample t, assuming the hypothesized population mean.

Calculate and show the p-value for the sample, assuming this sample t.
This is the probability of a getting the sample (or one with an equal or lower t value)
if the actual population mean is above or equal to 64.4033.

```python
pop_mean_hypo_1 = ...
sample_t_1 = ...
p_value_1 = ...
print("The assumed population mean is", pop_mean_hypo_2)
print("The sample t is", sample_t_2)
print("The p-value is", p_value_2)
```

```python
grader.check("q14")
```

Calculate and show the probability of getting the sample (or one even less likely),
assuming that the population mean is below 52.9166 or above 64.4033.

```python
prob_sample = ...
print("The probablity of getting the sample is", prob_sample)
```

```python
grader.check("q15")
```

### Analysis: Confidence Interval

Get 1,000,000 values from the standard t distribution for sample size 50.
Show a histogram of all the values (50 bins, range -4 to 4).

```python
seed(0)
dist_2 = Table().with_column('t', standard_t(..., ...))
dist_2.hist(bins=50, range=make_array(-4,4))
```

```python
grader.check("q21")
```

Set the confidence level to 95%.  This indicates that you focus on the middle 95% of the standard t distribution.

Calculate and show the probability of t values **below** the middle 95% of the standard t distribution.  
Calculate and show the probability of t values **above** the middle 95% of the standard t distribution. 

```python
confidence_level = ...
prob_lower = ...
prob_upper = ...
print("Our confidence level is", confidence_level)
print("Our probability ot t values below the middle 95% is", prob_lower)
print("Our probability ot t values below the middle 95% is", prob_upper)
```

```python
grader.check("q22")
```

Calculate and show the critical values corresponding to the probabilities.

Show a histogram of all the t values (50 bins, range -4 to 4) with values below the middle 95% highlighted.  
Show a histogram of all the t values (50 bins, range -4 to 4) with values above the middle 95% highlighted.

```python
cv_lower = percentile(..., dist.column('t'))
cv_upper = ...
dist.hist(bins=50, range=make_array(-4,4), left_end=-4, right_end=cv_lower)
dist.hist(bins=50, range=make_array(-4,4), left_end=cv_upper, right_end=4)
print("lower critical value", cv_lower)
print("upper critical value", cv_upper)
```

```python
grader.check("q23")
```

Calculate and show hypothetical population means corresponding to the lower and upper critical values, like this:

$$
\textit{critical value} = \cfrac{\textit{sample mean} - \textit{pop mean hypo}}{\left(\cfrac{\textit{sample standard deviation}}{\sqrt{\textit{sample size}}}\right)}
$$

<br>

$$
\textit{critical value} \times \left(\cfrac{\textit{sample standard deviation}}{\sqrt{\textit{sample size}}} \right) = \textit{sample mean} - \textit{pop mean hypo}  
$$

<br>

$$
\textit{critical value} \times \left(\cfrac{\textit{sample standard deviation}}{\sqrt{\textit{sample size}}} \right) - \textit{sample mean} = - \textit{pop mean hypo}
$$

<br>

$$
- \textit{critical value} \times \left(\cfrac{\textit{sample standard deviation}}{\sqrt{\textit{sample size}}} \right) + \textit{sample mean} = \textit{pop mean hypo}
$$

<br>

$$
\textit{pop mean hypo} =  \textit{sample mean} - \textit{critical value} \times \left(\cfrac{\textit{sample standard deviation}}{\sqrt{\textit{sample size}}} \right)
$$

```python
prob_sample = ...
print("The mean high is ", mean_high)
print("The mean low is ", mean_low)
```

```python
grader.check("q24")
```

Again, calculate and show a hypothetical population mean corresponding to the lower critical value, like this:

$$
\textit{pop mean hypo} =  \textit{sample mean} - \textit{lower critical value} \times \left(\cfrac{\textit{sample standard deviation}}{\sqrt{\textit{sample size}}} \right)
$$

<br>

Note that $\textit{upper critical value} = - \textit{lower critical value}$.  So now, calculate and show a hypothetical population mean corresponding to the upper critical value, like this:

$$
\textit{pop mean hypo} =  \textit{sample mean} + \textit{lower critical value} \times \left(\cfrac{\textit{sample standard deviation}}{\sqrt{\textit{sample size}}} \right)
$$

```python
mean_high = ...
mean_low = ...
print("The mean high is ", mean_high)
print("The mean low is ", mean_low)
```

```python
grader.check("q25")
```

Hence, the formula for confidence interval:

$$
\textit{pop mean hypo} =  \textit{sample mean} \pm \textit{critical value} \times \left(\cfrac{\textit{sample standard deviation}}{\sqrt{\textit{sample size}}} \right)
$$

### Analysis: Business Impact

Recall that the team comprises 10 analysts.

Calculate and show the labor (expressed in minutes) available in a 21-day month,
assuming all the analysts work 8 hours per day.

```python
number_of_analysts = ...
labor_available = ...
print("The number of analysts is ", number_of_analysts)
print("The labor available is ", labor_available)
```

```python
grader.check("q31")
```

Calculate and show how many claims you would expect to process completely in a month,
if the actual average labor per claim is the sample average labor per claim.

Calculate and show how many claims would be left unprocessed, assuming the team receives 2000 claims per month.  
Calculate and show how many claims would be left unprocessed, assuming the team receives 1700 claims per month.  
Calculate and show how many claims would be left unprocessed, assuming the team receives 1500 claims per month.

```python
claims_processed_q32 = ...
unprocessed_2000_q32 = ...
unprocessed_1700_q32 = ...
unprocessed_1500_q32 = ...
print("The number of claims processed is", claims_processed_q32)
print("Unprocessed claims, given 2000 claims per month ", unprocessed_2000_q32)
print("Unprocessed claims, given 1700 claims per month ", unprocessed_1700_q32)
print("Unprocessed claims, given 1500 claims per month ", unprocessed_1500_q32)
```

```python
grader.check("q32")
```

Calculate and show how many claims you would expect to process completely in a month,
if the actual average labor per claim is the lower end of the 95% confidence interval.

Calculate and show how many claims would be left unprocessed, assuming the team receives 2000 claims per month.  
Calculate and show how many claims would be left unprocessed, assuming the team receives 1700 claims per month.  
Calculate and show how many claims would be left unprocessed, assuming the team receives 1500 claims per month.

```python
claims_processed_q33 = ...
unprocessed_2000_q33 = ...
unprocessed_1700_q33 = ...
unprocessed_1500_q33 = ...
print("The number of claims processed is", claims_processed_q33)
print("Unprocessed claims, given 2000 claims per month ", unprocessed_2000_q33)
print("Unprocessed claims, given 1700 claims per month ", unprocessed_1700_q33)
print("Unprocessed claims, given 1500 claims per month ", unprocessed_1500_q33)
```

```python
grader.check("q33")
```

Calculate and show how many claims you would expect to process completely in a month,
if the actual average labor per claim is the upper end of the 95% confidence interval.

Calculate and show how many claims would be left unprocessed, assuming the team receives 2000 claims per month.  
Calculate and show how many claims would be left unprocessed, assuming the team receives 1700 claims per month.  
Calculate and show how many claims would be left unprocessed, assuming the team receives 1500 claims per month.

```python
claims_processed_q34 = ...
unprocessed_2000_q34 = ...
unprocessed_1700_q34 = ...
unprocessed_1500_q34 = ...
print("The number of claims processed is", claims_processed_q34)
print("Unprocessed claims, given 2000 claims per month ", unprocessed_2000_q34)
print("Unprocessed claims, given 1700 claims per month ", unprocessed_1700_q34)
print("Unprocessed claims, given 1500 claims per month ", unprocessed_1500_q34)
```

```python
grader.check("q34")
```

## Hypothesis Tests vs. Confidence Interval: The Math

For your reference, here is the math showing the relationship between hypothesis tests and a confidence interval.

DEFINITIONS

* $t$ is the sample t
* $\bar{x}$ is the sample mean
* $\hat{\mu}$ is a hypothetical population mean
* $s$ is the sample standard deviation
* $n$ is the sample size
<br><br>
* $cv_{\scriptsize{0.025 , n-1}}$ is the critical value at the lower 2.5% of t values in a standard t distribution taken at sample size n
* $cv_{\scriptsize{0.975 , n-1}}$ is the critical value at the lower 97.5% (upper 2.5%) of t values in a standard t distribution taken at sample size n
* $percentile_{\scriptsize{0.025}}(...)$ is the 2.5th percentile of a distribution
* $percentile_{\scriptsize{0.975}}(...)$ is the 97.5th percentile of a distribution
* $\vec{T}_{n-1}$ is a standard t distribution taken at sample size n
<br><br>

VIEWED AS HYPOTHESIS TESTS<br>
Set $\hat{\mu}$, calculate $t$, compare $t$ to $cv$

* $t = \large{\frac{\bar{x} - \hat{\mu}}{\frac{s}{\sqrt{n}}}}$
* $cv_{\scriptsize{0.025 , n-1}} = percentile_{\scriptsize{0.025}}(\vec{T}_{n-1})$
* $cv_{\scriptsize{0.975 , n-1}} = percentile_{\scriptsize{0.975}}(\vec{T}_{n-1}) = - cv_{\scriptsize{0.025 , n-1}}$
<br><br>

VIEWED AS A CONFIDENCE INTERVAL<br>
Set $t$, calculate $\hat{\mu}$ based on $cv$

* If $t = cv_{\scriptsize{0.025 , n-1}}$, then $\hat{\mu} = \bar{x} - cv_{\scriptsize{0.025 , n-1}} \frac{s}{\sqrt{n}}$
* If $t = cv_{\scriptsize{0.975 , n-1}}$, then $\hat{\mu} = \bar{x} + cv_{\scriptsize{0.025 , n-1}} \frac{s}{\sqrt{n}}$

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Assume you have this sample of scores: 17, 23, 20, 17, 40, 20. Calculate the confidence interval at the 95% confidence level (show as an array of 2 values).
(Answer is approximate.)

```python
confidence_level_c1 = 0.95
sample_c1 = make_array(17, 23, 20, 17, 40, 20)
sample_size_c1 = len(sample_c1)
sample_mean_c1 = mean(sample_c1)
sample_std_c1 = std(sample_c1, ddof=1)
dist_c1 = ...
prob_c1 = ...
cv_c1 = ...
mean_high_c1 = ...
mean_low_c1 = ...
ci_c1 = ...
ci_c1
```

```python
grader.check("qc_1")
```

**Question 2**. Assume you have this sample of scores: 17, 23, 20, 17, 40, 20. Calculate the confidence interval at the 90% confidence level (show as an array of 2 values).
(Answer is approximate.)

```python
confidence_level_c2 = 0.90
sample_c2 = make_array(17, 23, 20, 17, 40, 20)
sample_size_c2 = len(sample_c2)
sample_mean_c2 = mean(sample_c2)
sample_std_c2 = std(sample_c2, ddof=1)
dist_c2 = ...
prob_c2 = ...
cv_c2 = ...
mean_high_c2 = ...
mean_low_c2 = ...
ci_c2 = ...
ci_c2
```

```python
grader.check("qc_2")
```

```python
confidence_level_c3 = 0.99
sample_c3 = make_array(17, 23, 20, 17, 40, 20)
sample_size_c3 = len(sample_c3)
sample_mean_c3 = mean(sample_c3)
sample_std_c3  = std(sample_c3, ddof=1)
dist_c3 = ...
prob_c3 = ...
cv_c3 = ...
mean_high_c3 = ...
mean_low_c3 = ...
ci_c3 = ...
ci_c3
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised February 26, 2024
</span>
</p>

Tests for q11 (points: None)
```
>>> type(sample_size) is int and sample_size > 0 and (sample_size == 50)
True
```
```
>>> type(sample_mean) is float64 and sample_mean > 0 and (round(sample_mean) == 59.0)
True
```
```
>>> type(sample) is Table and sample.num_rows == 50 and (round(mean(sample.column('average labor'))) == 59.0)
True
```


Tests for q12 (points: None)
```
>>> round(mean(dist.column('t')), 1) == 0.0
True
```
```
>>> type(dist) is Table and dist.num_rows == 1000000 and (dist.num_columns == 1)
True
```


Tests for q13 (points: None)
```
>>> pop_mean_hypo_1 == 52.9166
True
```
```
>>> round(sample_t_1, 1) == 2.0
True
```
```
>>> round(p_value_1, 3) == 0.025
True
```


Tests for q14 (points: None)
```
>>> pop_mean_hypo_2 == 64.4033
True
```
```
>>> sample_t_2 < 0 and round(sample_t_2, 1) == -2.0
True
```
```
>>> p_value_2 > 0 and round(p_value_2, 3) == 0.025
True
```


Tests for q15 (points: None)
```
>>> prob_sample == 0.049921
True
```


Tests for q21 (points: None)
```
>>> round(dist_2.column('t').item(4)) == 1
True
```
```
>>> round(mean(dist_2.column('t'))) == 0.0
True
```
```
>>> type(dist_2) is Table and dist_2.num_rows == 1000000 and (dist_2.num_columns == 1)
True
```


Tests for q22 (points: None)
```
>>> confidence_level == 0.95
True
```
```
>>> round(prob_lower, 3) == 0.025
True
```
```
>>> assert prob_lower is not Ellipsis
>>> assert prob_upper is not Ellipsis
>>> prob_lower == prob_upper
True
```


Tests for q23 (points: None)
```
>>> cv_lower == -cv_upper
True
```
```
>>> round(cv_lower) == -2.0
True
```


Tests for q24 (points: None)
```
>>> type(mean_high) is float64 and mean_high > 0 and (round(mean_high, 1) == 64.4)
True
```
```
>>> type(mean_low) is float64 and mean_low > 0 and (round(mean_low, 1) == 52.9)
True
```


Tests for q25 (points: None)
```
>>> type(mean_high) is float64 and mean_high > 0 and (round(mean_high, 1) == 64.4)
True
```
```
>>> type(mean_low) is float64 and mean_low > 0 and (round(mean_low, 1) == 52.9)
True
```


Tests for q31 (points: None)
```
>>> type(number_of_analysts) is int and number_of_analysts > 0 and (number_of_analysts == 10)
True
```
```
>>> type(labor_available) is int and labor_available > 0 and (labor_available == 100800)
True
```


Tests for q32 (points: None)
```
>>> type(claims_processed_q32) is int and claims_processed_q32 > 0 and (claims_processed_q32 == 1718)
True
```
```
>>> type(unprocessed_2000_q32) is int and unprocessed_2000_q32 > 0 and (unprocessed_2000_q32 == 282)
True
```
```
>>> type(unprocessed_1700_q32) is int and type(unprocessed_1500_q32) is int and (unprocessed_1700_q32 == 0) and (unprocessed_1500_q32 == 0)
True
```


Tests for q33 (points: None)
```
>>> type(claims_processed_q33) is int and claims_processed_q33 > 0 and (claims_processed_q33 == 1904)
True
```
```
>>> type(unprocessed_2000_q33) is int and unprocessed_2000_q33 > 0 and (unprocessed_2000_q33 == 96)
True
```
```
>>> type(unprocessed_1700_q33) is int and type(unprocessed_1500_q33) is int and (unprocessed_1700_q33 == 0) and (unprocessed_1500_q33 == 0)
True
```


Tests for q34 (points: None)
```
>>> type(claims_processed_q34) is int and claims_processed_q34 > 0 and (claims_processed_q34 == 1565)
True
```
```
>>> type(unprocessed_2000_q34) is int and unprocessed_2000_q34 > 0 and (unprocessed_2000_q34 == 435)
True
```
```
>>> type(unprocessed_1700_q34) is int and unprocessed_1700_q34 > 0 and (unprocessed_1700_q34 == 135)
True
```
```
>>> type(unprocessed_1500_q34) is int and unprocessed_1500_q34 == 0
True
```


Tests for qc_1 (points: None)
```
>>> prob_c1 == (1 - confidence_level) / 2
True
```
```
>>> cv_c1 == percentile(prob_c1 * 100, dist_c1)
True
```
```
>>> mean_high_c1 == sample_mean_c1 - cv_c1 * sample_std_c1 / sqrt(sample_size_c1)
True
```
```
>>> mean_low_c1 == sample_mean_c1 + cv_c1 * sample_std_c1 / sqrt(sample_size_c1)
True
```
```
>>> assert ci_c1 is not Ellipsis
>>> all(ci_c1 == make_array(mean_low_c1, mean_high_c1))
True
```


Tests for qc_2 (points: None)
```
>>> prob_c2 == (1 - confidence_level_c2) / 2
True
```
```
>>> cv_c2 == percentile(prob_c2 * 100, dist_c2)
True
```
```
>>> mean_high_c2 == sample_mean_c2 - cv_c2 * sample_std_c2 / sqrt(sample_size_c2)
True
```
```
>>> mean_low_c2 == sample_mean_c2 + cv_c2 * sample_std_c2 / sqrt(sample_size_c2)
True
```
```
>>> assert ci_c2 is not Ellipsis
>>> all(ci_c2 == make_array(mean_low_c2, mean_high_c2))
True
```


Tests for qc_3 (points: None)
```
>>> prob_c3 == (1 - confidence_level_c3) / 2
True
```
```
>>> cv_c3 == percentile(prob_c3 * 100, dist_c3)
True
```
```
>>> mean_high_c3 == sample_mean_c3 - cv_c3 * sample_std_c3 / sqrt(sample_size_c3)
True
```
```
>>> mean_low_c3 == sample_mean_c3 + cv_c3 * sample_std_c3 / sqrt(sample_size_c3)
True
```
```
>>> assert ci_c3 is not Ellipsis
>>> all(ci_c3 == make_array(mean_low_c3, mean_high_c3))
True
```
