# Lab K in-class: Linear Regression

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

## Setup

```python
import warnings
from datascience import *
import numpy as np
from numpy import mean, median, std, isnan, unique
from statsmodels.formula.api import ols

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## More Investing in Residential Real Estate

### Business Decision

A real estate broker at a San Francisco firm wants to evaluate prices of new properties becoming available.  She needs a model to predict prices based on historical property sales data.

### Data

Retrieve data from file 'redfin.csv'.  Show the first few properties.

```python
redfin = ...
redfin
```

```python
grader.check("q1")
```

### Analysis: Descriptive Statistics & Data Visualization

Calculate and show the median, mean, and standard deviation of price.

```python
median_price = ...
mean_price = ...
std_price = ...

median_price
mean_price
std_price
```

```python
grader.check("q2_1")
```

Visualize the correlation between price and square feet as a scatterplot.

```python
x_col1 = ...
y_col1 = ...

redfin.select(x_col1,y_col1).scatter(x_col1)
```

```python
grader.check("q2_2")
```

### Analysis: Simple Linear Regression

#### Model

Build a simple linear regression model to predict price based on square feet.
Show the model goodness of fit (R^2).
Show the model parameter values (intercept and coefficient of square feet).

$$
\textit{price} \;=\; -702607.21 \;+\; 1526.17 \times \textit{square feet}
$$

```python
model_1 = ...
model_1.rsquared
model_1.params
i_1 = ...
c_1 = ...
i_1
c_1
```

```python
grader.check("q3_1")
```

Visualize the model goodness of fit as a scatterplot with best fit line.

```python
x_col2 = ...
y_col2 = ...

redfin.select(x_col2, y_col2).scatter(x_col2, fit_line=True)
```

```python
grader.check("q3_2")
```

#### Predictions

Use the model parameters to predict the price of a 14,000 square feet property.

```python
sqft_14000 = 14000
price_predicted_14000 = ...
price_predicted_14000
```

```python
grader.check("q4_1")
```

Use the model parameters to predict the price of a 3,000 square feet property.

```python
sqft_3000 = 3000
price_predicted_3000 = ...
price_predicted_3000
```

```python
grader.check("q4_2")
```

Use the model parameters to predict the price of a 600 square feet property.

```python
sqft_600 = 600
price_predicted_600 = ...
price_predicted_600
```

```python
grader.check("q4_3")
```

Use the model parameters to predict the price of a 400 square feet property.

```python
sqft_400 = ...
price_predicted_400 = ...
price_predicted_400
```

```python
grader.check("q4_4")
```

Use the predict method to predict the prices of 14,000, 3,000, 600, and 400 square feet properties.

```python
predictions_1 = ...
predictions_1
```

```python
grader.check("q4_5")
```

Use the predict method to predict the price of 3,000 square feet property.

```python
predictions_2 = ...
predictions_2
```

```python
grader.check("q4_6")
```

#### Back-Predictions

Back-predict prices of all the properties.

```python
all_predictions_1 = ...

redfin_predictions = redfin.with_column('price_predicted', all_predictions_1)
redfin_predictions
```

```python
grader.check("q5_1")
```

Visualize the prices and predicted prices vs. square feet as a scatterplot.

```python
cols_for_scatter_sqft_price = ...

redfin_predictions.select(cols_for_scatter_sqft_price).scatter('SQUARE_FEET')
```

```python
grader.check("q5_2")
```

### Analysis: Multiple Linear Regression

#### Data (for a new model)

Filter redfin_predictions to include only properties that are not missing values for BEDS.  
You can get the those properties like this: `data = data.where(~isnan(data.column('BEDS')))`

Filter redfin_predictions to include only properties that are not missing values for BATHS.  
You can get the those properties like this: `data = data.where(~isnan(data.column('BATHS')))`

Drop the price_predicted variable.

Show the first few properties.

```python
beds_and_baths = ...
beds_and_baths = ...
beds_and_baths = ...
beds_and_baths
```

```python
grader.check("q6_1")
```

#### Model

Build a multiple linear regression model to predict price based on square feet, number of beds, and number of baths.  
You can specify the outcome variable and predictor variables like this: `'PRICE ~ SQUARE_FEET + BEDS + BATHS'`

Show the model goodness of fit (R^2).

Show the model parameter values (intercept and coefficients).

$$
\textit{price} \;=\; -64046.77 \;+\; 1990.33 \times \textit{square feet} \;+\; -586503.50 \times \textit{beds} \;+\; 11648.71 \times \textit{baths}
$$

```python
model_2 = ...
model_2.rsquared
model_2.params
```

```python
grader.check("q6_2")
```

#### Predictions

a. Use the predict method to predict the price of 3,000 square feet, 2 bed, 2 bath property.  
b. Use the predict method to predict the price of 3,000 square feet, 3 bed, 2 bath property.  
c. Use the predict method to predict the price of 4,000 square feet, 2 bed, 2 bath property.  
d. Use the predict method to predict the price of 4,000 square feet, 3 bed, 2 bath property.

Refer to columns in all caps: 'SQUARE_FEET', 'BEDS', and 'BATHS'.

```python
preds_3000_2_2 = ...
preds_3000_3_2 = ...
preds_4000_2_2 = ...
preds_4000_3_2 = ...

preds_3000_2_2, preds_3000_3_2, preds_4000_2_2, preds_4000_3_2
```

```python
grader.check("q6_3")
```

#### Back-Predictions
Back-predict prices of all the properties.

```python
all_predictions_2 = ...

beds_and_baths_predictions = beds_and_baths.with_column('price_predicted', all_predictions_2)
beds_and_baths_predictions
```

```python
grader.check("q6_4")
```

Visualize the prices and predicted prices vs. square feet as a scatterplot.

```python
x_col3 = ...
y1_col3 = ...
y2_col3 = ...

beds_and_baths_predictions.select(x_col3, y1_col3, y2_col3).scatter(x_col3)
```

```python
grader.check("q6_5")
```

Visualize the prices and predicted prices vs. number of beds as a scatterplot.

```python
x_col4 = ...
y1_col4 = ...
y2_col4 = ...

beds_and_baths_predictions.select(x_col4, y1_col4, y2_col4).scatter(x_col4)
```

```python
grader.check("q6_6")
```

Visualize the prices and predicted prices vs. number of baths as a scatterplot.

```python
x_col5 = ...
y1_col5 = ...
y2_col5 = ...

beds_and_baths_predictions.select(x_col5, y1_col5, y2_col5).scatter(x_col5)
```

```python
grader.check("q6_7")
```

### Analysis: Multiple Linear Regression with Dummy Variables

#### Data (for a new model)

Drop the price_predicted variable from beds_and_baths_predictions.

```python
beds_and_baths_2 = ...
beds_and_baths_2
```

```python
grader.check("q7_1")
```

#### Dummy Variables

Show the domain of the PROPERTY_TYPE variable like this: `unique(data.column('PROPERTY_TYPE'))`

```python
domain = unique(beds_and_baths_2.column('PROPERTY_TYPE'))
domain
```

```python
grader.check("q7_2")
```

#### Model with Automatic Conversion to Dummy Variables

Build a multiple linear regression model to predict price based on square feet, number of beds, and number of baths,
whether or not the property is a single family residence, and whether or not the property is a townhouse.  Use PROPERTY_TYPE variable.

Show the model goodness of fit (R^2).

Show the model parameter values (intercept and coefficients).

$$
\textit{price} \;=\; -139693.79 \;+\; 1963.30 \times \textit{square feet} \;+\; -401079.11 \times \textit{beds} \;+\; -16637.97 \times \textit{baths} \;+\; -717094.49 \times \textit{sfr} \;+\; -466476.61 \times \textit{th}
$$

```python
model_3 = ...
model_3.rsquared
model_3.params
```

```python
grader.check("q7_3")
```

#### Predictions with Automatic Conversion to Dummy Variables

a. Use the predict method to predict the price of 3,000 square feet, 3 bed, 2 bath, Condo/Co-op property.  
b. Use the predict method to predict the price of 3,000 square feet, 3 bed, 2 bath, single family residence property.  
c. Use the predict method to predict the price of 3,000 square feet, 3 bed, 2 bath, townhouse property.

Refer to features in all caps: 'SQUARE_FEET', 'BEDS', and 'BATHS', but properties as: 'Condo/Co-op', 'Single Family Residential', and 'Townhouse'.

```python
preds_condo_coop = ...
preds_single = ...
preds_townhouse = ...

preds_condo_coop, preds_single, preds_townhouse
```

```python
grader.check("q7_4")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. According to the second multiple linear regression model (i.e., the one with a property type predictor variable), calculate the difference in price between a 4,000 square feet property and a 3,000 square feet property, all else being equal.

```python
diff_c1 = ...
diff_c1
```

```python
grader.check("qc_1")
```

**Question 2**. According to the second multiple linear regression model (i.e., the one with a property type predictor variable), calculate the difference in price between a 4,000 square feet 3 bathroom property and a 3,000 square feet 2 bathroom property, all else being equal.

```python
diff_c2 = ...
diff_c2
```

```python
grader.check("qc_2")
```

**Question 3**. According to the second multiple linear regression model (i.e., the one with a property type predictor variable), calculate the difference in price between a Single Family Residential property and a Condo/Co-op property, all else being equal.

```python
diff_c3 = ...
diff_c3
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised November 7, 2023
</span>
</p>

Tests for q1 (points: None)
```
>>> redfin.num_columns == 27
True
```
```
>>> redfin.num_rows == 637
True
```


Tests for q2_1 (points: None)
```
>>> round(median_price, 0) == 1299000
True
```
```
>>> round(mean_price, 0) == 1972153
True
```
```
>>> round(std_price, 0) == 2354321
True
```


Tests for q2_2 (points: None)
```
>>> x_col1 == 'SQUARE_FEET'
True
```
```
>>> y_col1 == 'PRICE'
True
```


Tests for q3_1 (points: None)
```
>>> cond1 = list(model_1.model.data.frame.labels) == list(redfin.labels)
>>> cond2 = model_1.model.data.frame.num_rows == redfin.num_rows
>>> cond3 = all((model_1.model.data.frame.row(i) == redfin.row(i) for i in range(redfin.num_rows)))
>>> cond1 and cond2 and cond3
False
```
```
>>> hasattr(model_1, 'rsquared') and hasattr(model_1, 'params') and (len(model_1.params) == 2)
True
```
```
>>> i_1 == model_1.params[0]
True
```
```
>>> c_1 == model_1.params[1]
True
```


Tests for q3_2 (points: None)
```
>>> x_col2 == 'SQUARE_FEET'
True
```
```
>>> y_col2 == 'PRICE'
True
```


Tests for q4_1 (points: None)
```
>>> round(price_predicted_14000, 0) == 20663159
True
```


Tests for q4_2 (points: None)
```
>>> round(price_predicted_3000, 0) == 3875771
True
```


Tests for q4_3 (points: None)
```
>>> round(price_predicted_600, 0) == 213068
True
```


Tests for q4_4 (points: None)
```
>>> round(price_predicted_400, 0) == -92157
True
```


Tests for q4_5 (points: None)
```
>>> all(predictions_1 == model_1.predict(Table().with_column('SQUARE_FEET', make_array(14000, 3000, 600, 400))))
True
```


Tests for q4_6 (points: None)
```
>>> all(predictions_2 == model_1.predict(Table().with_column('SQUARE_FEET', 3000)))
True
```


Tests for q5_1 (points: None)
```
>>> all(all_predictions_1 == model_1.predict(redfin))
True
```


Tests for q5_2 (points: None)
```
>>> cols_for_scatter_sqft_price == ('SQUARE_FEET', 'PRICE', 'price_predicted')
True
```


Tests for q6_1 (points: None)
```
>>> not isnan(beds_and_baths.column('BEDS')).any()
True
```
```
>>> not isnan(beds_and_baths.column('BATHS')).any()
True
```
```
>>> 'price_predicted' not in beds_and_baths.labels and beds_and_baths.num_rows > 0
True
```


Tests for q6_2 (points: None)
```
>>> cond1 = list(model_2.model.data.frame.labels) == list(beds_and_baths.labels)
>>> cond2 = model_2.model.data.frame.num_rows == beds_and_baths.num_rows
>>> cond3 = all((model_2.model.data.frame.row(i) == beds_and_baths.row(i) for i in range(beds_and_baths.num_rows)))
>>> cond1 and cond2 and cond3
False
```
```
>>> hasattr(model_2, 'rsquared') and hasattr(model_2, 'params') and (len(model_2.params) == 4)
True
```


Tests for q6_3 (points: None)
```
>>> all(preds_3000_2_2 == model_2.predict(Table().with_columns('SQUARE_FEET', 3000, 'BEDS', 2, 'BATHS', 2)))
True
```
```
>>> all(preds_3000_3_2 == model_2.predict(Table().with_columns('SQUARE_FEET', 3000, 'BEDS', 3, 'BATHS', 2)))
True
```
```
>>> all(preds_4000_2_2 == model_2.predict(Table().with_columns('SQUARE_FEET', 4000, 'BEDS', 2, 'BATHS', 2)))
True
```
```
>>> all(preds_4000_3_2 == model_2.predict(Table().with_columns('SQUARE_FEET', 4000, 'BEDS', 3, 'BATHS', 2)))
True
```


Tests for q6_4 (points: None)
```
>>> all(all_predictions_2 == model_2.predict(beds_and_baths))
True
```


Tests for q6_5 (points: None)
```
>>> x_col3 == 'SQUARE_FEET'
True
```
```
>>> y1_col3 == 'PRICE'
True
```
```
>>> y2_col3 == 'price_predicted'
True
```


Tests for q6_6 (points: None)
```
>>> x_col4 == 'BEDS'
True
```
```
>>> y1_col4 == 'PRICE'
True
```
```
>>> y2_col4 == 'price_predicted'
True
```


Tests for q6_7 (points: None)
```
>>> x_col5 == 'BATHS'
True
```
```
>>> y1_col5 == 'PRICE'
True
```
```
>>> y2_col5 == 'price_predicted'
True
```


Tests for q7_1 (points: None)
```
>>> 'price_predicted' not in beds_and_baths_2.labels
True
```


Tests for q7_2 (points: None)
```
>>> all(domain == unique(beds_and_baths_2.column('PROPERTY_TYPE')))
True
```


Tests for q7_3 (points: None)
```
>>> cond1 = list(model_3.model.data.frame.labels) == list(beds_and_baths_2.labels)
>>> cond2 = model_3.model.data.frame.num_rows == beds_and_baths_2.num_rows
>>> cond3 = all((model_3.model.data.frame.row(i) == beds_and_baths_2.row(i) for i in range(beds_and_baths_2.num_rows)))
>>> cond1 and cond2 and cond3
False
```
```
>>> hasattr(model_3, 'rsquared') and hasattr(model_3, 'params') and (len(model_3.params) == 6)
True
```


Tests for q7_4 (points: None)
```
>>> all(preds_condo_coop == model_3.predict(Table().with_columns('SQUARE_FEET', 3000, 'BEDS', 3, 'BATHS', 2, 'PROPERTY_TYPE', 'Condo/Co-op')))
True
```
```
>>> all(preds_single == model_3.predict(Table().with_columns('SQUARE_FEET', 3000, 'BEDS', 3, 'BATHS', 2, 'PROPERTY_TYPE', 'Single Family Residential')))
True
```
```
>>> all(preds_townhouse == model_3.predict(Table().with_columns('SQUARE_FEET', 3000, 'BEDS', 3, 'BATHS', 2, 'PROPERTY_TYPE', 'Townhouse')))
True
```


Tests for qc_1 (points: None)
```
>>> diff_c1 == model_3.params[3] * (4000 - 3000)
True
```


Tests for qc_2 (points: None)
```
>>> diff_c2 == model_3.params[3] * (4000 - 3000) + model_3.params[5] * (3 - 2)
True
```


Tests for qc_3 (points: None)
```
>>> diff_c3 == model_3.params[1] * (1 - 0)
True
```
