# Lab F: One-Sample t Hypothesis Test

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

## Setup

```python
import warnings
from datascience import Table, make_array, are, percentile
from numpy import mean, std, arange, sqrt, append, ndarray, float64, isclose
from numpy.random import seed, standard_t

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Customer Satisfaction Alarm - known population (unrealistic business situation)

### Business Decision

A customer service organization in a large multi-national conglomerate serves about $100,000$ customers.  Each month, an independent vendor interviews $10$ customers chosen at random for their levels of customer satisfaction.  The scores range from $0$ for highly dissatisfied to $100$ for highly satisfied.  The vendor then reports the average score for this sample of customers to the Vice President of Customer Service.  Historically, the average score for a sample of $10$ customers has usually been about $75$, which the Vice President considers acceptable.

For this analysis, suppose (somewhat unrealistically) that you know the satisfaction levels of all customers. 

### Data

**Question 0.1.** Retrieve the satisfaction levels of all customers from the file `'Population CSat1.csv'`. Show the population size, the population mean score, and a histogram of the population scores (100 bins, range 0 to 100).

```python
data = Table.read_table(...)
pop_mean = ...
data
print("There are", data.num_rows, "in the dataset.")
print("The mean of the population is", pop_mean)
data.hist(bins=100, range=make_array(0,100))
```

```python
grader.check("q01")
```

### 1. Analysis

#### Sample 

**Question 1.1.** Set sample size to `10`.

```python
sample_size = ...
sample_size
```

```python
grader.check("q11")
```

**Question 1.2.** Set the randomization seed to `4`, get a sample, show the sample scores, and show a histogram of the sample scores (100 bins, range 0 to 100).

```python
seed_value = ...
seed(seed_value)

sampled_data = ...

num_bins = ...
range_bins = ...
sampled_data.hist(bins=num_bins, range=range_bins)
```

```python
grader.check("q12")
```

**Question 1.3** Show the sample mean.

```python
sample_mean = ...
sample_mean
```

```python
grader.check("q13")
```

**Question 1.4** Show the sample standard deviation.

```python
sample_std = ...
sample_std
```

```python
grader.check("q14")
```

**Question 1.5** Calculate the sample t-statistic to test whether your sample mean is significantly different from the population mean of $75.00$.

Use the formula:

$$
t = \frac{\bar{x} - \text{pop. mean}}{s / \sqrt{n}}
$$

where:  
- $\bar{x}$ is the sample mean  
- $s$ is the sample standard deviation  
- $n$ is the sample size

```python
sample_t = ...
sample_t
```

```python
grader.check("q15")
```

#### 2. Distribution of Sample t's

**Question 2.1** Get a new sample of $10$ customers. Show the sample scores. Calculate and show the sample t.  Visualize as a histogram of the **sample scores** (100 bins, range 0 to 100).

Try this several times and notice that the shapes of the histograms are dissimilar.

```python
new_sample_i =  ...
sample_mean_i = ...
sample_std_i = ...
sample_t_i = (sample_mean_i - pop_mean) / (sample_std_i / sqrt(sample_size))

new_sample_i
sample_t_i
new_sample_i.hist(bins=100, range=make_array(0,100))
```

```python
grader.check("q21")
```

**Question 2.2** Get $5$ new samples of $10$ customers each.  Calculate and show the sample t's.  Visualize as a histogram of the **sample t's** (50 bins, range -4 to 4).

Try this several times and notice that the shapes of the histograms are dissimilar.

```python
dist_array_5 = make_array()

for i in ...:
    sample_i = ...
    sample_mean_i = ...
    sample_std_i = ...
    sample_t_i = ...
    dist_array_5 = append(..., ...)
dist_5 = Table().with_column('t', dist_array_5)

dist_5
dist_5.hist(bins=50, range=make_array(-4,4))
```

```python
grader.check("q22")
```

**Question 2.3** Get $100,000$ new samples of $10$ customers each.  Calculate and show the first few sample t's.  Visualize as a histogram of the sample `t`'s (50 bins, range -4 to 4).

Try this several times and notice that the shapes of the histograms are similar.

```python
dist_array_100k = make_array()

for i in ...:
    sample_i = ...
    sample_mean_i = ...
    sample_std_i = ...
    sample_t_i = ...                  # caution: potential divide by zero
    dist_array_100k = append(..., ...)
dist = Table().with_column('t', dist_array_100k)

dist
dist.hist(bins=50, range=make_array(-4,4))
```

```python
grader.check("q23")
```

**Question 2.4** Calculate and show the probability of getting a sample `t` $\le -0.6687$ (when population mean is $75.00$ and sample size is $10$).

Visualize as a histogram of standard `t`'s with the area corresponding to the probability highlighted (50 bins, range -4 to 4).

```python
prob = ...
print("Our sample t is", sample_t)
print("Our probability of observing a t as small, or even smaller than", sample_t, "is", prob)
dist.hist(bins=50, range=make_array(-4,4), left_end=-4, right_end=sample_t)
```

```python
grader.check("q24")
```

There was a $26.14$\% probability of getting a sample like the one we did.  That's not a surprising result.

## 3. Customer Satisfaction Alarm - unknown population (realistic business situation)

### Business Decision

As in Part 2, our customer service organization has hired a vendor to interview some of its customers monthly.  For this analysis, suppose (somewhat realistically) that you do NOT know the satisfaction levels of all customers.

This month, the average of 10 customer satisfaction scores is far below $75$.  The Vice President of Customer Service must decide whether to initiate action to address low customer satisfaction.

### Data

You get a sample with scores that look like this:  
$$ 63.1142, 86.2235, 75.3150, 62.5061, 73.3992, 67.3861, 53.3442, 67.4657, 76.6753, 67.5086 $$

```python
observed_sample = Table().with_column('score', make_array(63.1142, 86.2235, 75.3150, 62.5061, 73.3992, 67.3861, 53.3442, 67.4657, 76.6753, 67.5086))
observed_sample
```

### Analysis

#### Hypothesis

Hypothesize that the population mean $\ge$ $75.00$.  This would be business as usual and not require action.

```python
pop_mean_hypo = 75.00
pop_mean_hypo
```

#### Sample

**Question 3.1** Show the sample size, sample mean, sample standard deviation, sample `t`, and a histogram of the sample scores (100 bins, range 0 to 100).

```python
sample_size_obs = ...
sample_mean_obs = ...
sample_std_obs = ...
sample_t_obs = (sample_mean_obs - pop_mean_hypo) / (sample_std_obs / sqrt(sample_size_obs))
print("Our sample size is", sample_size_obs)
print("Our sample mean is", sample_mean_obs)
print("Our sample standard deviation is", sample_std_obs)
print("Our sample t is", sample_t_obs)
observed_sample.hist(bins=100, range=make_array(0,100))
```

```python
grader.check("q31")
```

Because the sample mean is less than the hypothesized population mean, we suspect that the real population mean might be less than the hypothesized population mean.  Note that what we suspect is the opposite of the hypothesis. 

#### Standard t Distribution

**Question 3.2** Get $1,000,000$ values from the standard `t` distribution for sample size $10$.  To do this, use `standard_t(sample_size - 1, 1000000)`.  Visualize as a histogram of all the standard `t`'s. (50 bins, range -4 to 4).

Treat the standard `t` distribution as the population `t` distribution.

```python
degrees_freedom = ...
dist_array_obs = standard_t(degrees_freedom, ...)
dist_obs = Table().with_column('t', dist_array_obs)
dist_obs
dist_obs.hist(bins=50, range=make_array(-4,4))
```

```python
grader.check("q32")
```

#### p-Value

**Question 3.3** Show the sample `t`.  Calculate and show the probability that the sample `t` $\le -1.9859$ when `pop mean` $\ge 75.00$ (this is the p-value).  Visualize as a histogram of standard t's with the area corresponding to the probability highlighted (50 bins, range -4 to 4).

```python
num_below_sample_t = ...
p_value = num_below_sample_t / dist_obs.num_rows
print("Our sample t was found to be", sample_t_obs)
print("Our p-value is", p_value)
dist_obs.hist(bins=50, range=make_array(-4,4), left_end=-4, right_end=sample_t)
```

```python
grader.check("q33")
```

#### Conclusion at Significance Level $0.05$

**Question 3.4** Set and show the significance level $0.05$.  Calculate and show the critical value at significance level $0.05$ based on the standard t distribution for sample size $10$.  Note that the suspected population mean is **below** the hypothesized population mean.  Visualize as a histogram of standard t's with the area corresponding to the significance level highlighted (50 bins, range -4 to 4).

```python
sig_level_05 = ...
cv = percentile(sig_level_05*100, ...)
print("Our significance level is", sig_level_05)
print("Our critical value is", cv)
dist_obs.hist(bins=50, range=make_array(-4,4), left_end=-4, right_end=cv)
```

```python
grader.check("q34")
```

**Question 3.5** Calculate and show that you should conclude the hypothesis is not correct (i.e., conclude population mean < 75.00), at significance level 0.05.

```python
do_we_reject_null = p_value <= ... and sample_t_obs <= ...
print("Do we reject the null hypothesis?", do_we_reject_null)
```

```python
grader.check("q35")
```

If average customer satisfaction of all customers were really $75$ or higher, then there would be only a 3.39% probability of getting a sample like the one we did.  Although that's possible, it's very unlikely, so the Vice President of Customer Service decides to initiate action to address low customer satisfaction.

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Assume a significance level of 0.01, rather than 0.05. Calculate the critical value.

```python
sig_level_c1 = 0.01
cv_c1 = ...
cv_c1
```

```python
grader.check("qc_1")
```

**Question 2**. Using your critical value from Question 1, calculate whether the manager should conclude that mean score is above or equal to 75.
(Answer is True or False.)

```python
conclusion_c2 = ...
conclusion_c2
```

```python
grader.check("qc_2")
```

**Question 3**. AAssume a significance level of 0.10, rather than 0.05. Calculate whether the manager should conclude that mean score is above or equal to 75.
(Answer is True or False.)

```python
sig_level_1 = 0.10
cv_c3 = percentile(sig_level_1*100, dist_obs.column('t'))
conclusion = cv_c3 < sample_t_obs
conclusion
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised October 5, 2024
</span>
</p>

Tests for q01 (points: None)
```
>>> type(data) is Table and data.num_rows == 99911 and isclose(pop_mean, 74.98, atol=0.01)
True
```
```
>>> isclose(data.column(0).item(0), 78.02, atol=0.01)
True
```


Tests for q11 (points: None)
```
>>> type(sample_size) is int and sample_size > 0 and (sample_size == 10)
True
```


Tests for q12 (points: None)
```
>>> type(seed_value) is int and seed_value > 0 and (seed_value == 4)
True
```
```
>>> sampled_data.num_rows == 10 and len(sampled_data.columns) == 1 and isclose(mean(sampled_data.column(0)), 73.2, atol=0.1)
True
```
```
>>> type(num_bins) is int and num_bins > 0 and (num_bins == 100)
True
```
```
>>> type(range_bins) is ndarray and len(range_bins) == 2 and (range_bins.item(0) == 0) and (range_bins.item(1) == 100)
True
```


Tests for q13 (points: None)
```
>>> type(sample_mean) is float64 and isclose(sample_mean, 73.2, atol=0.1)
True
```


Tests for q14 (points: None)
```
>>> type(sample_std) is float64 and isclose(sample_std, 8.19, atol=0.01)
True
```


Tests for q15 (points: None)
```
>>> isclose(sample_t, -0.668, atol=0.001)
True
```


Tests for q21 (points: None)
```
>>> new_sample_i.num_rows == 10 and len(new_sample_i.columns) == 1
True
```
```
>>> type(sample_mean_i) is float64 and -10 <= sample_t_i <= 10
True
```


Tests for q22 (points: None)
```
>>> sample_i.num_rows == 10 and len(sample_i.columns) == 1
True
```
```
>>> sum(dist_array_5 >= -10) == 5 and sum(dist_array_5 <= 10) == 5
True
```


Tests for q23 (points: None)
```
>>> sum(dist_array_100k >= -10) == 100000 and sum(dist_array_100k <= 10) == 100000
True
```


Tests for q24 (points: None)
```
>>> type(prob) is float and isclose(prob, 0.26, atol=0.01)
True
```


Tests for q31 (points: None)
```
>>> sample_size_obs * 100 == 1000
True
```
```
>>> isclose(sample_mean_obs, 69.3, atol=0.01) and type(sample_mean_obs) is float64 and (sample_mean_obs > 0)
True
```
```
>>> isclose(sample_std_obs, 9.08, atol=0.01) and type(sample_mean_obs) is float64 and (sample_mean_obs > 0)
True
```
```
>>> round(sample_t_obs) == -2 and type(sample_std_obs) is float64 and (sample_std_obs > 0)
True
```


Tests for q32 (points: None)
```
>>> dist_obs.num_rows == 1000000 and len(dist_obs.columns) == 1
True
```
```
>>> round(mean(dist_array_obs)) == 0
True
```


Tests for q33 (points: None)
```
>>> type(p_value) is float and p_value <= 0.1 and (p_value >= 0.01)
True
```


Tests for q34 (points: None)
```
>>> sig_level_05 == 0.05 and type(sig_level_05) is float
True
```
```
>>> isclose(cv, -1.836, atol=0.01)
True
```


Tests for q35 (points: None)
```
>>> do_we_reject_null
True
```


Tests for qc_1 (points: None)
```
>>> cv_c1 == percentile(sig_level_c1 * 100, dist_obs.column('t'))
True
```


Tests for qc_2 (points: None)
```
>>> conclusion_c2 == True
True
```


Tests for qc_3 (points: None)
```
>>> conclusion == False
True
```
