# Lab A: Data Management

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

## Setup

```python
from datascience import Table

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

## Competing with a Big Box Retailer

### Business Decision

A general retailer that competes with Best Buy wants to determine how well Best Buy is run.  It hopes to craft its own corporate strategy to exploit any Best Buy weakness it might discover. (A company's corporate strategy is its decision about where to apply its limited resources.)

### Data

Retrieve Best Buy data, show the number of observations and the first few observations.

```python
data = ...
...
data
```

```python
grader.check("q1")
```

### Analysis

Show the first year and last year that revenue was recorded.  
Show the time span (in years) over which revenue was recorded. Do not round to nearest integer.

```python
first_year = ...
last_year = ...
span = ...

first_year
last_year
span
```

```python
grader.check("q2_1")
```

Show the quarterly revenue amounts.

```python
revenue = ...
revenue
```

```python
grader.check("q2_2")
```

Calculate the sum of Best Buy's revenue over the time span.

```python
total_revenue = ...
total_revenue
```

```python
grader.check("q2_3")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Show the lowest annual revenue within the time span.

```python
min_revenue = ...
min_revenue
```

```python
grader.check("qc_1")
```

**Question 2**. Calculate the average annual revenue over the time span.

```python
average_annual_revenue = ...
average_annual_revenue
```

```python
grader.check("qc_2")
```

**Question 3**. What is the average cost of goods sold (COGS)?

```python
average_cogs = ...
average_cogs
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised August 22, 2023
</span>
</p>

Tests for q1 (points: None)
```
>>> data.num_rows == 83
True
```
```
>>> data.labels == ('Year', 'Quarter', 'Revenue', 'Cost of Goods Sold', 'Gross Profit (million)')
True
```


Tests for q2_1 (points: None)
```
>>> first_year == 1995
True
```
```
>>> last_year == 2015
True
```
```
>>> span == 20.75
True
```


Tests for q2_2 (points: None)
```
>>> len(revenue) == 83
True
```


Tests for q2_3 (points: None)
```
>>> round(total_revenue) == 607015
True
```


Tests for qc_1 (points: None)
```
>>> min_revenue == 1274.7
True
```


Tests for qc_2 (points: None)
```
>>> round(average_annual_revenue, 2) == 29253.71
True
```


Tests for qc_3 (points: None)
```
>>> round(average_cogs, 2) == 5487.49
True
```
