# Lab G: ANOVA Test

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

## Setup

```python
import warnings
from datascience import Table, make_array, are, percentile
import numpy as np
from numpy import mean, unique, array, var
from numpy.random import f

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Shipping Hazardous Materials

### Business Decision

A large logistics company wants to determine whether three different training delivery modes have different effects on employees' ability to pass a test on shipping hazardous materials.  The three delivery modes are in-class live instructor, online asynchronous recorded instructor, and individualized tutors.

Assume that you are the Vice President of Operations.  Individualized tutoring is the most expensive delivery mode.  Online is the least expensive delivery mode.  Is there a difference in the effectiveness of the delivery modes?  Which delivery mode(s) should you use?

### Data

Retrieve 3 samples of employee test scores, one sample from each delivery mode. 

Show the treatment count (number of delivery modes $c$), unit count (number of employees $n$),
and the samples (in one two-column table).

```python
data = ...

c = ...
n = ...

c
n
data.show()
```

```python
grader.check("q1")
```

### Analysis

Hypothesize that the three delivery mode population mean scores are all equal.  The alternative to this hypothesis is that the three delivery mode population mean scores are not all equal.

```python
# no code necessary
```

For convenience, extract each sample into its own array ($x_1$, $x_2$, $x_3$). 

```python
x1 = ...
x2 = ...
x3 = ...

x1
x2
x3
```

```python
grader.check("q2_1")
```

Calculate and show the weighted average of sample means like this:

$$
\textit{mmx} =
\cfrac{(\textit{sample size}_1 \times \textit{sample mean}_1) +
       (\textit{sample size}_2 \times \textit{sample mean}_2) +
       (\textit{sample size}_3 \times \textit{sample mean}_3)}
      {n}
$$

```python
mmx = ...
mmx
```

```python
grader.check("q2_2")
```

Calculate and show the "treatment sum of squares" like this:

$$
sst = 
\left(\textit{sample size}_1 \times (\textit{sample mean}_1 - \textit{mmx})^2\right) +
\left(\textit{sample size}_2 \times (\textit{sample mean}_2 - \textit{mmx})^2\right) +
\left(\textit{sample size}_3 \times (\textit{sample mean}_3 - \textit{mmx})^2\right)
$$

```python
sst = ...
sst
```

```python
grader.check("q2_3")
```

Calculate and show the "treatment mean of squares" (mean variance between groups) like this:

$$
\textit{mst} = 
\cfrac{\textit{sst}}{c-1}
$$



```python
mst = ...
mst
```

```python
grader.check("q2_4")
```

Calculate and show the "error sum of squares" like this:

$$
\textit{sse} =
\sum{(\textit{sample}_1-\textit{sample mean}_1)^2} + \sum{(\textit{sample}_2-\textit{sample mean}_2)^2} + \sum{(\textit{sample}_3-\textit{sample mean}_3)^2}
$$

```python
sse = ...
sse
```

```python
grader.check("q2_5")
```

Calculate and show the "error mean of squares" like this:

$$
\textit{mse}  =
\cfrac{\textit{sse}}{n-c}
$$

(This is the mean variance within groups.)

```python
mse = ...
mse
```

```python
grader.check("q2_6")
```

Calculate and show the multi-sample f like this:


$$
\textit{multi-sample f} =
\cfrac{\textit{mst}}{\textit{mse}}
$$

```python
multi_sample_f = ...
multi_sample_f
```

```python
grader.check("q2_7")
```

Calculate and show the upper and lower degrees of freedom for the samples.
These numbers will be used to choose a specific f distribution.

$$
\begin{array}{rcl}
\textit{df upper} &=& c-1 \\  
\textit{df lower} &=& n-c
\end{array}
$$

```python
df_upper = ...
df_lower = ...

df_upper
df_lower
```

```python
grader.check("q2_8")
```

Get 1,000,000 values from the f distribution for the appropriate degrees of freedom.

Show a few of the values and a histogram of all the values (50 bins, range 0 to 10).

```python
dist_array = ...
dist = Table().with_column('f', dist_array)
dist

dist_hist = ...
dist_hist
```

```python
grader.check("q2_9")
```

Calculate and show the probability that the multi-sample f $\ge$ 1.94 when the population mean scores
are all equal (this is the p-value).
Also show the multi-sample f and histogram of standard f distribution with the area corresponding to the
probability highlighted.

```python
p_value_anova = ...
values_for_hist_anova = ...

multi_sample_f
p_value_anova
dist.hist(bins=50, range=make_array(0,10), left_end=values_for_hist_anova, right_end=10)
```

```python
grader.check("q2_10")
```

Calculate and show the critical value at significance level 0.05.
Also show the significance level and histogram of standard f distribution with the area corresponding to
the significance level highlighted.

```python
sig_level_anova = 0.05
cv_anova = ...

sig_level_anova
cv_anova
dist.hist(bins=50, range=make_array(0,10), left_end=cv_anova, right_end=10)
```

```python
grader.check("q2_11")
```

Show that you should assume the hypothesis is correct (all population mean scores are equal), at significance level 0.05. You should only use comparison symbols (>, <, =) in the blank. Define reason_1 and reason_2 as the full string expressing the statement. For example, reason_1 = "p_value = sig_level". Don't include any extra spaces.

```python
# Reason 1: "p_value ___ sig_level"
anova_reason_1 = ...

# Reason 2: "multi_sample_f ___ cv"
anova_reason_2 = ...

anova_reason_1, anova_reason_2
```

```python
grader.check("q2_12")
```

Show the ANOVA table for analysis.

```python
sum_squares_values = ...
df_values = ...
mean_squares_values = ...
f_values = ...

anova_table = Table().with_columns('source of variation', make_array('between groups','within groups','total'),
                     'sum squares', sum_squares_values,
                     'df', df_values,
                     'mean squares', mean_squares_values,
                     'f', f_values)

anova_table
```

```python
grader.check("q2_13")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Assume a significance level of 0.01, rather than 0.05. Calculate the critical value.

```python
sig_level_one = 0.01
cv_c1 = ...
cv_c1
```

```python
grader.check("qc_1")
```

**Question 2**. Assume a significance level of 0.01, rather than 0.05. Calculate whether the manager should conclude that all 3 delivery modes are equally effective.
(Answer is True or False.)

```python
sig_level_two = 0.01
cv_c2 = ...
conclusion = ...
conclusion
```

```python
grader.check("qc_2")
```

**Question 3**. Assume a significance level of 0.10, rather than 0.05. Calculate whether the manager should conclude that all 3 delivery modes are equally effective.
(Answer is True or False.)

```python
sig_level_three = 0.10
cv_c3 = ...
conclusion_2 = ...
conclusion_2
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised October 5, 2023
</span>
</p>

Tests for q1 (points: None)
```
>>> set(data.labels) >= {'delivery', 'score'}
True
```
```
>>> c == 3
True
```
```
>>> n == 14
True
```


Tests for q2_1 (points: None)
```
>>> all(x1 == array([85, 72, 83, 80]))
True
```
```
>>> all(x2 == array([80, 84, 81, 78, 82]))
True
```
```
>>> all(x3 == array([82, 80, 85, 90, 88]))
True
```


Tests for q2_10 (points: None)
```
>>> round(p_value_anova, 2) == 0.19
True
```
```
>>> values_for_hist_anova == multi_sample_f
True
```


Tests for q2_11 (points: None)
```
>>> round(cv_anova, 0) == 4
True
```


Tests for q2_12 (points: None)
```
>>> (anova_reason_1 == 'p_value > sig_level') | (anova_reason_1 == 'p_value>sig_level')
True
```
```
>>> (anova_reason_2 == 'multi_sample_f < cv') | (anova_reason_2 == 'multi_sample_f<cv')
True
```


Tests for q2_13 (points: None)
```
>>> all(sum_squares_values == make_array(sst, sse, sst + sse))
True
```
```
>>> all(df_values == make_array(df_upper, df_lower, df_upper + df_lower))
True
```
```
>>> all(mean_squares_values == make_array(mst, mse, None))
True
```
```
>>> all(f_values == make_array(multi_sample_f, None, None))
True
```


Tests for q2_2 (points: None)
```
>>> round(mmx, 0) == 82
True
```


Tests for q2_3 (points: None)
```
>>> round(sst, 0) == 66
True
```


Tests for q2_4 (points: None)
```
>>> round(mst, 0) == 33
True
```


Tests for q2_5 (points: None)
```
>>> round(sse, 0) == 186
True
```


Tests for q2_6 (points: None)
```
>>> round(mse, 0) == 17
True
```


Tests for q2_7 (points: None)
```
>>> round(multi_sample_f, 0) == 2
True
```


Tests for q2_8 (points: None)
```
>>> df_upper == 2
True
```
```
>>> df_lower == 11
True
```


Tests for q2_9 (points: None)
```
>>> col_values = dist.column(0)
>>> all(col_values >= 0) and mean(col_values) > 0 and (len(dist.column(0)) == 1000000)
True
```
```
>>> data_col = dist.column(0)
>>> column_is_floats = all((isinstance(x, float) for x in data_col[:100]))
>>> column_is_positive = all((x >= 0 for x in data_col[:100]))
>>> column_variance = var(data_col) > 0
>>> column_is_floats and column_is_positive and column_variance and (dist_hist is None)
True
```


Tests for qc_1 (points: None)
```
>>> cv_c1 == percentile((1 - sig_level_one) * 100, dist.column('f'))
True
```


Tests for qc_2 (points: None)
```
>>> conclusion == True
True
```


Tests for qc_3 (points: None)
```
>>> conclusion_2 == True
True
```
