# Lab G: Two-Sample t Hypothesis Test

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

## Setup

```python
import warnings
from datascience import Table, make_array, are, percentile
from numpy import mean, std, sqrt, isfinite
from numpy.random import standard_t


# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Precision Manufacturing Quality

### Business Decision

A manufacturing organization uses 2 special 3D printers to produce precision parts, specificallly 10cm-diameter tubes.  Recently, the machine operators have become suspicious that the first machine produces tubes on average slightly larger than the second machine on average does.  If so, that would necessitate expensive action to get the machines aligned.

Assume that you are the Vice President of Manufacturing.  Does the second machine produce on average tubes of smaller diameter than the first machine does?  Should you invest in aligning the machines?

### Data

Machine #1 produces a sample of 12 tubes with these diameters:  
10.03, 10.04, 10.02, 9.98, 9.99, 10.04, 10.01, 10.04, 10.05, 9.95, 9.96, 9.99

Show the sample size, sample mean, sample standard deviation, and
a histogram of the sample diameters (10 bins, range 9.95 to 10.1).

```python
sample_1 = Table().with_column('score', make_array(10.03, 10.04, 10.02, 9.98, 9.99, 10.04,
                                                   10.01, 10.04, 10.05, 9.95, 9.96, 9.99))

sample_size_1 = ...
sample_mean_1 = ...
sample_std_1  = ...

sample_size_1
sample_mean_1
sample_std_1
sample_1.hist(bins=10, range=make_array(9.95,10.1))
```

```python
grader.check("q1_1")
```

Machine #2 produces a sample of 12 tubes with these diameters:  
10.03, 9.75, 10.01, 9.75, 9.99, 10.01, 10.01, 10.08, 10.08, 9.95, 9.96, 9.99

Show the sample size, sample mean, sample standard deviation, and a histogram of the sample diameters (10 bins, range 9.95 to 10.1).

```python
sample_2 = Table().with_column('score', make_array(10.03, 9.75, 10.01, 9.75, 9.99, 10.01,
                                                   10.01, 10.08, 10.08, 9.95, 9.96, 9.99))

sample_size_2 = ...
sample_mean_2 = ...
sample_std_2  = ...

sample_size_2
sample_mean_2
sample_std_2
sample_2.hist(bins=10, range=make_array(9.95,10.1))
```

```python
grader.check("q1_2")
```

### Analysis

Hypothesize that the difference between population means (pop_mean_1 - pop_mean_2) is $\leq$ 0.
The alternative to this hypothesis is that the difference between population means is $\gt$ 0.

```python
pop_mean_diff_hypo = 0
pop_mean_diff_hypo
```

Calculate and show the two-sample t @ hypothesized population means difference 0.

$$
\textit{two-sample t} = \cfrac{(\textit{sample mean}_1 - \textit{sample mean}_2) - \textit{pop mean diff hypo}}
                              {\sqrt{\cfrac{{\textit{sample std}_1}^2}{\textit{sample size}_1} +
                                     \cfrac{{\textit{sample std}_2}^2}{\textit{sample size}_2}}}
$$

```python
two_sample_t = ( (sample_mean_1 - sample_mean_2) - pop_mean_diff_hypo) / \
                 ...

two_sample_t
```

```python
grader.check("q2_1")
```

Calculate and show the degrees of freedom for two samples.  This number will be used to choose a specific standard t distribution.

$$
\textit{df} =
\cfrac{\left(\cfrac{{\textit{sample std}_1}^2}{\textit{sample size}_1} + \cfrac{{\textit{sample std}_2}^2}{\textit{sample size}_2}\right)^2}
{\left(\left(\cfrac{1}{\textit{sample size}_1-1}\right) \times \left(\cfrac{{\textit{sample std}_1}^2}{\textit{sample size}_1}\right) ^2 \right)    +
\left(\left(\cfrac{1}{\textit{sample size}_2-1}\right) \times \left(\cfrac{{\textit{sample std}_2}^2}{\textit{sample size}_2}\right) ^2 \right)}$$

Note, if the sample standard deviations are the same, then you can use a simpler formula for degrees of freedom.

$$
\textit{df} = (\textit{sample size}_1 - 1) + (\textit{sample size}_2 - 1)
$$

```python
df = ((sample_std_1**2 / sample_size_1) + (sample_std_2**2 / sample_size_2))**2 / \
     ((1/(sample_size_1-1))*(sample_std_1**2/sample_size_1)**2 + (1/(sample_size_2-1))*(sample_std_2**2/sample_size_2)**2)
df
```

```python
grader.check("q2_2")
```

Get 1,000,000 values from the standard t distribution for the appropriate degrees of freedom.

Show a few of the values and a histogram of all the values (50 bins, range -4 to 4).

```python
dist_array = ...
dist = Table().with_column('t', dist_array)

dist
dist.hist(bins=50, range=make_array(-4,4))
```

```python
grader.check("q2_3")
```

Calculate and show the probability that the two-sample t $\ge$ 1.2398 when the difference
in population means is $\le$ 0 (this is the p-value).
Also show the two-sample t and histogram of standard t distribution with the area
corresponding to the probability highlighted.

```python
p_value = ...
values_for_hist = ...

two_sample_t
p_value
dist.hist(bins=50, range=make_array(-4,4), left_end=values_for_hist, right_end=4)
```

```python
grader.check("q2_4")
```

Calculate and show the critical value at significance level 0.05.
Note that the suspected difference between population means is ABOVE the hypothesized difference between
population means (0).
Also show the significance level and histogram of standard t distribution with the area corresponding
to the significance level highlighted.

```python
sig_level = 0.05
cv = ...

sig_level
cv
dist.hist(bins=50, range=make_array(-4,4), left_end=cv, right_end=4)
```

```python
grader.check("q2_5")
```

In a string, fill in the blank to show why you should assume the hypothesis is correct (difference in population means is $\le$ 0), at significance level 0.05. You should only use comparison symbols (>, <, =) in the blank. Define reason_1 and reason_2 as the full string expressing the statement. For example, reason_1 = "p_value = sig_level". Don't include any extra spaces.

```python
# Reason 1: "p_value ___ sig_level"
reason_1 = ...

# Reason 2: "two_sample_t ___ cv"
reason_2 = ...

reason_1, reason_2
```

```python
grader.check("q2_6")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Assume a significance level of 0.01, rather than 0.05. Calculate the critical value.

```python
sig_level_01 = 0.01
cv_c1 = ...
cv_c1
```

```python
grader.check("qc_1")
```

**Question 2**. Assume a significance level of 0.01, rather than 0.05. Calculate whether the manager should conclude that machine #1 is drilling diameters smaller than machine #2 is drilling.
(Answer is True or False.)

```python
sig_level_01 = 0.01
cv_c2 = ...
conclusion = ...
conclusion
```

```python
grader.check("qc_2")
```

**Question 3**. Assume a significance level of 0.10, rather than 0.05. Calculate whether the manager should conclude that machine #1 is drilling diameters smaller than machine #2 is drilling.
(Answer is True or False.)

```python
sig_level_1 = 0.10
cv_c3 = ...
conclusion_2 = ...
conclusion_2
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised October 5, 2023
</span>
</p>

Tests for q1_1 (points: None)
```
>>> sample_size_1 == 12
True
```
```
>>> round(sample_mean_1, 0) == 10
True
```
```
>>> round(sample_std_1, 2) == 0.03
True
```


Tests for q1_2 (points: None)
```
>>> sample_size_2 == 12
True
```
```
>>> round(sample_mean_2, 0) == 10
True
```
```
>>> round(sample_std_2, 2) == 0.11
True
```


Tests for q2_1 (points: None)
```
>>> round(two_sample_t, 2) == 1.24
True
```


Tests for q2_2 (points: None)
```
>>> round(df, 0) == 13
True
```


Tests for q2_3 (points: None)
```
>>> isinstance(dist, Table) and dist.num_columns == 1 and (dist.num_rows == 1000000) and all(isfinite(dist.column(0)))
True
```


Tests for q2_4 (points: None)
```
>>> round(p_value, 2) == 0.12
True
```
```
>>> values_for_hist == two_sample_t
True
```


Tests for q2_5 (points: None)
```
>>> round(cv, 0) == 2
True
```


Tests for q2_6 (points: None)
```
>>> (reason_1 == 'p_value > sig_level') | (reason_1 == 'p_value>sig_level')
True
```
```
>>> (reason_2 == 'two_sample_t < cv') | (reason_2 == 'two_sample_t<cv')
True
```


Tests for qc_1 (points: None)
```
>>> cv_c1 == percentile((1 - sig_level_01) * 100, dist.column('t'))
True
```


Tests for qc_2 (points: None)
```
>>> conclusion == True
True
```


Tests for qc_3 (points: None)
```
>>> conclusion_2 == True
True
```
