# Lab E: Simulation

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

## Setup

```python
import warnings
from datascience import Table, make_array, percentile, are
from numpy import mean, minimum, ndarray, allclose
from numpy.random import choice

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Hot Dog Vendor - without investment decision

### Business Situation

A vendor sells hot dogs at sporting events.  For each event, the vendor orders 5,000 hot dogs and pays for them in advance.  Each hot dog costs the vendor \\$3 and, if there is demand for it, can be sold for \\$5.  To assist in estimating demand, the vendor can consult data about demand at the prior 1,000 events.

### Data

Retrieve historical demand data, show the number of observations and the first few observations.

```python
data = ...
...
data
```

```python
grader.check("q1")
```

### Analysis

Calculate and show the mean, minimum, and maximum historical demand.  Visualize historical demand as a 10-bin histogram.

```python
demand_history = ...

demand_mean = ...
demand_min = ...
demand_max = ...

demand_hist = Table().with_column("demand", demand_history).hist()

print(demand_mean, demand_min, demand_max)
demand_hist
```

```python
grader.check("q2_1")
```

Set the unit price to \\$5, unit cost to \\$3, and order size to 5,000.
Calculate and show the investment required to make an order.

```python
unit_price = ...
unit_cost  = ...
order = ...
invest = ...

invest
```

```python
grader.check("q2_2")
```

Assume demand at an event is 5,000 units.  Calculate and show what the profit would be.

```python
demand_5000 = 5000 # units
sold_5000 = ...
profit_5000 = ...

profit_5000
```

```python
grader.check("q2_3")
```

Simulate 5 events.  Assume each event has demand equal to demand at some randomly chosen prior event.  
Visualize the distribution of profit across these 5 events as a 10-bin histogram, range -20,000 to 20,000.

```python
demand_5_events = ...
sold_5_events = minimum(order, demand_5_events) # units
profit_5_events = ...

profit_table = Table().with_column('profit', profit_5_events)
profit_hist = ...
profit_hist
```

```python
grader.check("q2_4")
```

Simulate 1,000,000 events.  Assume each event has demand equal to demand at some randomly chosen prior event.  
Keep track of the would-be profit at each event, stored in an array.  
Visualize the distribution of profits across these 1,000,000 events as a 10-bin histogram, range -20,000 to 20,000.

```python
demand = ...
sold = minimum(order, demand) # units
profit = ...

profit_table = Table().with_column('profit', profit)
profit_hist = ...
profit_hist
```

```python
grader.check("q2_5")
```

Calculate and show the expected profit (i.e., mean profit).

```python
mean_profit = ...
mean_profit
```

```python
grader.check("q2_6")
```

Calculate and show the best case profit.

```python
best_case_profit = ...
best_case_profit
```

```python
grader.check("q2_7")
```

Calculate and show the probability of a loss.

```python
probability_of_loss = ...
probability_of_loss
```

```python
grader.check("q2_8")
```

Calculate and show the worst case profit.

```python
worst_case_profit = ...
worst_case_profit
```

```python
grader.check("q2_9")
```

Construct and show a table with columns for
order, invest, expected profit, best case, probability of loss, probability of best case, worst case.

```python
result_5000 = Table().with_columns('order', order,
                                   'invest', invest,
                                   'expected profit', mean(profit),
                                   'best case', max(profit),
                                   'prob loss', profit_table.where('profit', are.below(0)).num_rows/profit_table.num_rows,
                                   'prob best case', profit_table.where('profit', max(profit)).num_rows/profit_table.num_rows,
                                   ...

result_5000
```

```python
grader.check("q2_10")
```

## Hot Dog Vendor - with investment decision

### Business Decision

A vendor sells hot dogs at sporting events.  For each event, the vendor must decide how many hot dogs to order and pay for them in advance.  Each hotdog costs the vendor \\$3 and, if there is demand for it, can be sold for \\$5.  To assist in estimating demand, the vendor can consult data about demand at the prior 1,000 events.

### Data

Use the historical data retrieved in Section 2. 

### Analysis

Set and show the unit price to \\$5 and unit cost to \\$3.

```python
unit_price = ...
unit_cost  = ...
unit_price
unit_cost
```

```python
grader.check("q3_1")
```

Set the order size to **4,000** and calculate the investment required.

Simulate 1,000,000 events.  Assume each event has demand equal to demand at some randomly chosen prior event.
Keep track of the would-be profit at each event, stored in an array.

Construct and show a table with columns for
order, invest, expected profit, best case, probability of loss, probability of best case, worst case.

Visualize the distribution of profits across these 1,000,000 events as a 10-bin histogram, range -20,000 to 20,000.

```python
order_invest_4000 = ...
invest_with_4000 = ...

demand_with_4000 = ...
sold_with_4000 = minimum(order_invest_4000, demand_with_4000) # units
profit_with_4000 = ...

profit_table_4000 = Table().with_column('profit', profit_with_4000)

result_4000 = Table().with_columns('order', order_invest_4000,
                                   'invest', invest_with_4000,
                                   'expected profit', mean(profit_with_4000),
                                   'best case', max(profit_with_4000),
                                   'prob loss', profit_table_4000.where('profit', are.below(0)).num_rows/profit_table_4000.num_rows,
                                   'prob best case', profit_table_4000.where('profit', max(profit_with_4000)).num_rows/profit_table_4000.num_rows,
                                   ...
result_4000

profit_hist_4000 = ...
profit_hist_4000
```

```python
grader.check("q3_2")
```

Set the order size to **5,000** and calculate the investment required.

Simulate 1,000,000 events.  Assume each event has demand equal to demand at some randomly chosen prior event.
Keep track of the would-be profit at each event, stored in an array.

Construct and show a table with columns for
order, invest, expected profit, best case, probability of loss, probability of best case, worst case.

Visualize the distribution of profits across these 1,000,000 events as a 10-bin histogram, range -20,000 to 20,000.

```python
order_invest_5000 = ...
invest_with_5000 = ...

demand_with_5000 = ...
sold_with_5000 = minimum(order_invest_5000, demand_with_5000) # units
profit_with_5000 = ...

profit_table_5000 = Table().with_column('profit', profit_with_5000)

result_5000 = Table().with_columns('order', order_invest_5000,
                                   'invest', invest_with_5000,
                                   'expected profit', mean(profit_with_5000),
                                   'best case', max(profit_with_5000),
                                   'prob loss', profit_table_5000.where('profit', are.below(0)).num_rows/profit_table_5000.num_rows,
                                   'prob best case', profit_table_5000.where('profit', max(profit_with_5000)).num_rows/profit_table_5000.num_rows,
                                   ...
result_5000

profit_hist_5000 = ...
profit_hist_5000
```

```python
grader.check("q3_3")
```

Set the order size to **6,000** and calculate the investment required.

Simulate 1,000,000 events.  Assume each event has demand equal to demand at some randomly chosen prior event.
Keep track of the would-be profit at each event, stored in an array.

Construct and show a table with columns for
order, invest, expected profit, best case, probability of loss, probability of best case, worst case.

Visualize the distribution of profits across these 1,000,000 events as a 10-bin histogram, range -20,000 to 20,000.

```python
order_invest_6000 = ...
invest_with_6000 = ...

demand_with_6000 = ...
sold_with_6000 = minimum(order_invest_6000, demand_with_6000) # units
profit_with_6000 = ...

profit_table_6000 = Table().with_column('profit', profit_with_6000)

result_6000 = Table().with_columns('order', order_invest_6000,
                                   'invest', invest_with_6000,
                                   'expected profit', mean(profit_with_6000),
                                   'best case', max(profit_with_6000),
                                   'prob loss', profit_table_6000.where('profit', are.below(0)).num_rows/profit_table_6000.num_rows,
                                   'prob best case', profit_table_6000.where('profit', max(profit_with_6000)).num_rows/profit_table_6000.num_rows,
                                   ...
result_6000

profit_hist_6000 = ...
profit_hist_6000
```

```python
grader.check("q3_4")
```

Set the order size to **7,000** and calculate the investment required.

Simulate 1,000,000 events.  Assume each event has demand equal to demand at some randomly chosen prior event.
Keep track of the would-be profit at each event, stored in an array.

Construct and show a table with columns for
order, invest, expected profit, best case, probability of loss, probability of best case, worst case.

Visualize the distribution of profits across these 1,000,000 events as a 10-bin histogram, range -20,000 to 20,000.

```python
order_invest_7000 = ...
invest_with_7000 = ...

demand_with_7000 = ...
sold_with_7000 = minimum(order_invest_7000, demand_with_7000) # units
profit_with_7000 = ...

profit_table_7000 = Table().with_column('profit', profit_with_7000)

result_7000 = Table().with_columns('order', order_invest_7000,
                                   'invest', invest_with_7000,
                                   'expected profit', mean(profit_with_7000),
                                   'best case', max(profit_with_7000),
                                   'prob loss', profit_table_7000.where('profit', are.below(0)).num_rows/profit_table_7000.num_rows,
                                   'prob best case', profit_table_7000.where('profit', max(profit_with_7000)).num_rows/profit_table_7000.num_rows,
                                   ...
result_7000

profit_hist_7000 = ...
profit_hist_7000
```

```python
grader.check("q3_5")
```

Combine the results into a single table and show the table.

```python
result = result_4000
result = ...
result = ...
result = ...

result
```

```python
grader.check("q3_6")
```

Visualize the results as a lineplot with expected profit, best case, and worst case vs. order decision.

```python
...
```

```python
grader.check("q3_7")
```

Visualize the results as a lineplot with probability of loss vs. order decision.

```python
...
```

```python
grader.check("q3_8")
```

Calculate and show the order decision that maximizes expected profit.

```python
best_expected_profit = ...
order_best_expected_profit = ...
order_best_expected_profit
```

```python
grader.check("q3_9")
```

Calculate and show the order decision that minimizes probability of a loss.

```python
best_prob_loss = ...
order_best_prob_loss = ...
order_best_prob_loss
```

```python
grader.check("q3_10")
```

Show the order decision that maximizes best case profit.

```python
best_best_case = ...
order_best_best_case = ...
order_best_best_case
```

```python
grader.check("q3_11")
```

Show the order decision that maximizes worst case profit.

```python
best_worst_case = ...
order_best_worst_case = ...
order_best_worst_case
```

```python
grader.check("q3_12")
```

Calculate and show the best case profit when decision is to order 4,000.  
Calculate and show the probability of the best case profit when decision is to order 4,000.  
Calculate and show the best case profit when decision is to order 7,000.  
Calculate and show the probability of the best case profit when decision is to order 7,000.

```python
best_4000 = ...
prob_best_4000 = ...

best_7000 = ...
prob_best_7000 = ...

best_4000, prob_best_4000, best_7000, prob_best_7000
```

```python
grader.check("q3_13")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Simulate 1,000,000 events for an order size of 7,000 and calculate the probability that profit exceeds $9,000. (Answer will be approximate.)

```python
order_7000 = 7000
invest_7000 = ...
demand_7000 = ...
sold_7000 = ...
profit_7000 = ...
prob_profit_above_9000 = ...
prob_profit_above_9000
```

```python
grader.check("qc_1")
```

**Question 2**. Simulate 1,000,000 events for an order size of 7,000, and calculate the lower bound of the highest 5% of possible profits. (This is known as the upside.)
(Answer will be approximate.)

```python
order_c2 = 7000
invest_c2 = order_c2 * unit_cost
demand_c2 = ...
sold_c2 = ...
profit_c2 = ...
upside_c2 = ...
upside_c2
```

```python
grader.check("qc_2")
```

3. Using the result table, determine the order quantity that has the highest worst-case profit.

```python
order_best_worst_case_5000 = ...
order_best_worst_case_5000
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised September 19, 2024
</span>
</p>

Tests for q1 (points: None)
```
>>> data.num_rows == 1000
True
```
```
>>> data.labels == ('Demand',)
True
```


Tests for q2_1 (points: None)
```
>>> demand_mean == 5019.176
True
```
```
>>> demand_min == 1662
True
```
```
>>> demand_max = 7984
```
```
>>> all(demand_history == data.column('Demand'))
True
```


Tests for q2_10 (points: None)
```
>>> row = result_5000.row(0)
>>> expected = [order, invest, mean(profit), max(profit), profit_table.where('profit', are.below(0)).num_rows / profit_table.num_rows, profit_table.where('profit', max(profit)).num_rows / profit_table.num_rows, min(profit)]
>>> row.item('order') == expected[0] and row.item('invest') == expected[1] and (row.item('expected profit') == expected[2]) and (row.item('best case') == expected[3]) and (row.item('prob loss') == expected[4]) and (row.item('prob best case') == expected[5]) and (row.item('worst case') == expected[6])
True
```


Tests for q2_2 (points: None)
```
>>> invest == 15000
True
```


Tests for q2_3 (points: None)
```
>>> profit_5000 == 10000
True
```


Tests for q2_4 (points: None)
```
>>> isinstance(demand_5_events, (list, ndarray)) and len(demand_5_events) == 5
True
```
```
>>> all(profit_5_events == sold_5_events * unit_price - order * unit_cost)
True
```


Tests for q2_5 (points: None)
```
>>> isinstance(demand, (list, ndarray)) and len(demand) == 1000000
True
```
```
>>> all(profit == sold * unit_price - order * unit_cost)
True
```


Tests for q2_6 (points: None)
```
>>> mean_profit == mean(profit)
True
```


Tests for q2_7 (points: None)
```
>>> best_case_profit == max(profit)
True
```


Tests for q2_8 (points: None)
```
>>> probability_of_loss == profit_table.where('profit', are.below(0)).num_rows / profit_table.num_rows
True
```


Tests for q2_9 (points: None)
```
>>> worst_case_profit == min(profit)
True
```


Tests for q3_1 (points: None)
```
>>> unit_price == 5
True
```
```
>>> unit_cost == 3
True
```


Tests for q3_10 (points: None)
```
>>> probs_of_loss = [(result_4000.column(0)[0], result_4000.column(4)[0]), (result_5000.column(0)[0], result_5000.column(4)[0]), (result_6000.column(0)[0], result_6000.column(4)[0]), (result_7000.column(0)[0], result_7000.column(4)[0])]
>>> min_loss = probs_of_loss[0][1]
>>> correct_order = probs_of_loss[0][0]
>>> for order_with_invest_dec, prob in probs_of_loss[1:]:
...     if prob < min_loss:
...         min_loss = prob
...         correct_order = order_with_invest_dec
>>> order_best_prob_loss == correct_order
True
```


Tests for q3_11 (points: None)
```
>>> best_cases = [(result_4000.column(0)[0], result_4000.column(3)[0]), (result_5000.column(0)[0], result_5000.column(3)[0]), (result_6000.column(0)[0], result_6000.column(3)[0]), (result_7000.column(0)[0], result_7000.column(3)[0])]
>>> max_best_case = best_cases[0][1]
>>> correct_order = best_cases[0][0]
>>> for order_with_invest_dec, best_case in best_cases[1:]:
...     if best_case > max_best_case:
...         max_best_case = best_case
...         correct_order = order_with_invest_dec
>>> order_best_best_case == correct_order
True
```


Tests for q3_12 (points: None)
```
>>> worst_cases = [(result_4000.column(0)[0], result_4000.column(6)[0]), (result_5000.column(0)[0], result_5000.column(6)[0]), (result_6000.column(0)[0], result_6000.column(6)[0]), (result_7000.column(0)[0], result_7000.column(6)[0])]
>>> max_worst_case = worst_cases[0][1]
>>> correct_order = worst_cases[0][0]
>>> for order_with_invest_dec, wc in worst_cases[1:]:
...     if wc > max_worst_case:
...         max_worst_case = wc
...         correct_order = order_with_invest_dec
>>> order_best_worst_case == correct_order
True
```


Tests for q3_13 (points: None)
```
>>> row_4000 = result.where(0, 4000).row(0)
>>> row_7000 = result.where(0, 7000).row(0)
>>> correct_best_4000 = row_4000[3]
>>> best_4000 == correct_best_4000
True
```
```
>>> row_4000 = result.where(0, 4000).row(0)
>>> row_7000 = result.where(0, 7000).row(0)
>>> correct_prob_best_4000 = row_4000[5]
>>> prob_best_4000 == correct_prob_best_4000
True
```
```
>>> row_4000 = result.where(0, 4000).row(0)
>>> row_7000 = result.where(0, 7000).row(0)
>>> correct_best_7000 = row_7000[3]
>>> best_7000 == correct_best_7000
True
```
```
>>> row_4000 = result.where(0, 4000).row(0)
>>> row_7000 = result.where(0, 7000).row(0)
>>> correct_prob_best_7000 = row_7000[5]
>>> prob_best_7000 == correct_prob_best_7000
True
```


Tests for q3_2 (points: None)
```
>>> order_invest_4000 == 4000 and invest_with_4000 == order_invest_4000 * unit_cost
True
```
```
>>> correct_length = len(demand_with_4000) == 1000000
>>> valid_values = all((d in demand_history for d in demand_with_4000[:100]))
>>> sold_correct = all(sold_with_4000 == minimum(order_invest_4000, demand_with_4000))
>>> profit_correct = allclose(profit_with_4000, sold_with_4000 * unit_price - order_invest_4000 * unit_cost)
>>> correct_length and valid_values and sold_correct and profit_correct
True
```
```
>>> expected = [order_invest_4000, invest_with_4000, mean(profit_with_4000), max(profit_with_4000), profit_table_4000.where('profit', are.below(0)).num_rows / profit_table_4000.num_rows, profit_table_4000.where('profit', max(profit_with_4000)).num_rows / profit_table_4000.num_rows, min(profit_with_4000)]
>>> row = result_4000.row(0)
>>> row.item('order') == expected[0] and row.item('invest') == expected[1] and (row.item('expected profit') == expected[2]) and (row.item('best case') == expected[3]) and (row.item('prob loss') == expected[4]) and (row.item('prob best case') == expected[5]) and (row.item('worst case') == expected[6])
True
```


Tests for q3_3 (points: None)
```
>>> order_invest_5000 == 5000 and invest_with_5000 == order_invest_5000 * unit_cost
True
```
```
>>> correct_length = len(demand_with_5000) == 1000000
>>> valid_values = all((d in demand_history for d in demand_with_5000[:100]))
>>> sold_correct = all(sold_with_5000 == minimum(order_invest_5000, demand_with_5000))
>>> profit_correct = allclose(profit_with_5000, sold_with_5000 * unit_price - order_invest_5000 * unit_cost)
>>> correct_length and valid_values and sold_correct and profit_correct
True
```
```
>>> expected_values = [order_invest_5000, invest_with_5000, mean(profit_with_5000), max(profit_with_5000), profit_table_5000.where('profit', are.below(0)).num_rows / profit_table_5000.num_rows, profit_table_5000.where('profit', max(profit_with_5000)).num_rows / profit_table_5000.num_rows, min(profit_with_5000)]
>>> actual_row = result_5000.row(0)
>>> actual_row.item('order') == expected_values[0] and actual_row.item('invest') == expected_values[1] and (actual_row.item('expected profit') == expected_values[2]) and (actual_row.item('best case') == expected_values[3]) and (actual_row.item('prob loss') == expected_values[4]) and (actual_row.item('prob best case') == expected_values[5]) and (actual_row.item('worst case') == expected_values[6])
True
```


Tests for q3_4 (points: None)
```
>>> order_invest_6000 == 6000 and invest_with_6000 == order_invest_6000 * unit_cost
True
```
```
>>> correct_length = len(demand_with_6000) == 1000000
>>> valid_values = all((d in demand_history for d in demand_with_6000[:100]))
>>> sold_correct = all(sold_with_6000 == minimum(order_invest_6000, demand_with_6000))
>>> profit_correct = allclose(profit_with_6000, sold_with_6000 * unit_price - order_invest_6000 * unit_cost)
>>> correct_length and valid_values and sold_correct and profit_correct
True
```
```
>>> expected_values = [order_invest_6000, invest_with_6000, mean(profit_with_6000), max(profit_with_6000), profit_table_6000.where('profit', are.below(0)).num_rows / profit_table_6000.num_rows, profit_table_6000.where('profit', max(profit_with_6000)).num_rows / profit_table_6000.num_rows, min(profit_with_6000)]
>>> actual_row = result_6000.row(0)
>>> actual_row.item('order') == expected_values[0] and actual_row.item('invest') == expected_values[1] and (actual_row.item('expected profit') == expected_values[2]) and (actual_row.item('best case') == expected_values[3]) and (actual_row.item('prob loss') == expected_values[4]) and (actual_row.item('prob best case') == expected_values[5]) and (actual_row.item('worst case') == expected_values[6])
True
```


Tests for q3_5 (points: None)
```
>>> order_invest_7000 == 7000 and invest_with_7000 == order_invest_7000 * unit_cost
True
```
```
>>> correct_length = len(demand_with_7000) == 1000000
>>> valid_values = all((d in demand_history for d in demand_with_7000[:100]))
>>> sold_correct = all(sold_with_7000 == minimum(order_invest_7000, demand_with_7000))
>>> profit_correct = allclose(profit_with_7000, sold_with_7000 * unit_price - order_invest_7000 * unit_cost)
>>> correct_length and valid_values and sold_correct and profit_correct
True
```
```
>>> expected_values = [order_invest_7000, invest_with_7000, mean(profit_with_7000), max(profit_with_7000), profit_table_7000.where('profit', are.below(0)).num_rows / profit_table_7000.num_rows, profit_table_7000.where('profit', max(profit_with_7000)).num_rows / profit_table_7000.num_rows, min(profit_with_7000)]
>>> actual_row = result_7000.row(0)
>>> actual_row.item('order') == expected_values[0] and actual_row.item('invest') == expected_values[1] and (actual_row.item('expected profit') == expected_values[2]) and (actual_row.item('best case') == expected_values[3]) and (actual_row.item('prob loss') == expected_values[4]) and (actual_row.item('prob best case') == expected_values[5]) and (actual_row.item('worst case') == expected_values[6])
True
```


Tests for q3_6 (points: None)
```
>>> correct_row_count = result.num_rows == 4
>>> correct_orders = all((row.item('order') in [4000, 5000, 6000, 7000] for row in result.rows))
>>> correct_row_count and correct_orders
True
```


Tests for q3_7 (points: None)
```
>>> t = result.select('order', 'expected profit', 'best case', 'worst case')
>>> t.num_rows == 4 and t.labels == ('order', 'expected profit', 'best case', 'worst case')
True
```


Tests for q3_8 (points: None)
```
>>> t = result.select('order', 'prob loss')
>>> t.num_rows == 4 and t.labels == ('order', 'prob loss')
True
```


Tests for q3_9 (points: None)
```
>>> expected_profits = [(result_4000.column(0)[0], result_4000.column(2)[0]), (result_5000.column(0)[0], result_5000.column(2)[0]), (result_6000.column(0)[0], result_6000.column(2)[0]), (result_7000.column(0)[0], result_7000.column(2)[0])]
>>> max_profit = expected_profits[0][1]
>>> correct_order = expected_profits[0][0]
>>> for order_with_invest_dec, profit_with_invest_dec in expected_profits[1:]:
...     if profit_with_invest_dec > max_profit:
...         max_profit = profit_with_invest_dec
...         correct_order = order_with_invest_dec
>>> order_best_expected_profit == correct_order
True
```


Tests for qc_1 (points: None)
```
>>> invest_7000 == order_7000 * unit_cost
True
```
```
>>> all(sold_7000 == minimum(order_7000, demand_7000))
True
```
```
>>> all(profit_7000 == sold_7000 * unit_price - invest_7000)
True
```
```
>>> prob_profit_above_9000 == (profit_7000 > 9000).mean()
True
```


Tests for qc_2 (points: None)
```
>>> all(sold_c2 == minimum(order_c2, demand_c2))
True
```
```
>>> all(profit_c2 == sold_c2 * unit_price - invest_c2)
True
```
```
>>> all(profit_c2 == sold_c2 * unit_price - invest_c2)
True
```
```
>>> upside_c2 == percentile(95, profit_c2)
True
```


Tests for qc_3 (points: None)
```
>>> order_best_worst_case_5000 == 4000
True
```
