# Lab D: Data Visualization

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

## Setup

```python
import warnings
from datascience import Table, make_array, are
from numpy import mean, std, arange, cumsum, float64, int64, ndarray

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Fundraising for a Political Campaign 

### Business Decision

The Director of Fundraising in California for a US presidential candidate is considering where best to focus the campaign's effort during the early period January to September of the year prior to the election.  To inform her decision, she would like insight into how fundraising has gone for other candidates in the past. 

### Data

The donor dataset here is from the US Federal Election Commission (FEC) and itemizes individual contributions from California to federal campaign committees for US presidential candidates received January 2019 to September 2019 (refer: US Federal Election Commission https://www.fec.gov/data/browse-data/?tab=bulk-data).

Retrieve the donor dataset, filter to include only transactions where amount is between \\$0 and \\$2,200, show the number of transactions and the first few transactions.

```python
data = Table.read_table('donations.csv', low_memory=False)
data = data.where('AMOUNT', are.between_or_equal_to(0, 2200))
data.num_rows
data
```

### 1. Statistics by Party

**Question 1.1.** Calculate the number of donations.  Calculate the mean \\$ donated.  Calculate the standard deviation of \\$ donated.  Calculate the total \\$ donated.

```python
num_donations = ...
mean_donations = ...
sd_donations = ...
sum_donations = ...
print("Number of donations:", num_donations)
print("Mean donation amount:", mean_donations)
print("Standard deviation of donation amounts:", sd_donations)
print("Total amount donated:", sum_donations)
```

```python
grader.check("q11")
```

**Question 1.2.** Calculate the number of donations for each party.

```python
agg_party_count = ...
agg_party_count
```

```python
grader.check("q12")
```

**Question 1.3.** Calculate the mean $ donated for each party. 

```python
agg_party_mean = ...
agg_party_mean
```

```python
grader.check("q13")
```

**Question 1.4.** Calculate the $ donated for each party.

```python
agg_party_sum = ...
agg_party_sum
```

```python
grader.check("q14")
```

**Question 1.5.** Visualize the number of donations for each party as a bar chart.

```python
agg_party_mean = ...
agg_party_count_bh
```

```python
grader.check("q15")
```

<!-- BEGIN QUESTION -->

**Question 1.6.** Visualize the mean $ donated for each party as a bar chart.

```python
agg_party_mean_barh = ...
agg_party_mean_barh
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.7.** Visualize the $ donated for each party as a bar chart.

```python
agg_party_sum_barh = ...
agg_party_sum_barh
```

<!-- END QUESTION -->

### 2. Statistics by Candidate

**Question 2.1.** Visualize the number of donations for the top-10 high-count candidates as a bar chart, sorted highest to lowest.

```python
top_10 = ...
top_10.barh('...')
```

```python
grader.check("q21")
```

**Question 2.2.** Visualize the mean \\$ donated for the top-10 high-mean candidates as a bar chart, sorted highest to lowest.

```python
top_10_amount_mean = ...
top_10_amount_mean.barh('...')
```

```python
grader.check("q22")
```

**Question 2.3.** Visualize the \\$ donated for the top-10 high-\\$ candidates as a bar chart, sorted highest to lowest.

```python
top_10_amount_sum = ...
top_10_amount_sum.barh('...')
```

```python
grader.check("q23")
```

### 3. Statistics by Occupation

**Question 3.1.** Calculate mean \\$ donated to Democrats for each occupation that donated to Democrats.

```python
agg_dem = ...
agg_dem
```

```python
grader.check("q31")
```

**Question 3.2.** Calculate mean \\$ donated to Republicans for each occupation that donated to Republicans.

```python
agg_rep = ...
agg_rep
```

```python
grader.check("q32")
```

**Question 3.3.** Construct a table to compare mean \\$ donated to Democrats vs. mean \\$ donated to Republicans for each occupation (from among occupations that donated to both Democrats and Republicans).

```python
agg = ...
agg
```

```python
grader.check("q33")
```

**Question 3.4.** Find the 10 occupations with the highest mean \\$ donated to Democrats (from among occupations that donated to both Democrats and Republicans).  These are the top-10 Democrat donor occupations.  Use a bar chart to visualize mean \\$ donated to Democrats and mean \\$ donated to Republicans for each of the top-10 Democrat donor occupations, sorted by mean \\$ donated to Democrats.

```python
agg_sort_dem = ...
agg_sort_dem.barh('...')
```

```python
grader.check("q34")
```

**Question 3.5.** Find the 10 occupations with the highest mean \\$ donated to Republicans (from among occupations that donated to both Democrats and Republicans).  These are the top-10 Republican donor occupations.  Use a bar chart to visualize mean \\$ donated to Democrats and mean \\$ donated to Republicans for each of the top-10 Republican donor occupations, sorted by mean \\$ donated to Republicans.

```python
agg_sort_rep = ...
agg_sort_rep.barh('...')
```

```python
grader.check("q35")
```

### 4. Distribution of Donations

For this analysis of distribution of donations, consider donations between \\$0 and \\$100.

**Question 4.1.** Construct a table to show counts of donations partitioned into 10 evenly spaced bins.  Note the bin variable indicates an amount range, the AMOUNT count variable indicates a number of donations.

```python
bin_table = ...
bin_table.show()
```

```python
grader.check("q41")
```

**Question 4.2.** Given the table above, visualize the counts of donations, partitioned into 10 evenly spaced bins, as a bar chart.

```python
bin_quantity = ...
bin_table.bar(bin_quantity)
```

```python
grader.check("q42")
```

**Question 4.3.** Visualize the counts of donations, partitioned into 10 evenly spaced bins, as a histogram.  Note that the bar chart and histogram visualizations match, except for the bar center positioning.  

```python
quantity = ...
num_bins = ...
bin_ranges = ...
data.hist(quantity, normed=False, bins=num_bins, range=bin_ranges)
```

```python
grader.check("q43")
```

**Question 4.4.** Visualize the counts of donation amounts between \\$0 and \\$100, partitioned into 1000 evenly spaced bins, as a histogram.  Note that too many bins results in a non-useful histogram.

```python
num_bins_1000 = ...
data.hist(quantity, normed=False, bins=num_bins_1000, range=bin_ranges)
```

```python
grader.check("q44")
```

**Question 4.5.** Visualize the counts of donation amounts between \\$0 and \\$100, partitioned into 3 evenly spaced bins, as a histogram.  Note that too few bins results in a non-useful histogram.

```python
num_bins_3 = ...
data.hist(quantity, normed=False, bins=num_bins_3, range=bin_ranges)
```

```python
grader.check("q45")
```

**Question 4.6.** Visualize the counts of donation amounts between \\$0 and \\$100 to Democrats, partitioned into 10 evenly spaced bins, as a histogram.

```python
dem_only = ..
num_bins_10 = ...
dem_only.hist(quantity, normed=False, bins=num_bins_10, range=bin_ranges)
```

```python
grader.check("q46")
```

**Question 4.7.** Visualize the counts of donation amounts between \\$0 and \\$100 to Republicans, partitioned into 10 evenly spaced bins, as a histogram.

```python
rep_only = ..
num_bins_10 = ...
rep_only.hist(quantity, normed=False, bins=num_bins_10, range=bin_ranges)
```

```python
grader.check("q47")
```

**Question 4.8.** Visualize the counts of donations, partitioned into 10 evenly spaced bins, as a histogram.  Scale the vertical axis so that area of the histogram equals one. 

```python
data.hist(quantity, normed=..., bins=num_bins_10, range=bin_ranges)
```

```python
grader.check("q48")
```

**Question 4.9.** Calculate the probability that a donation is between \\$0 and \\$10 given the donations between \\$0 and \\$100. Also calculate the area of the first bar of the previous histogram.

```python
num_less_than_10 = ...
num_between_0_100 = ...
conditional_prob = num_less_than_10 / num_between_0_100
print("Probability that a donation is between $0 and $10 given the donations between $0 and $100:", conditional_prob)
```

```python
grader.check("q49")
```

**Question 4.10.** Calculate the height of the first bar of the previous histogram. 

```python
width_bar = ...
height = ...
print("Height of the first bar in the histogram (in % per unit):", height)
```

```python
grader.check("q410")
```

### 5. Accumulation of Donations

**Question 5.1.** Construct a table to show the cumulative $ donated, by date.

```python
data_sorted_by_date = ...
data = data_sorted_by_date.with_column('CUMSUM', ...)
data
```

```python
grader.check("q51")
```

**Question 5.2.** Visualize the cumulative $ donated, by date, as a scatter plot.

```python
only_date_cumsum =...
only_date_cumsum.scatter(...)
```

```python
grader.check("q52")
```

**Question 5.3.** Visualize the cumulative $ donated, by date, as a line plot.

```python
only_date_cumsum...
```

```python
grader.check("q53")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Show which party raised the largest $ donated.

```python
party_totals = ...
highest_party = ...
print("The", highest_party, "party raised the most!")
```

```python
grader.check("qc_1")
```

**Question 2**. Calculate the total amount donated for the top 10 richest candidates.

```python
top10_total = ...
top10_total
```

```python
grader.check("qc_2")
```

**Question 3**. Calculate the total amount donated by the top 2 occupations.

```python
top2_total = ...
top2_total
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised September 10, 2023
</span>
</p>

Tests for q11 (points: None)
```
>>> type(num_donations) is int and num_donations > 0 and (num_donations == 264119)
True
```
```
>>> type(mean_donations) is float64 and mean_donations > 0 and (round(mean_donations, 2) == 101.47)
True
```
```
>>> type(sd_donations) is float64 and sd_donations >= 0 and (round(sd_donations, 2) == 221.02)
True
```
```
>>> type(sum_donations) is int64 and sum_donations > 0 and (sum_donations == 26801008)
True
```


Tests for q12 (points: None)
```
>>> bool(sum(agg_party_count[0] == ['DEM', 'IND', 'LIB', 'OTH', 'REP', 'UN', 'UNK']) / 7)
True
```
```
>>> agg_party_count.column('PARTY').item(0) == 'DEM' and agg_party_count.column('count').item(0) == 251145
True
```
```
>>> min(agg_party_count.column('count')) == 1
True
```


Tests for q13 (points: None)
```
>>> bool(sum(agg_party_mean[0] == ['DEM', 'IND', 'LIB', 'OTH', 'REP', 'UN', 'UNK']) / 7)
True
```
```
>>> agg_party_mean.column('PARTY').item(0) == 'DEM' and round(agg_party_mean.column('AMOUNT mean').item(0), 3) == 101.884
True
```
```
>>> type(agg_party_mean) is Table
True
```
```
>>> len(agg_party_mean.columns) == 2
True
```


Tests for q14 (points: None)
```
>>> bool(sum(agg_party_sum[0] == ['DEM', 'IND', 'LIB', 'OTH', 'REP', 'UN', 'UNK']) / 7)
True
```
```
>>> agg_party_sum.column('PARTY').item(0) == 'DEM' and agg_party_sum.column('AMOUNT sum').item(0) == 25587612
True
```
```
>>> len(agg_party_sum.columns) == 2
True
```
```
>>> len(agg_party_sum.rows) == 7
True
```


Tests for q15 (points: None)
```
>>> agg_party_count_bh is not True
True
```


Tests for q21 (points: None)
```
>>> len(top_10.rows) == 10 and len(top_10.columns) == 2
True
```
```
>>> top_10.column('CAND_NAME').item(0) == 'SANDERS, BERNARD' and top_10.column('count').item(0) == 93913
True
```
```
>>> top_10.column('CAND_NAME').item(-1) == 'WILLIAMSON, MARIANNE' and top_10.column('count').item(-1) == 4632
True
```


Tests for q22 (points: None)
```
>>> top_10_amount_mean
CAND_NAME               | AMOUNT mean
DE BLASIO, BILL         | 1068.18
SANFORD, MARSHALL HON   | 725
GLEIBERMAN, BEN         | 628.571
CHARLES, MARK R.        | 562.5
OJEDA, RICHARD NEECE II | 512.5
WALSH, JOE              | 464.286
MOULTON, SETH           | 395.568
GILLIBRAND, KIRSTEN     | 376.282
HOROWITZ, AMI           | 375.75
DELANEY, JOHN K.        | 373.577
```
```
>>> len(top_10_amount_mean.rows) == 10 and len(top_10_amount_mean.columns) == 2
True
```
```
>>> top_10_amount_mean.column('CAND_NAME').item(0) == 'DE BLASIO, BILL' and round(top_10_amount_mean.column('AMOUNT mean').item(0), 2) == 1068.18
True
```
```
>>> top_10_amount_mean.column('CAND_NAME').item(-1) == 'DELANEY, JOHN K.' and round(top_10_amount_mean.column('AMOUNT mean').item(-1), 2) == 373.58
True
```


Tests for q23 (points: None)
```
>>> len(top_10_amount_sum.rows) == 10 and len(top_10_amount_sum.columns) == 2
True
```
```
>>> top_10_amount_sum.column('CAND_NAME').item(0) == 'BUTTIGIEG, PETE' and top_10_amount_sum.column('AMOUNT sum').item(0) == 6240894
True
```
```
>>> top_10_amount_sum.column('CAND_NAME').item(-1) == 'WILLIAMSON, MARIANNE' and top_10_amount_sum.column('AMOUNT sum').item(-1) == 500255
True
```


Tests for q31 (points: None)
```
>>> len(agg_dem.rows) == 7838 and len(agg_dem.columns) == 2
True
```
```
>>> agg_dem.column('OCCUPATION').item(0) == '(NULL)' and round(agg_dem.column('DEM mean').item(0), 1) == 178.5
True
```


Tests for q32 (points: None)
```
>>> len(agg_rep.rows) == 1070 and len(agg_rep.columns) == 2
True
```
```
>>> agg_rep.column('OCCUPATION').item(0) == 'ACCOUNT EXECUTIVE' and agg_rep.column('REP mean').item(0) == 50
True
```


Tests for q33 (points: None)
```
>>> agg
OCCUPATION            | DEM mean | REP mean
ACCOUNT EXECUTIVE     | 86.6923  | 50
ACCOUNT MANAGER       | 83.0488  | 31
ACCOUNTANT            | 74.9879  | 49.7143
ACCOUNTING            | 49.551   | 60.2174
ACCOUNTING SPECIALIST | 20       | 20
ACTIVE DUTY           | 89.5714  | 250
ACTOR                 | 165.39   | 25
ADMIN                 | 46.0405  | 5
ADMINISTRATION        | 69.6757  | 82.6923
ADMINISTRATIVE        | 94.2174  | 35
... (508 rows omitted)
```
```
>>> len(agg_rep.rows)
1070
```
```
>>> len(agg.rows) == 518 and len(agg.columns) == 3
True
```
```
>>> agg.column('OCCUPATION').item(0) == 'ACCOUNT EXECUTIVE' and round(agg.column('DEM mean').item(0), 1) == 86.7 and (round(agg.column('REP mean').item(0), 1) == 50.0)
True
```


Tests for q34 (points: None)
```
>>> agg_sort_dem.column('REP mean').item(0)
1000.0
```
```
>>> len(agg_sort_dem.rows) == 10 and len(agg_sort_dem.columns) == 3
True
```
```
>>> agg_sort_dem.column('OCCUPATION').item(0) == 'SOLAR ENERGY' and agg_sort_dem.column('DEM mean').item(0) == 1250.0
True
```


Tests for q35 (points: None)
```
>>> agg_sort_rep.column('REP mean').item(0)
1000.0
```
```
>>> len(agg_sort_rep.rows) == 10 and len(agg_sort_rep.columns) == 3
True
```
```
>>> agg_sort_rep.column('OCCUPATION').item(0) == 'BUSINESS PARTNER' and agg_sort_rep.column('REP mean').item(0) == 1000.0
True
```


Tests for q41 (points: None)
```
>>> len(bin_table.columns)
2
```
```
>>> type(bin_table) is Table and bin_table.num_rows == 11 and (len(bin_table.columns) == 2)
True
```
```
>>> bin_table.column(0).item(0) == 0 and bin_table.column(1).item(0) == 41619
True
```


Tests for q410 (points: None)
```
>>> type(width_bar) is int and width_bar > 0 and (width_bar == 10)
True
```
```
>>> type(height) is float and height >= 0 and (round(height, 3) == 0.019)
True
```


Tests for q42 (points: None)
```
>>> type(bin_quantity) is str
True
```
```
>>> len(bin_quantity) == 3 and len(bin_quantity) > 0
True
```
```
>>> bin_quantity == 'bin'
True
```


Tests for q43 (points: None)
```
>>> type(quantity) is str and len(quantity) > 0 and (quantity == 'AMOUNT')
True
```
```
>>> type(num_bins) is int and num_bins > 0 and (num_bins ** 2 == 100)
True
```
```
>>> type(bin_ranges) is ndarray and len(bin_ranges) == 2 and (bin_ranges.item(0) == 0) and (bin_ranges.item(1) == 100)
True
```


Tests for q44 (points: None)
```
>>> type(num_bins_1000) is int and num_bins_1000 > 0 and (num_bins_1000 * 1000 == 1000000)
True
```


Tests for q45 (points: None)
```
>>> type(num_bins_3) is int and num_bins_3 > 0 and (round(num_bins_3 * 4 * 8 / 10, 2) == 9.6)
True
```


Tests for q46 (points: None)
```
>>> type(dem_only) is Table and dem_only.num_rows == 251145 and (len(dem_only.columns) == 7)
True
```
```
>>> dem_only.column('CAND_NAME').item(0) == 'DELANEY, JOHN K.' and dem_only.column('AMOUNT').item(0) == 25
True
```
```
>>> mean(dem_only.column('PARTY') == 'DEM') == 1
True
```


Tests for q47 (points: None)
```
>>> type(rep_only) is Table and rep_only.num_rows == 12946 and (len(rep_only.columns) == 7)
True
```
```
>>> rep_only.column('CAND_NAME').item(0) == 'SANFORD, MARSHALL HON' and rep_only.column('AMOUNT').item(0) == 400
True
```
```
>>> mean(rep_only.column('PARTY') == 'REP') == 1
True
```


Tests for q48 (points: None)
```
>>> is_normed is True
True
```


Tests for q49 (points: None)
```
>>> type(num_less_than_10) is int and num_less_than_10 >= 0 and (num_less_than_10 == 41619)
True
```
```
>>> type(num_between_0_100) is int and num_between_0_100 >= 0 and (num_between_0_100 == 223959)
True
```
```
>>> type(conditional_prob) is float and conditional_prob >= 0 and (round(conditional_prob, 3) == 0.186)
True
```


Tests for q51 (points: None)
```
>>> type(data) is Table and data.num_rows == 264119 and (len(data.columns) == 8)
True
```
```
>>> type(data_sorted_by_date) is Table and data_sorted_by_date.column(0).item(0) == 'TRUMP, DONALD J.' and (data_sorted_by_date.column(5).item(0) == '2019-01-01')
True
```
```
>>> data.column('CUMSUM').item(0) == 50 and data.column('OCCUPATION').item(2) == 'ENTREPRENEUR'
True
```


Tests for q52 (points: None)
```
>>> type(only_date_cumsum) is Table and only_date_cumsum.num_rows == 264119 and (len(only_date_cumsum.columns) == 2)
True
```


Tests for q53 (points: None)
```
>>> type(only_date_cumsum) is Table and only_date_cumsum.num_rows == 264119 and (len(only_date_cumsum.columns) == 2)
True
```


Tests for qc_1 (points: None)
```
>>> highest_party == 'DEM'
True
```


Tests for qc_2 (points: None)
```
>>> top10_total == sum(data.select('CAND_NAME', 'AMOUNT').group('CAND_NAME', sum).sort('AMOUNT sum', descending=True).take(arange(10)).column('AMOUNT sum'))
True
```


Tests for qc_3 (points: None)
```
>>> top2_total == sum(data.select('OCCUPATION', 'AMOUNT').group('OCCUPATION', sum).sort('AMOUNT sum', descending=True).take(0, 1).column('AMOUNT sum'))
True
```
