# Lab J3: Chi-Squared Test for Independence

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

## Setup

```python
import warnings
from datascience import Table, make_array, percentile, are
from numpy import mean, unique, isclose
from numpy.random import chisquare

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Commerce Website and Customer Demographics

### Business Decision

An online retailer of cameras and phones is considering how to customize its website to drive more revenue.  It suspects that visitors to the website with higher incomes have different preferences than visitors with lower incomes do.  If they do, then it will treat visitors to information primarily about the product category they prefer based on their incomes.  Otherwise, it will treat all visitors to information about both product categories.  Is preference for a product category independent of income level?

### Data

**Question 0.1** Retrieve a sample of purchases, show the number of observations and the first few observations

```python
data = Table().read_table('...')
print("There are", data.num_rows, "rows in the dataset.")
data
```

```python
grader.check("q01")
```

### Analysis

**Question 1.1** Show the count of each category-income pair.

```python
data_pair = data.groups(...)
data_pair
```

```python
grader.check("q11")
```

**Question 1.2** Calculate and show the count of each category.

```python
data_category = data.group(...)
data_category
```

```python
grader.check("q12")
```

**Question 1.3** Calculate and show the count of each income.

```python
data_category = data.group(...)
data_income
```

```python
grader.check("q13")
```

**Question 1.4** Construct a table that shows category, income, count of category-income pair, and probability of category-income pair.

```python
n = ...
# we will group the data by category and income as we did earlier
data_agg = data.groups(make_array('category','income'))
data_agg = data_agg.relabel(..., 'count pair')
data_agg = data_agg.with_column('prob pair', ...)
data_agg
```

```python
grader.check("q14")
```

**Question 1.5** Extend the table to show count of income and probability of income.

```python
# since we need data_agg for future analysis we make the join and assign the result to 
# the variable data_agg_prob_income
data_agg = data_agg.join(...)
data_agg = data_agg_prob_income.relabel(..., 'count income')
data_agg = data_agg.with_column('prob income',...)
data_agg
```

```python
grader.check("q15")
```

**Question 1.6** Extend the table to show count of category and probability of category.

```python
# since we need data_agg for future analysis we make the join and assign the result to 
# the variable data_agg_prob
data_agg = data_agg.join(...)
data_agg = data_agg.relabel(..., 'count category')
data_agg = data_agg.with_column('prob category',...)
data_agg
```

```python
grader.check("q16")
```

**Question 1.7** Hypothesize that category and income are independent of each other.

Recall the definition of independence:

$$
P(A \cap B) = P(A) \times P(B)
$$

So, if category and income are independent of each other, then the probability of a specific category and a specific income is equal to the probability of that category times the probability of that income:

$$
P((\textit{category}=x) \cap (\textit{income}=y)) = P(category=x) \times P(income=y)
$$

Extend the table to show the hypothesized indepedence.

```python
# since we need data_agg for future analysis we make the join and assign the result to 
# the variable data_agg_hyp_independence
data_agg  = data_agg.with_column('prob pair hypo', ...)
data_agg 
```

```python
grader.check("q17")
```

**Question 1.8** Calculate and show the expected count of each category-income pair, based on the hypothesized probabilities.

```python
# since we need data_agg for future analysis we make the join and assign the result to 
# the variable data_agg
data_agg  = data_agg.with_column('count pair expected', ...)
data_agg
```

```python
grader.check("q18")
```

**Question 1.9** Calculate and show the sample chi-squared like this:  

$$
\chi^2 = \sum \cfrac{\left(\textit{actual} - \textit{expected}\right)^2}{\textit{expected} }
$$

```python
actual = data_agg.column(...)
expected = ...
sample_chisquared = ...
print("Sample Chi-squared statistic:", sample_chisquared)
```

```python
grader.check("q19")
```

**Question 1.10** Get 1,000,000 values from the standard chi-squared distribution for the appropriate degrees of freedom.

$$
\textit{df} = (\textit{number of unique values in variable \#1} - 1) \times (\textit{number of unique values in variable \#2} - 1)
$$ 

Show the degrees of freedom, a few of the values, and a histogram of all the values (50 bins, range 0 to 35).

```python
r = len(...)
c = len(...)
df = (r-1)*(c-1)
dist_array = chisquare(df, 1000000)
print("Degrees of freedom:", df)
dist = Table().with_column('chisquared', dist_array)
dist.hist(bins=50, range=make_array(0,35))
```

```python
grader.check("q110")
```

**Question 1.11** Calculate and show the probability of getting a sample with the sample chi-squared (or above), given that the hypothesis is correct (this is the p-value).
Also show the sample chi-squared and a histogram of the standard chi-squared distribution with the area corresponding to the probability highlighted.

```python
p_value = ... / dist.num_rows
print("Sample Chi-squared statistic:", sample_chisquared)
print("P-value:", p_value)
dist.hist(bins=50, range=make_array(0,35), left_end=sample_chisquared, right_end=35) 
```

```python
grader.check("q111")
```

**Question 1.12** Calculate and show the critical value at significance level 0.05.
Also show the significance level and a histogram of the standard chi-squared distribution with the area corresponding to the significance level highlighted.

```python
sig_level = ...
cv = percentile((..., dist.column('chisquared'))
print("Significance level:", sig_level)
print("Critical Value:", cv)
dist.hist(bins=50, range=make_array(0,35), left_end=cv, right_end=35)
```

```python
grader.check("q112")
```

**Question 1.13** Calculate and show what to conclude about the hypothesis at significance level 0.05.

```python
reject_null_1 = p_value <= 
reject_null_2 = ...sample_chisquared >=
print("We choose to Reject the null hypothesis ", reject_null_1)
print("We choose to Reject the null hypothesis ", reject_null_2)
```

```python
grader.check("q113")
```

Reject the hypothesis.  Conclude that customer preference for product category does depend on income level.  Take action to customize the website accordingly. 

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Assume a significance level of 0.01, rather than 0.05. Calculate the critical value.

```python
significance_level = 0.01
cv_c1 = ...
cv_c1
```

```python
grader.check("qc_1")
```

**Question 2**. Assume a significance level of 0.01, rather than 0.05. Calculate whether the manager should conclude that customer preference for product category and income level are independent of each other.
(Answer is True or False.)

```python
significance_level = 0.01
cv_c2 = ...
conclusion = ...
conclusion
```

```python
grader.check("qc_2")
```

**Question 3**. Assume a significance level of 0.010, rather than 0.05. Calculate whether the manager should conclude that customer preference for product category and income level are independent of each other.
(Answer is True or False.)

```python
sig_level_c3 = 0.10
cv_c3 = ...
conclusion_2 = ...
conclusion_2
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised October 27, 2023
</span>
</p>

Tests for q01 (points: None)
```
>>> data.num_rows
555
```
```
>>> type(data) is Table and data.num_rows == 555 and (data.num_columns == 3)
True
```
```
>>> data.column(1).item(0) == 'phone' and data.column(2).item(0) == 'from_25k_to_50k'
True
```


Tests for q11 (points: None)
```
>>> mean(data_pair.column('count')) == 55.5
True
```


Tests for q110 (points: None)
```
>>> df == 4
True
```
```
>>> r == 2 and c == 5
True
```


Tests for q111 (points: None)
```
>>> p_value < 0.001
True
```


Tests for q112 (points: None)
```
>>> sig_level == 0.05
True
```
```
>>> isclose(cv, 9.494, atol=0.1)
True
```


Tests for q113 (points: None)
```
>>> reject_null_1 and reject_null_2
True
```


Tests for q12 (points: None)
```
>>> type(data_category) is Table and data_category.num_rows == 2 and (data_category.num_columns == 2)
True
```
```
>>> mean(data_category.column('count')) == 277.5
True
```
```
>>> sum(data_category.column(0) == make_array('camera', 'phone')) == 2
True
```


Tests for q13 (points: None)
```
>>> type(data_income) is Table and data_income.num_rows == 5 and (data_income.num_columns == 2)
True
```
```
>>> mean(data_income.column('count')) == 111.0
True
```
```
>>> sum(data_income.column(0) == make_array('above_100k', 'below_25k', 'from_25k_to_50k', 'from_50k_to_75k', 'from_75k_to_100K')) == 5
True
```


Tests for q14 (points: None)
```
>>> mean(data_agg.column('count pair')) == 55.5
True
```
```
>>> round(mean(data_agg.column('prob pair')), 2) == 0.1
True
```


Tests for q15 (points: None)
```
>>> mean(data_agg.column('count income')) == 111.0
True
```


Tests for q16 (points: None)
```
>>> mean(data_agg.column('prob category')) == 0.5
True
```


Tests for q17 (points: None)
```
>>> mean(data_agg.column('prob pair hypo')) == 0.1
True
```


Tests for q18 (points: None)
```
>>> type(data_agg) is Table and data_agg.num_rows == 10 and (len(data_agg.columns) == 10)
True
```
```
>>> data_agg.labels == ('category', 'income', 'count pair', 'prob pair', 'count income', 'prob income', 'count category', 'prob category', 'prob pair hypo', 'count pair expected')
True
```
```
>>> mean(data_agg.column('count pair expected')) == 55.5
True
```


Tests for q19 (points: None)
```
>>> round(sample_chisquared, 2) == 33.92
True
```


Tests for qc_1 (points: None)
```
>>> cv_c1 == percentile((1 - significance_level) * 100, dist.column('chisquared'))
True
```


Tests for qc_2 (points: None)
```
>>> conclusion == False
True
```


Tests for qc_3 (points: None)
```
>>> conclusion_2 == False
True
```
