**THIS SOLUTION DIFFERS SLIGHTLY FROM THE SOLUTION SHOWN IN THE LECTURE VIDEO.**  
**IT USES A DATA FILE COMPRISING 637 OBSERVATIONS, RATHER THAN 638 OBSERVATIONS.**

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# Lab B: Probability & Average

## Setup

```python
# Import some useful functions
import warnings
from datascience import Table, are
from numpy import mean, median, isclose, percentile


# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Hide some unnecessary warning messages
warnings.filterwarnings("ignore")
```

## Shopping for Real Estate Investments

### Business Decision

A real estate brokerage has been engaged by an investor interested in understanding the San Francisco bay area residential real estate market.  Characterize the San Francisco bay area residential real estate market to inform a decision about what to buy. 

### Data

Retrieve real estate data, show the number of observations and the first few observations.

```python
data = ...
observations = ...
observations, data
```

```python
grader.check("q1")
```

### Analysis

#### Mean & Median

Calculate the mean price.

```python
# Calculate the mean price (use the mean function, don't call the data object "mean" !!!)
mean_price = ...
mean_price

# Calculate the mean price (use the sum function and divide by the appropriate number)
mean_price = ...
mean_price
```

```python
grader.check("q2")
```

Calculate the median price.

```python
median_price = ...
median_price
```

```python
grader.check("q3")
```

Calculate the minimum price.

```python
min_price = ...
min_price
```

```python
grader.check("q4")
```

Calculate the maximum price.

```python
max_price = max(data.column('PRICE'))
max_price
```

```python
grader.check("q5")
```

#### Probability

Sanity check: Calculate the probability that a price is equal to the mean price.  
Sanity check: Calculate the probability that a price is equal to the mean price (rounded down to the nearest dollar).  
Sanity check: Calculate the probability that a price is equal to the mean price (rounded up to the nearest dollar).

```python
# Calculate the probability that a price is equal to the mean price
price_equal = ...

# Calculate the probability that a price is equal to the mean price (rounded down to the nearest dollar, use the int function)
price_eq_round_down = ...

# Calculate the probability that a price is equal to the mean price (rounded up to the nearest dollar, use the int function)
price_eq_round_up = ...
```

```python
grader.check("q6")
```

Calculate the probability that a price is less than or equal to the mean price.

```python
probability = ...
probability
```

```python
grader.check("q7")
```

Sanity check: Calculate the probability that a price is less than or equal to median price.

```python
...
```

```python
grader.check("q8")
```

#### Percentile

Calculate the 25th percentile of prices.  
Calculate the 50th percentile of prices.  
Calculate the 75th percentile of prices.

```python
# Calculate the 25th percentile of prices (use the percentile function) 
p25 = ...
p25

# Calculate the 50th percentile of prices (use the percentile function)
p50 = ...
p50

# Calculate the 75th percentile of prices (use the percentile function)
p75 = ...
p75
```

```python
grader.check("q9")
```

Sanity check: Calculate the probability that a price is less than or equal to 25th percentile of prices.  
Sanity check: Calculate the probability that a price is less than or equal to 50th percentile of prices.  
Sanity check: Calculate the probability that a price is less than or equal to 75th percentile or prices.

```python
# Calculate the probability that a price is less than or equal to 25th percentile of prices
...

# Calculate the probability that a price is less than or equal to 50th percentile of prices
...

# Calculate the probability that a price is less than or equal to 75th percentile or prices
...
```

```python
grader.check("q10")
```

Sanity check: Calculate the probability that a price is in 1st quartile of prices (min to percentile 25th).  
Sanity check: Calculate the probability that a price is in 2nd quartile of prices (just above 25th percentile to 50th percentile).  
Sanity check: Calculate the probability that a price is in 3rd quartile of prices (just above 50th percentile to 75th percentile).  
Sanity check: Calculate the probability that a price is in 4th quartile of prices (just above 75th percentile to max).

```python
# Calculate the probability that a price is in 1st quartile of prices (min to percentile 25th)
data_q1 = ...
...

# Calculate the probability that a price is in 2nd quartile of prices (percentile 25th to 50th)
data_q2 = ...
...

# Calculate the probability that a price is in 3rd quartile of prices (percentile 50th to 75th)
data_q3 = ...
...

# Calculate the probability that a price is in 4th quartile of prices (percentile 75th to max)
data_q4 = ...
...
```

```python
data_q1.num_rows
```

```python
grader.check("q11")
```

Calculate the mean price of 1st quartile of prices (min to 25th percentile).  
Calculate the mean price of 2nd quartile of prices (just above 25th percentile to 50th percentile).  
Calculate the mean price of 3rd quartile of prices (just above 50th percentile to 75th percentile).  
Calculate the mean price of 4th quartile of prices (just above 75th percentile to max).

```python
# Calculate the mean price of 1st quartile of prices (min to percentile 25th)
...

# Calculate the mean price of 2nd quartile of prices (just above percentile 25th to 50th)
...

# Calculate the mean price of 3rd quartile of prices (just above percentile 50th to 75th)
...

# Calculate the mean price of 4th quartile of prices (just above percentile 75th to max)
...
```

```python
grader.check("q12")
```

Calculate the median price of 1st quartile of prices (min to 25th percentile).  
Calculate the median price of 2nd quartile of prices (just above 25th percentile to 50th percentile).  
Calculate the median price of 3rd quartile of prices (just above 50th percentile to 75th percentile).  
Calculate the median price of 1st quartile of prices (just above 75th percentile to max).

```python
# Calculate the median price of 1st quartile of prices (min to percentile 25th)
...

# Calculate the median price of 2nd quartile of prices (just above percentile 25th to 50th)
...

# Calculate the median price of 3rd quartile of prices (just above percentile 50th to 75th)
...

# Calculate the median price of 1st quartile of prices (just above percentile 75th to max)
...
```

```python
grader.check("q13")
```

#### Multiple Variables

Calculate the mean area (SQUARE_FEET) of properties with price in 1st quartile of prices.  
Calculate the mean area (SQUARE_FEET) of properties with price in 2nd quartile of prices.  
Calculate the mean area (SQUARE_FEET) of properties with price in 3rd quartile of prices.  
Calculate the mean area (SQUARE_FEET) of properties with price in 4th quartile of prices.

```python
# Calculate the mean area (SQUARE_FEET) of properties with price in 1st quartile of prices
...

# Calculate the mean area (SQUARE_FEET) of properties with price in 2nd quartile of prices
...

# Calculate the mean area (SQUARE_FEET) of properties with price in 3rd quartile of prices
...

# Calculate the mean area (SQUARE_FEET) of properties with price in 4th quartile of prices
...
```

```python
grader.check("q14")
```

```python
sq_feet_1st = ...
mean(sq_feet_1st.column('PRICE'))
```

```python
grader.check("q15")
```

#### Joint Probability

Calculate the probability that a property has an area smaller than 1000 sqft and has a price in the 2nd quartile of prices.

```python
...
```

```python
grader.check("q16")
```

Calculate the probability that a property has an area smaller than 2000 sqft and has a price in the 4th quartile of prices.

```python
...
```

```python
grader.check("q17")
```

Calculate the probability that a property has an area larger than 1500 sqft and has a price in the 1st quartile of prices.

```python
...
```

```python
grader.check("q18")
```

```python
...
```

```python
grader.check("q19")
```

```python
...
```

```python
grader.check("q20")
```

```python
...
```

```python
grader.check("q21")
```

```python
...
```

```python
grader.check("q22")
```

```python
data_notq1 = ...

...
```

```python
grader.check("q23")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Calculate the probability that a property is in San Francisco, given that its price is in the top 25% of prices.

```python
sf_high_price = data_q4...
print("The number of rows in data_q4:", data_q4.num_rows)
print("Probability in SF, given it's price is in the top 25% of prices: ", sf_high_price)
```

```python
grader.check("qc_1")
```

**Question 2**. Calculate the probability that a property is in San Francisco, given that its price is below the median price.

```python
low_price = data.where(..., ...)
sf_given_low = ...
print("Probability in SF, given it's price is below the median: ", sf_given_low)
```

```python
grader.check("qc_2")
```

**Question 3**. Calculate the probability that a property is either larger than 2000 sqft or built more recently than year 2010.

```python
condition = ...
print("Probability a home is 2000+ sqft, or newer than 2010: ", condition)
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised January 28, 2025
</span>
</p>

Tests for q1 (points: None)
```
>>> observations == data.num_rows
True
```


Tests for q10 (points: None)
```
>>> data.where('PRICE', are.below_or_equal_to(p25)).num_rows / data.num_rows >= 0.25
True
```
```
>>> data.where('PRICE', are.below_or_equal_to(p50)).num_rows / data.num_rows >= 0.5
True
```
```
>>> data.where('PRICE', are.below_or_equal_to(p75)).num_rows / data.num_rows >= 0.75
True
```


Tests for q11 (points: None)
```
>>> data_q1.num_rows == 161
True
```
```
>>> data_q1.num_rows / data.num_rows >= 0.25
True
```
```
>>> data_q2.num_rows == 160
True
```
```
>>> data_q2.num_rows / data.num_rows >= 0.25
True
```
```
>>> data_q3.num_rows == 158
True
```
```
>>> data_q3.num_rows / data.num_rows >= 0.24
True
```
```
>>> data_q4.num_rows == 158
True
```
```
>>> data_q4.num_rows / data.num_rows >= 0.24
True
```


Tests for q12 (points: None)
```
>>> round(mean(data_q1.column('PRICE'))) >= 70500
True
```
```
>>> round(mean(data_q2.column('PRICE'))) >= 1090000
True
```
```
>>> round(mean(data_q3.column('PRICE'))) >= 1600000
True
```
```
>>> round(mean(data_q4.column('PRICE'))) >= 4500000
True
```


Tests for q13 (points: None)
```
>>> bool(median(data_q1.column('PRICE')) == 749000)
True
```
```
>>> bool(median(data_q2.column('PRICE')) == 1090000)
True
```
```
>>> bool(median(data_q3.column('PRICE')) == 1599500)
True
```
```
>>> median(data_q3.column('PRICE')) == 1599500
True
```


Tests for q14 (points: None)
```
>>> round(mean(data_q1.column('SQUARE_FEET'))) >= 917
True
```
```
>>> round(mean(data_q2.column('SQUARE_FEET'))) >= 1376
True
```
```
>>> round(mean(data_q3.column('SQUARE_FEET'))) >= 1642
True
```
```
>>> round(mean(data_q4.column('SQUARE_FEET'))) >= 3096
True
```


Tests for q15 (points: None)
```
>>> sq_feet_1st.num_rows == 160
True
```
```
>>> round(mean(sq_feet_1st.column('PRICE'))) >= 820000
True
```


Tests for q16 (points: None)
```
>>> data_q2.where('SQUARE_FEET', are.below(1000)).num_rows / data.num_rows <= 0.1
True
```


Tests for q17 (points: None)
```
>>> data_q4.where('SQUARE_FEET', are.below(2000)).num_rows / data.num_rows <= 0.1
True
```


Tests for q18 (points: None)
```
>>> data_q1.where('SQUARE_FEET', are.above(1500)).num_rows / data.num_rows <= 0.1
True
```


Tests for q19 (points: None)
```
>>> data_q2.where('SQUARE_FEET', are.below(1000)).num_rows / data_q2.num_rows <= 0.5
True
```


Tests for q2 (points: None)
```
>>> assert isinstance(mean_price, (int, float))
>>> expected = mean(data.column('PRICE'))
>>> assert isclose(mean_price, expected)
```


Tests for q20 (points: None)
```
>>> data_q4.where('SQUARE_FEET', are.below(2000)).num_rows / data_q4.num_rows <= 0.5
True
```


Tests for q21 (points: None)
```
>>> data_q1.where('SQUARE_FEET', are.above(1500)).num_rows / data_q1.num_rows <= 0.1
True
```


Tests for q22 (points: None)
```
>>> (data_q1.num_rows + data.where('SQUARE_FEET', are.below(1000)).num_rows - data_q1.where('SQUARE_FEET', are.below(1000)).num_rows) / data.num_rows <= 0.5
True
```


Tests for q23 (points: None)
```
>>> data_notq1.num_rows == 476
True
```
```
>>> (data_notq1.num_rows + data.where('SQUARE_FEET', are.below(1000)).num_rows - data_notq1.where('SQUARE_FEET', are.below(1000)).num_rows) / data.num_rows <= 1
True
```


Tests for q3 (points: None)
```
>>> def test_median_price():
...     assert isinstance(median_price, (int, float))
...     expected = median(data.column('PRICE'))
...     assert isclose(median_price, expected)
```


Tests for q4 (points: None)
```
>>> def test_min_price():
...     assert isinstance(min_price, (int, float))
...     expected = min(data.column('PRICE'))
...     assert isclose(min_price, expected)
```


Tests for q5 (points: None)
```
>>> def test_max_price():
...     assert isinstance(max_price, (int, float))
...     expected = max(data.column('PRICE'))
...     assert isclose(max_price, expected)
```


Tests for q6 (points: None)
```
>>> def test_price_eq():
...     assert isinstance(price_equal, (int, float))
...     expected = data.where('PRICE', mean_price).num_rows / data.num_rows
...     assert isclose(price_equal, expected)
```
```
>>> def test_price_round_down():
...     assert isinstance(price_eq_round_down, (int, float))
...     expected = data.where('PRICE', int(mean_price)).num_rows / data.num_rows
...     assert isclose(price_eq_round_down, expected)
```
```
>>> def test_price_round_up():
...     assert isinstance(price_eq_round_up, (int, float))
...     expected = data.where('PRICE', int(mean_price) + 1).num_rows / data.num_rows
...     assert isclose(price_eq_round_up, expected)
```


Tests for q7 (points: None)
```
>>> import math
>>> math.isclose(probability, 0.75, abs_tol=0.01)
True
```


Tests for q8 (points: None)
```
>>> round(data.where('PRICE', are.below_or_equal_to(median_price)).num_rows / data.num_rows) == 0.5
False
```


Tests for q9 (points: None)
```
>>> p25 <= 900000
True
```
```
>>> p50 <= 1300000
True
```
```
>>> p75 <= 2000000
True
```


Tests for qc_1 (points: None)
```
>>> round(sf_high_price, 3) == 0.949
True
```


Tests for qc_2 (points: None)
```
>>> round(sf_given_low, 3) == 0.716
True
```


Tests for qc_3 (points: None)
```
>>> round(condition, 3) == 0.377
True
```
