# Lab J1: Chi-Squared Test for Uniform Probabilities

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

## Setup

```python
import warnings
from datascience import Table, make_array, percentile, are
from numpy import mean, ndarray, isclose
from numpy.random import chisquare

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Manufacturing Cars & Lemons

### Business Decision

A car manufacturing organization produces cars 5 days per week, Monday through Friday.  We might expect that the day of the week that cars are produced has no relationship to the number of defects in those cars.  However, car buyers are sometimes warned that cars made on a Monday or a Friday have more defects on average, i.e., they are more likely to be "lemons". 

### Data

Here are the number of defects per day from a one week sample:  
Monday 32, Tuesday 22, Wednesday 26, Thursday 19, Friday 30 

Show the sample information as a table.

```python
data = Table().with_columns('dow', ...,
                            'defects',...)
data
```

```python
grader.check("q01")
```

### Analysis

Hypothesize that for the population of all cars manufactured, the proportion of defects occurring over a week is the same for every day of the week.  That would imply that the probability that a defect occurs on any specific day of the week is 20%.  The alternative to this hypothesis is that the proportion of defects occurring on some one day is different than for some
other day of the week.

```python
prob_hypo = ...
data = data.with_column('prob_hypo', ...)

print("Hypothesis probability is", prob_hypo)
data
```

```python
grader.check("q11")
```

Calculate and show the expected number of defects for each day of the week,
based on the hypothesized probability of a defect occurring on that day of the week. 

```python
data = data.with_column('defects_expected', prob_hypo * ...)
data
```

```python
grader.check("q12")
```

Calculate and show the sample chi-squared like this:  

$$
\chi^2 = \sum \cfrac{\left(\textit{actual} - \textit{expected}\right)^2}{\textit{expected} }
$$

```python
defects = ...
defects_expected = ...

diff = ...
diff_squared = ...
relative_diff_squared = ...

sample_chisquared = sum(relative_diff_squared)
data = data.with_columns('diff', diff, 'diff**2', diff_squared, 'relative diff**2', relative_diff_squared)
data
print("Sample chisquared is", sample_chisquared)
```

```python
grader.check("q13")
```

Get 1,000,000 values from the standard chi-squared distribution for the appropriate degrees of freedom.

$$
\textit{df} = \textit{number of unique values} - 1
$$

Show the degrees of freedom, a few of the values, and a histogram of all the values (50 bins, range 0 to 25).

```python
degrees_of_freedom = ...

dist_array = chisquare(..., 1000000)
dist = Table().with_column('chisquared', ...)
print("Degrees of freedom is", degrees_of_freedom)
dist
dist.hist(bins=50, range=make_array(0,25))
```

```python
grader.check("q14")
```

Calculate and show the probability of getting a sample with the sample chi-squared (or above), given that the hypothesis is correct (this is the p-value).
Also show the sample chi-squared and a histogram of the standard chi-squared distribution with the area corresponding to the probability highlighted.

```python
p_value = dist.where('chisquared', ...).num_rows / ...
print("Sample chisquared is", sample_chisquared)
print("Our p-value is", p_value)
dist.hist(bins=50, range=make_array(0,25), left_end=sample_chisquared, right_end=25)
```

```python
grader.check("q15")
```

Calculate and show the critical value at significance level 0.05.
Also show the significance level and a histogram of the standard chi-squared distribution with the area corresponding to the significance level highlighted.

```python
sig_level = ...
cv = percentile(..., dist.column('chisquared'))
print("Our significance level is", sig_level)
print("Our critical value is", cv)
dist.hist(bins=50, range=make_array(0,25), left_end=cv, right_end=25)
```

```python
grader.check("q16")
```

Calculate and show what to conclude about the hypothesis at significance level 0.05

```python
reject_the_null_1 = p_value <= ...
reject_the_null_2 = sample_chisquared >= ...
print("Do we reject the null hypothesis?", reject_the_null_1)
print("Do we reject the null hypothesis?", reject_the_null_2)
```

```python
grader.check("q17")
```

Accept the hypothesis. Conclude that there is no difference in number of defects based on day of the week.

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Assume a significance level of 0.01, rather than 0.05. Calculate the critical value.

```python
sig_level_01 = 0.01
cv_c1 = ...
cv_c1
```

```python
grader.check("qc_1")
```

**Question 2**. Assume a significance level of 0.01, rather than 0.05. Calculate whether the manager should conclude that day of week has no effect on the defect rate. (Answer is True or False.)

```python
sig_level_01 = 0.01
cv_c2 = ...
conclusion = ...
conclusion
```

```python
grader.check("qc_2")
```

**Question 3**. Assume a significance level of 0.10, rather than 0.05. Calculate whether the manager should conclude that day of week has no effect on the defect rate. (Answer is True or False.)

```python
sig_level_c3 = 0.10
cv_c3 = ...
conclusion_c3 = ...
conclusion_c3
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised October 27, 2023
</span>
</p>

Tests for q01 (points: None)
```
>>> type(data) is Table and data.column(0).item(0) == 'Monday' and (data.column(1).item(0) == 32)
True
```


Tests for q11 (points: None)
```
>>> type(prob_hypo) is float and prob_hypo > 0 and (prob_hypo == 0.2)
True
```
```
>>> type(data) is Table and data.column(0).item(3) == 'Thursday' and (data.column(1).item(3) == 19)
True
```
```
>>> mean(data.column('prob_hypo')) == 0.2
True
```


Tests for q12 (points: None)
```
>>> type(data) is Table and data.column(0).item(0) == 'Monday' and (data.column(3).item(0) == 25.8)
True
```
```
>>> mean(data.column('defects_expected')) == 25.8
True
```


Tests for q13 (points: None)
```
>>> type(defects) is ndarray and defects.shape == (5,) and (defects[0] == 32) and (defects[1] == 22) and (defects[2] == 26) and (defects[3] == 19) and (defects[4] == 30)
True
```
```
>>> type(diff) is ndarray and diff.shape == (5,) and (round(diff[1], 1) == -3.8) and (round(diff[2], 1) == 0.2) and (round(diff[4], 1) == 4.2)
True
```
```
>>> type(data) is Table and data.num_rows == 5 and (data.num_columns == 7) and (round(data.column(5).item(0), 2) == 38.44) and (round(data.column(4).item(0), 2) == 6.2)
True
```
```
>>> round(sample_chisquared, 1) == 4.5
True
```


Tests for q14 (points: None)
```
>>> degrees_of_freedom == data.num_rows - 1
True
```
```
>>> dist_array.shape == (1000000,) and type(dist_array) is ndarray
True
```


Tests for q15 (points: None)
```
>>> isclose(p_value, 0.34, rtol=0.01)
True
```


Tests for q16 (points: None)
```
>>> sig_level == 0.05
True
```
```
>>> round(cv, 1) == 9.5
True
```


Tests for q17 (points: None)
```
>>> reject_the_null_1 == reject_the_null_2
True
```
```
>>> reject_the_null_1 == 0
True
```


Tests for qc_1 (points: None)
```
>>> cv_c1 == percentile((1 - sig_level_01) * 100, dist.column('chisquared'))
True
```


Tests for qc_2 (points: None)
```
>>> conclusion == True
True
```


Tests for qc_3 (points: None)
```
>>> conclusion_c3 == True
True
```
