# Lab J2: Chi-Squared Test for Specific Probabilities

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

## Setup

```python
import warnings
from datascience import Table, make_array, percentile, are
from numpy import isclose
from numpy.random import chisquare

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Customer Preference for Microwave Oven Brand

### Business Decision

An appliance retail chain wants to test whether customer preference for brand of microwave oven is the same in Milwaukee as in Cleveland.

### Data

**Question 0.1** Here are the probabilities that customers prefer various brands in Cleveland:  

Brand 1: 20%, Brand 2: 35%, Brand 3: 30%, Brand 4: 15% 

Show the sample information as a table.

```python
data_cleveland = Table().with_columns('brand', ...,
                            'prob',...)
data_cleveland
```

```python
grader.check("q01")
```


**Question 0.2** Here are the numbers of sales of various brands in Milwaukee during a specific time period:  

Brand 1: 102, Brand 2: 121, Brand 3: 120, Brand 4: 57 

Show the sample information as a table.

```python
data_milwaukee = Table().with_columns('brand', ...,
                            'sales',...)
data_milwaukee
```

```python
grader.check("q02")
```

### Analysis

**Question 1.1** Hypothesize that the preferences for brands in Milwaukee are the same as those in Cleveland,
so that the probability of a sale of any specific brand in Milwaukee is equal to the probability of a sale of that brand in Cleveland.
The alternative to this hypothesis is that the preferences in Milwaukee are different than in Cleveland.

```python
data = data_milwaukee
data = ...
data
```

```python
grader.check("q11")
```

**Question 1.2** Calculate and show the expected number of sales of each brand, based on the hypothesized probability a sale of that brand. 

```python
sales = data.column(...)
prob_hypo = data.column(...)
sales_expected = ...
data = data.with_column(... , ...)
data
```

```python
grader.check("q12")
```

**Question 1.3** Calculate and show the sample chi-squared like this:  

$$
\chi^2 = \sum \cfrac{\left(\textit{actual} - \textit{expected}\right)^2}{\textit{expected} }
$$

```python
diff = ...
diff_squared = diff**2
relative_diff_squared = ...

data = data.with_columns(...)

sample_chisquared = ...
data
print("Our sample Chi-Squared is ", sample_chisquared)
```

```python
grader.check("q13")
```

**Question 1.4** Get $1,000,000$ values from the standard chi-squared distribution for the appropriate degrees of freedom.

$$
\textit{df} = \textit{number of unique values} - 1
$$

Show the degrees of freedom, a few of the values, and a histogram of all the values (50 bins, range 0 to 25).

```python
degrees_freedom = ...
dist_array = chisquare(..., ...)
dist = ...
print("Degrees of freedom: ", degrees_freedom)
dist
dist.hist(bins=50, range=make_array(0,25))
```

```python
grader.check("q14")
```

**Question 1.5** Calculate and show the probability of getting a sample with the sample chi-squared (or above), given that the hypothesis is correct (this is the p-value).
Also show the sample chi-squared and a histogram of the standard chi-squared distribution with the area corresponding to the probability highlighted.

```python
p_value = ... / dist.num_rows
dist.hist(bins=50, range=make_array(0,25), left_end=sample_chisquared, right_end=25)
print("We calculate the p-value as: ", p_value)
print("Sample Chi-Squared: ", sample_chisquared)
```

```python
grader.check("q15")
```

**Question 1.6** Calculate and show the critical value at significance level 0.05.
Also show the significance level and a histogram of the standard chi-squared distribution with the area corresponding to the significance level highlighted.

```python
p_value = ... / dist.num_rows
print("Our Signifance Level is ", sig_level)
print("Our Critical Value is ", cv)
dist.hist(bins=50, range=make_array(0,25), left_end=cv, right_end=25)
```

```python
grader.check("q16")
```

**Question 1.7** Calculate and show what to conclude about the hypothesis at significance level 0.05.

```python
reject_null_1 = p_value <= 
reject_null_2 = ...sample_chisquared >=
print("We choose to Reject the null hypothesis ", reject_null_1)
print("We choose to Reject the null hypothesis ", reject_null_2)
```

```python
grader.check("q17")
```

Reject the hypothesis. Conclude that the preferences for brands in Cleveland is different than in Milwaukee.

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. Assume a significance level of 0.01, rather than 0.05. Calculate the critical value.

```python
sig_level_01 = 0.01
cv_c2 = ...
cv_c2
```

```python
grader.check("qc_1")
```

**Question 2**. Assume a significance level of 0.01, rather than 0.05. Calculate whether the manager should conclude that customer preferences for microwave brand is the same in Milwaukee and Cleveland.
(Answer is True or False.)

```python
sig_level_01 = 0.01
cv_c2 = ...
conclusion = ...
conclusion
```

```python
grader.check("qc_2")
```

**Question 3**. Assume a significance level of 0.10, rather than 0.05. Calculate whether the manager should conclude that customer preferences for microwave brand is the same in Milwaukee and Cleveland.
(Answer is True or False.)

```python
sig_level_c3 = 0.10
cv_c3 = ...
conclusion_2 = ...
conclusion_2
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised October 27, 2023
</span>
</p>

Tests for q01 (points: None)
```
>>> type(data_cleveland) is Table and data_cleveland.num_rows == 4 and (data_cleveland.num_columns == 2)
True
```
```
>>> sum(data_cleveland.column(1) == make_array(0.2, 0.35, 0.3, 0.15)) == 4
True
```
```
>>> sum(data_cleveland.column('prob')) == 1.0
True
```


Tests for q02 (points: None)
```
>>> type(data_milwaukee) is Table and data_milwaukee.num_rows == 4 and (data_cleveland.num_columns == 2)
True
```
```
>>> sum(data_milwaukee.column(1) == make_array(102, 121, 120, 57)) == 4
True
```
```
>>> sum(data_milwaukee.column('sales')) == 400
True
```


Tests for q11 (points: None)
```
>>> type(data) is Table and data.column(1).item(2) == 120 and (data.column(2).item(0) == 0.2)
True
```
```
>>> sum(data.column(1) == make_array(102, 121, 120, 57)) == 4
True
```
```
>>> sum(data.column('prob_hypo')) == 1
True
```


Tests for q12 (points: None)
```
>>> type(data) is Table and data.column(3).item(2) == 120 and (data.column(3).item(0) == 80)
True
```
```
>>> sum(data.column(3) == make_array(80, 140, 120, 60)) == 4
True
```


Tests for q13 (points: None)
```
>>> type(data) is Table and data.num_rows == 4 and (data.num_columns == 7) and isclose(data.column(6).item(1), 2.579, rtol=0.01) and isclose(data.column(6).item(0), 6.05, rtol=0.01)
True
```
```
>>> sum(data.column(5) == make_array(484, 361, 0, 9)) == 4
True
```


Tests for q14 (points: None)
```
>>> type(data) is Table and data.num_rows == 4 and (data.num_columns == 7) and isclose(data.column(6).item(1), 2.579, rtol=0.01) and isclose(data.column(6).item(0), 6.05, rtol=0.01)
True
```
```
>>> sum(data.column(5) == make_array(484, 361, 0, 9)) == 4
True
```


Tests for q15 (points: None)
```
>>> p_value < 0.033
True
```


Tests for q16 (points: None)
```
>>> sig_level == 0.05
True
```
```
>>> isclose(cv, 7.835, rtol=0.01)
True
```


Tests for q17 (points: None)
```
>>> reject_null_1 == reject_null_1
True
```
```
>>> reject_null_1 == 1
True
```


Tests for qc_1 (points: None)
```
>>> cv_c2 == percentile((1 - sig_level_01) * 100, dist.column('chisquared'))
True
```


Tests for qc_2 (points: None)
```
>>> conclusion == True
True
```


Tests for qc_3 (points: None)
```
>>> conclusion_2 == False
True
```
