# Lab L: Nonlinearity

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

## Setup

```python
import warnings
from datascience import *
from numpy import log, exp
from statsmodels.formula.api import ols

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Plot styling
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Your Own Jet Airplane

### Business Decision

Your company wants to buy a used Boeing 747.  This is the most popular airplane on the planet so many used 747s are available.

The data here describe 288 recent sales of Boeing 747s. For each sale, you have the transaction price (in millions $) and the total number of miles traveled at time of sale (in millions).

### Data

Retrieve data from file 'Airplane_Purchases.csv'.  Show the first few sales.

```python
airplane_purchases = ...
airplane_purchases
```

```python
grader.check("q1")
```

### Analysis: Linear Regression

Build a linear regression model to predict price based on mileage.  
Show the model goodness of fit (R^2).  
Show the model parameter values (intercept and coefficient).

```python
model_1 = ...
model_1.rsquared
model_1.params
```

```python
grader.check("q2_1")
```

Use the model to predict the sale prices.  
Show the first few predictions.

```python
all_predictions_1 = ...

airplane_purchases_preds = airplane_purchases.with_column('price_predicted', all_predictions_1)
airplane_purchases_preds
```

```python
grader.check("q2_2")
```

Visualize the model performance as a scatterplot of actual price and predicted price vs. mileage.

```python
airplane_purchases_preds.scatter('mileage')
```

Show the errors between the sale actual prices and predicted prices.

```python
price_errors = ...

airplane_purchases_errors = airplane_purchases_preds.with_column('error', price_errors)
airplane_purchases_errors
```

```python
grader.check("q2_3")
```

Visualize the errors as a scatterplot.

```python
airplane_purchases_errors.select('mileage', 'error').scatter('mileage')
```

Use the model to predict the price of a Boeing 747 that has traveled 50 million miles.

```python
my_mileage = 50
boeing_price_pred_1 = ...
boeing_price_pred_1
```

```python
grader.check("q2_4")
```

### Analysis: Linear-Log Regression

Reset dataset to include mileage and price variables only.

```python
airplane_purchases = airplane_purchases.select('mileage','price')
```

Add a variable for log mileage.  You can use the `log()` function.  
Show the first few observations of the resulting dataset. 

```python
log_mileage_column = ...
airplane_purchases_mileages = airplane_purchases.with_column('log_mileage', log_mileage_column)
airplane_purchases_mileages
```

```python
grader.check("q3_1")
```

Build a linear regression model to predict price based on log mileage.  
Show the model goodness of fit (R^2).  
Show the model parameter values (intercept and coefficient).

```python
model_2 = ...
model_2.rsquared
model_2.params
```

```python
grader.check("q3_2")
```

Use the model to predict the sale prices.  
Show the first few predictions.

```python
all_predictions_2 = ...

airplane_mileage_preds = airplane_purchases_mileages.with_column('price_predicted', all_predictions_2)
airplane_mileage_preds
```

```python
grader.check("q3_3")
```

Visualize the model performance as a scatterplot of actual price and predicted price vs. log mileage.

```python
cols_for_scatter_1 = ...
airplane_mileage_preds.select(cols_for_scatter_1).scatter('log_mileage')
```

```python
grader.check("q3_4")
```

Visualize the model performance as a scatterplot of actual price and predicted price vs. mileage.

```python
cols_for_scatter_2 = ...
airplane_mileage_preds.select(cols_for_scatter_2).scatter('mileage')
```

```python
grader.check("q3_5")
```

Show the errors between the sale actual prices and predicted prices.

```python
mileage_price_errors = ...
airplane_mileage_errors = airplane_mileage_preds.with_column('error', mileage_price_errors)
airplane_mileage_errors
```

```python
grader.check("q3_6")
```

Visualize the errors as a scatterplot.

```python
cols_for_scatter_3 = 'mileage', 'error'
airplane_mileage_errors.select(cols_for_scatter_3).scatter('mileage')
```

Use the model to predict the price of a Boeing 747 that has traveled 50 million miles.

```python
mileage = 50
log_mileage = ...
boeing_price_pred_2 = ...
boeing_price_pred_2
```

```python
grader.check("q3_7")
```

### Analysis: Log-Linear Regression

Reset dataset to include mileage and price variables only.

```python
airplane_purchases = airplane_purchases.select('mileage','price')
```

Add a variable for log price.  You can use the `log()` function.  
Show the first few observations of the resulting dataset. 

```python
log_price = ...
airplane_purchases_log_price = airplane_purchases.with_column('log_price', log_price)
airplane_purchases_log_price
```

```python
grader.check("q4_1")
```

Build a linear regression model to predict log price based on mileage.  
Show the model goodness of fit (R^2).  
Show the model parameter values (intercept and coefficient).

```python
model_3 = ...
model_3.rsquared
model_3.params
```

```python
grader.check("q4_2")
```

Use the model to predict the sale log prices.
Show the first few predictions.

```python
all_predictions_3 = ...
airplane_log_price_preds = airplane_purchases_log_price.with_column('log_price_predicted', all_predictions_3)
airplane_log_price_preds
```

```python
grader.check("q4_3")
```

Visualize the model performance as a scatterplot of actual log price and predicted log price vs. mileage.

```python
cols_for_scatter_4 = ...
airplane_log_price_preds.select(cols_for_scatter_4).scatter('mileage')
```

```python
grader.check("q4_4")
```

Predict the sale prices.  Show the first few predictions. You can use the `exp()` function.

```python
sale_preds = ...

airplane_log_price_preds = airplane_log_price_preds.with_column('price_predicted', sale_preds)
airplane_log_price_preds
```

```python
grader.check("q4_5")
```

Visualize the model performance as a scatterplot of actual price and predicted price vs. mileage.

```python
cols_for_scatter_5 = ...
airplane_log_price_preds.select(cols_for_scatter_5).scatter('mileage')
```

```python
grader.check("q4_6")
```

Show the errors between the sale actual prices and predicted prices.

```python
log_price_errors = airplane_log_price_preds.column('price')-airplane_log_price_preds.column('price_predicted')
airplane_log_price_errors = airplane_log_price_preds.with_column('error', log_price_errors)
airplane_log_price_errors
```

```python
grader.check("q4_7")
```

Visualize the errors as a scatterplot.

```python
airplane_log_price_errors.select('mileage', 'error').scatter('mileage')
```

Use the model to predict the price of a Boeing 747 that has traveled 50 million miles.

```python
mileage = 50
mileage_pred = ...
boeing_price_pred_3 = exp(mileage_pred)
boeing_price_pred_3
```

```python
grader.check("q4_8")
```

# OPTIONAL Challenge Questions
You can not get a 100% in the course without getting these questions correct but you do not need to get them correct to pass the course and get the certificate.

**Question 1**. According to the log-linear model and the dataset (airplane_log_price_errors), calculate the range of predicted price errors for airplanes with fewer than 40 million miles.

```python
data_below40 = ...
range_c1 = ...
range_c1
```

```python
grader.check("qc_1")
```

**Question 2**. According to the log-linear model and the dataset (airplane_log_price_errors), calculate the range of predicted price errors for airplanes with more than 40 million miles.

```python
data_above40 = ...
range_c2 = ...
range_c2
```

```python
grader.check("qc_2")
```

**Question 3**. According to the log-linear model and the dataset (airplane_log_price_errors), show the mileage of the airplane with the largest predicted price absolute error.

```python
c3_data = ...
x = ...
x
```

```python
grader.check("qc_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised November 10, 2023
</span>
</p>

Tests for q1 (points: None)
```
>>> airplane_purchases.labels == ('mileage', 'price') and airplane_purchases.num_rows == 288
True
```


Tests for q2_1 (points: None)
```
>>> cond1 = list(model_1.model.data.frame.labels) == list(airplane_purchases.labels)
>>> cond2 = model_1.model.data.frame.num_rows == airplane_purchases.num_rows
>>> cond3 = all((model_1.model.data.frame.row(i) == airplane_purchases.row(i) for i in range(airplane_purchases.num_rows)))
>>> cond1 and cond2 and cond3
True
```
```
>>> hasattr(model_1, 'rsquared') and hasattr(model_1, 'params') and (len(model_1.params) == 2)
True
```


Tests for q2_2 (points: None)
```
>>> all(all_predictions_1 == model_1.predict(airplane_purchases))
True
```


Tests for q2_3 (points: None)
```
>>> all(price_errors == airplane_purchases_preds.column('price') - airplane_purchases_preds.column('price_predicted'))
True
```


Tests for q2_4 (points: None)
```
>>> all(boeing_price_pred_1 == model_1.predict(Table().with_column('mileage', my_mileage)))
True
```


Tests for q3_1 (points: None)
```
>>> all(log_mileage_column == log(airplane_purchases.column('mileage')))
True
```


Tests for q3_2 (points: None)
```
>>> cond1 = list(model_2.model.data.frame.labels) == list(airplane_purchases_mileages.labels)
>>> cond2 = model_2.model.data.frame.num_rows == airplane_purchases_mileages.num_rows
>>> cond3 = all((model_2.model.data.frame.row(i) == airplane_purchases_mileages.row(i) for i in range(airplane_purchases_mileages.num_rows)))
>>> cond1 and cond2 and cond3
True
```
```
>>> hasattr(model_2, 'rsquared') and hasattr(model_2, 'params') and (len(model_2.params) == 2)
True
```


Tests for q3_3 (points: None)
```
>>> all(all_predictions_2 == model_2.predict(airplane_purchases_mileages))
True
```


Tests for q3_4 (points: None)
```
>>> cols_for_scatter_1 == ('log_mileage', 'price', 'price_predicted')
True
```


Tests for q3_5 (points: None)
```
>>> cols_for_scatter_2 == ('mileage', 'price', 'price_predicted')
True
```


Tests for q3_6 (points: None)
```
>>> all(mileage_price_errors == airplane_mileage_preds.column('price') - airplane_mileage_preds.column('price_predicted'))
True
```


Tests for q3_7 (points: None)
```
>>> log_mileage == log(50)
True
```
```
>>> all(boeing_price_pred_2 == model_2.predict(Table().with_column('log_mileage', log_mileage)))
True
```


Tests for q4_1 (points: None)
```
>>> all(log_price == log(airplane_purchases.column('price')))
True
```


Tests for q4_2 (points: None)
```
>>> cond1 = list(model_3.model.data.frame.labels) == list(airplane_purchases_log_price.labels)
>>> cond2 = model_3.model.data.frame.num_rows == airplane_purchases_log_price.num_rows
>>> cond3 = all((model_3.model.data.frame.row(i) == airplane_purchases_log_price.row(i) for i in range(airplane_purchases_log_price.num_rows)))
>>> cond1 and cond2 and cond3
True
```
```
>>> hasattr(model_3, 'rsquared') and hasattr(model_3, 'params') and (len(model_3.params) == 2)
True
```


Tests for q4_3 (points: None)
```
>>> all(all_predictions_3 == model_3.predict(airplane_purchases_log_price))
True
```


Tests for q4_4 (points: None)
```
>>> cols_for_scatter_4 == ('mileage', 'log_price', 'log_price_predicted')
True
```


Tests for q4_5 (points: None)
```
>>> all(sale_preds == exp(airplane_log_price_preds.column('log_price_predicted')))
True
```


Tests for q4_6 (points: None)
```
>>> cols_for_scatter_5 == ('mileage', 'price', 'price_predicted')
True
```


Tests for q4_7 (points: None)
```
>>> all(log_price_errors == airplane_log_price_preds.column('price') - airplane_log_price_preds.column('price_predicted'))
True
```


Tests for q4_8 (points: None)
```
>>> all(mileage_pred == model_3.predict(Table().with_column('mileage', mileage)))
True
```
```
>>> all(boeing_price_pred_3 == exp(mileage_pred))
True
```


Tests for qc_1 (points: None)
```
>>> range_c1 == max(data_below40.column('error')) - min(data_below40.column('error'))
True
```


Tests for qc_2 (points: None)
```
>>> range_c2 == max(data_above40.column('error')) - min(data_above40.column('error'))
True
```


Tests for qc_3 (points: None)
```
>>> x == c3_data.sort('absolute error', descending=True).take(0).column('mileage').item(0)
True
```
